```python
# Initialize Otter
import otter
grader = otter.Notebook("ws2.ipynb")
```

### Workshop 2

```python
# Setup
# You may import math if needed in your solutions.
import math
```

## **Question 1**: Conditional Logic

### Q1a: Charge Classification
A charged particle has charge q (in Coulombs). Write a function `classify_charge(q)` that returns a string classification:
- "positive" if q > 0
- "negative" if q < 0  
- "neutral" if q = 0

```python
def classify_charge(q):
    # BEGIN SOLUTION NO PROMPT
    if q > 0:
        return "positive"
    elif q < 0:
        return "negative"
    else:
        return "neutral"
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q1a")
```

### Q1b: Electric Field Strength Classification
Write a function `classify_electric_field(E)` that classifies the strength of an electric field E (in N/C):
- "weak" if 0 ≤ |E| < 1000
- "moderate" if 1000 ≤ |E| < 10000  
- "strong" if 10000 ≤ |E| < 100000
- "extreme" if |E| ≥ 100000

Note: Use the absolute value of E since field strength since it's always positive, regardless of direction.

```python
def classify_electric_field(E):
    # BEGIN SOLUTION NO PROMPT
    field_strength = abs(E)
    if field_strength < 1000:
        return "weak"
    elif field_strength < 10000:
        return "moderate"
    elif field_strength < 100000:
        return "strong"
    else:
        return "extreme"
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q1b")
```

### Q1c: Electromagnetic Force Calculator
Building on your previous functions, write a function `calculate_electromagnetic_force(q1, q2, distance, k=8.99e9)` that calculates and classifies the electromagnetic force between two charges.

The function should:
- Calculate the force using Coulomb's law: F = k|q1*q2|/r²
- Use your `classify_charge` function to determine the charges' types
- Return a tuple (force_magnitude, force_type) where:
  - `force_magnitude` is the absolute force value in Newtons
  - `force_type` is:
    - "repulsive" if both charges have the same sign (both positive or both negative)
    - "attractive" if charges have opposite signs  
    - "no_force" if either charge is neutral

Note: k is Coulomb's constant (8.99×10⁹ N⋅m²/C²)

```python
def calculate_electromagnetic_force(q1, q2, distance, k=8.99e9):
    # BEGIN SOLUTION NO PROMPT
    # Get charge classifications
    charge1_type = classify_charge(q1)
    charge2_type = classify_charge(q2)
    
    # Check if either charge is neutral
    if charge1_type == "neutral" or charge2_type == "neutral":
        return (0, "no_force")
    
    # Calculate force magnitude using Coulomb's law
    force_magnitude = k * abs(q1 * q2) / (distance ** 2)
    
    # Determine force type based on charge signs
    if (charge1_type == "positive" and charge2_type == "positive") or \
       (charge1_type == "negative" and charge2_type == "negative"):
        force_type = "repulsive"
    else:
        force_type = "attractive"
    
    return (force_magnitude, force_type)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    # Hint: Use your classify_charge function from Q1a
    """; # END PROMPT
```

```python
grader.check("q1c")
```

## **Question 2**: Dictionaries and Iteration - Periodic Table Analysis

You may find `dict.keys()` and `dict.values()` useful.

```python
# Periodic table dictionary: atomic number -> element name
periodic_table = {
    1: "Hydrogen", 2: "Helium", 3: "Lithium", 4: "Beryllium", 5: "Boron",
    6: "Carbon", 7: "Nitrogen", 8: "Oxygen", 9: "Fluorine", 10: "Neon",
    11: "Sodium", 12: "Magnesium", 13: "Aluminum", 14: "Silicon", 15: "Phosphorus",
    16: "Sulfur", 17: "Chlorine", 18: "Argon", 19: "Potassium", 20: "Calcium",
    21: "Scandium", 22: "Titanium", 23: "Vanadium", 24: "Chromium", 25: "Manganese",
    26: "Iron", 27: "Cobalt", 28: "Nickel", 29: "Copper", 30: "Zinc",
    31: "Gallium", 32: "Germanium", 33: "Arsenic", 34: "Selenium", 35: "Bromine",
    36: "Krypton", 37: "Rubidium", 38: "Strontium", 39: "Yttrium", 40: "Zirconium",
    41: "Niobium", 42: "Molybdenum", 43: "Technetium", 44: "Ruthenium", 45: "Rhodium",
    46: "Palladium", 47: "Silver", 48: "Cadmium", 49: "Indium", 50: "Tin",
    51: "Antimony", 52: "Tellurium", 53: "Iodine", 54: "Xenon", 55: "Cesium",
    56: "Barium", 57: "Lanthanum", 58: "Cerium", 59: "Praseodymium", 60: "Neodymium",
    61: "Promethium", 62: "Samarium", 63: "Europium", 64: "Gadolinium", 65: "Terbium",
    66: "Dysprosium", 67: "Holmium", 68: "Erbium", 69: "Thulium", 70: "Ytterbium",
    71: "Lutetium", 72: "Hafnium", 73: "Tantalum", 74: "Tungsten", 75: "Rhenium",
    76: "Osmium", 77: "Iridium", 78: "Platinum", 79: "Gold", 80: "Mercury",
    81: "Thallium", 82: "Lead", 83: "Bismuth", 84: "Polonium", 85: "Astatine",
    86: "Radon", 87: "Francium", 88: "Radium", 89: "Actinium", 90: "Thorium",
    91: "Protactinium", 92: "Uranium", 93: "Neptunium", 94: "Plutonium", 95: "Americium",
    96: "Curium", 97: "Berkelium", 98: "Californium", 99: "Einsteinium", 100: "Fermium",
    101: "Mendelevium", 102: "Nobelium", 103: "Lawrencium"
}
```

### Q2a: Element Name Lookup
Write a function `get_element_info(atomic_numbers)` that takes a list of atomic numbers and returns a list of element names using the `periodic_table` dictionary.

The function should:
- Take a list of atomic numbers (integers from 1 to 103)
- Use the `periodic_table` dictionary to look up each element name
- Return a list of element names in the same order as the input
- If an atomic number is not in the dictionary (outside range 1-103), include "Unknown" in the result list

Example: `get_element_info([1, 6, 8])` should return `["Hydrogen", "Carbon", "Oxygen"]`

```python
def get_element_info(atomic_numbers):
    # BEGIN SOLUTION NO PROMPT
    element_names = []
    for atomic_number in atomic_numbers:
        if atomic_number in periodic_table:
            element_names.append(periodic_table[atomic_number])
        else:
            element_names.append("Unknown")
    return element_names
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q2a")
```

### Q2b: Atomic Number Lookup
Write a function `get_atomic_numbers(element_names)` that takes a list of element names (strings) and returns a list of their corresponding atomic numbers using the `periodic_table` dictionary.

The function should:
- Take a list of element names (strings like "Hydrogen", "Carbon", etc.)
- Use the `periodic_table` dictionary to find each element's atomic number
- Return a list of atomic numbers in the same order as the input
- If an element name is not found in the dictionary, include -1 in the result list

Example: `get_atomic_numbers(["Hydrogen", "Carbon", "Oxygen"])` should return `[1, 6, 8]`

Note: This will be case sensitive, so `"oxygen"` adds -1 to the list.

Hint: You'll need to search through the dictionary to find which key has the matching value.

```python
def get_atomic_numbers(element_names):
    # BEGIN SOLUTION NO PROMPT
    atomic_numbers = []
    for element_name in element_names:
        # Search through the dictionary to find the atomic number
        found = False
        for atomic_num, name in periodic_table.items():
            if name == element_name:
                atomic_numbers.append(atomic_num)
                found = True
                break
        if not found:
            atomic_numbers.append(-1)
    return atomic_numbers
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    # Hint: For each name, search through periodic_table.items() to find the matching atomic number
    """; # END PROMPT
```

```python
grader.check("q2b")
```

## **Question 3**: Metrics on R^n

### Q3a: Euclidean Distance
Implement the Euclidean distance between two points in $\mathbb{R}^n$. The Euclidean distance is defined as:

$$d_{\text{euclidean}}(\mathbf{p}, \mathbf{q}) = \sqrt{\sum_{i=1}^{n} (p_i - q_i)^2}$$

where $\mathbf{p} = (p_1, p_2, \ldots, p_n)$ and $\mathbf{q} = (q_1, q_2, \ldots, q_n)$ are two points in $\mathbb{R}^n$.

Write a function `euclidean_distance(point1, point2)` that calculates the Euclidean distance between two points in n-dimensional space.

The function should:
- Take two tuples representing points in n-dimensional space (e.g., `(1, 2)` for 2D, `(1, 2, 3)` for 3D)
- Calculate the Euclidean distance using the formula above
- Return the distance as a float
- Work for any number of dimensions (as long as both points have the same dimension)

Example: `euclidean_distance((0, 0), (3, 4))` should return `5.0`

```python
def euclidean_distance(point1, point2):
    # BEGIN SOLUTION NO PROMPT
    sum_of_squares = 0
    for i in range(len(point1)):
        diff = point1[i] - point2[i]
        sum_of_squares += diff ** 2
    return math.sqrt(sum_of_squares)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q3a")
```

### Q3b: Manhattan Distance
Implement the taxicab distance (also known as Manhattan distance) between two points in $\mathbb{R}^n$. The taxicab distance is defined as:

$$d_{\text{taxicab}}(\mathbf{p}, \mathbf{q}) = \sum_{i=1}^{n} |p_i - q_i|$$

where $\mathbf{p} = (p_1, p_2, \ldots, p_n)$ and $\mathbf{q} = (q_1, q_2, \ldots, q_n)$ are two points in $\mathbb{R}^n$.

Write a function `taxicab_distance(point1, point2)` that calculates the taxicab distance between two points in $\mathbb{R}^n$.

The function should:
- Take two tuples representing points in n-dimensional space (e.g., `(1, 2)` for 2D, `(1, 2, 3)` for 3D)
- Calculate the taxicab distance using the formula: d = Σ|x_i - y_i|
- Return the distance as a float
- Work for any number of dimensions (as long as both points have the same dimension)

Example: `taxicab_distance((0, 0), (3, 4))` should return `7.0`

```python
def taxicab_distance(point1, point2):
    # BEGIN SOLUTION NO PROMPT
    distance = 0
    for i in range(len(point1)):
        distance += abs(point1[i] - point2[i])
    return float(distance)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q3b")
```

### Q3c: Distance Threshold
Implement a function that filters points based on their distance from a reference point using a specified metric.

Write a function `filter_points_by_distance(point, distance, point_list, metric)` that returns all points in `point_list` that are at most `distance` away from `point`, using the specified distance metric.

The function should:
- Take a reference point (tuple/list), maximum distance (float), list of candidate points, and metric string
- Use the specified metric to calculate distances: "euclidean" or "taxicab"
- Return a list of points that are within the specified distance (inclusive)
- The returned list should maintain the original order of points from `point_list`

Parameters:
- `point`: Reference point (tuple/list)
- `distance`: Maximum distance threshold (float)
- `point_list`: List of candidate points to filter
- `metric`: String specifying distance metric ("euclidean" or "taxicab")

Example: `filter_points_by_distance((0, 0), 5.0, [(1, 1), (3, 4), (6, 8)], "euclidean")` should return `[(1, 1), (3, 4)]`

```python
def filter_points_by_distance(point, distance, point_list, metric):
    # BEGIN SOLUTION NO PROMPT
    filtered_points = []
    
    for candidate_point in point_list:
        if metric == "euclidean":
            dist = euclidean_distance(point, candidate_point)
        elif metric == "taxicab":
            dist = taxicab_distance(point, candidate_point)
        else:
            # Invalid metric, skip this point
            continue
            
        if dist <= distance:
            filtered_points.append(candidate_point)
    
    return filtered_points
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q3c")
```

## **Question 4**: Matrix Computations with SciPy

In this question, you'll explore matrix operations using SciPy's linear algebra module. You'll need to import and use functions from `scipy.linalg` and represent matrices as NumPy arrays.

**Documentation**: [SciPy Linear Algebra Documentation](https://docs.scipy.org/doc/scipy/reference/linalg.html)

**Matrix Multiplication**: You can multiply matrices using `np.matmul(A, B)` or the `@` operator: `A @ B` when the matrices are represented by numpy arrays.

```python
# Import necessary libraries for matrix computations
import numpy as np
import scipy.linalg
```

### Q4a: Matrix Operations and Properties

Write a function `analyze_matrix(A)` that computes several important properties of a square matrix using `scipy.linalg` functions.

The function should:
- Take a square matrix `A` (as a NumPy array)
- Calculate the determinant of A
- Calculate the trace of A (sum of diagonal elements)
- Calculate the matrix norm of A (Frobenius Norm)
- Return a tuple `(determinant, trace, norm)` as floats

Look through the scipy.linalg documentation to find the appropriate functions for these calculations.

Example: For a 2×2 identity matrix, the function should return `(1.0, 2.0, sqrt(2))`

```python
def analyze_matrix(A):
    # BEGIN SOLUTION NO PROMPT
    determinant = scipy.linalg.det(A)
    trace = np.trace(A)
    norm = scipy.linalg.norm(A)
    return (float(determinant), float(trace), float(norm))
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    # Hint: Look for det, trace, and norm functions in scipy.linalg documentation
    """; # END PROMPT
```

```python
grader.check("q4a")
```

### Q4b: Eigenvalues and Eigenvectors

Write a function `find_dominant_eigenvalue(A)` that finds the eigenvalue with the largest absolute value (magnitude) and its corresponding eigenvector.

The function should:
- Take a square matrix `A` (as a NumPy array)
- Calculate all eigenvalues and eigenvectors of A
- Find the eigenvalue with the largest absolute value (the "dominant" eigenvalue)
- Return a tuple `(dominant_eigenvalue, corresponding_eigenvector)` where both are real numbers/arrays
- If the dominant eigenvalue/eigenvector is complex, return its real part (You may find the numpy array methods .iscomplexobj and .real useful here)
- The eigenvector should be normalized (unit length) (look into whether this is default for scipy function you end up using)
- In the degenerate case, return just one of the eigenvectors

Look through the scipy.linalg documentation to find functions for eigenvalue/eigenvector computation.

Example: For the matrix `[[3, 1], [0, 2]]`, the dominant eigenvalue should be 3.0

```python
def find_dominant_eigenvalue(A):
    # BEGIN SOLUTION NO PROMPT
    eigenvalues, eigenvectors = scipy.linalg.eig(A)
    
    # Find index of eigenvalue with largest absolute value
    dominant_idx = np.argmax(np.abs(eigenvalues))
    
    # Get the dominant eigenvalue and corresponding eigenvector
    dominant_eigenvalue = eigenvalues[dominant_idx]
    dominant_eigenvector = eigenvectors[:, dominant_idx]
    
    # Convert to real if complex (take real part)
    if np.iscomplexobj(dominant_eigenvalue):
        dominant_eigenvalue = dominant_eigenvalue.real
    if np.iscomplexobj(dominant_eigenvector):
        dominant_eigenvector = dominant_eigenvector.real
    
    return (float(dominant_eigenvalue), dominant_eigenvector)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q4b")
```

### Q4c: Linear Regression using Matrix Operations

Write a function `linear_regression(X, y)` that performs linear regression using matrix operations to find the best-fit coefficients.

**Matrix Formula for Linear Regression:**
$$\boldsymbol{\beta} = (\mathbf{X}^T \mathbf{X})^{-1} \mathbf{X}^T \mathbf{y}$$

Where:
- $\mathbf{X}$ is the design matrix (already includes a column of ones for the intercept)
- $\mathbf{y}$ is the target vector
- $\boldsymbol{\beta}$ contains the coefficients (intercept and slopes)

The function should:
- Take design matrix `X` (n×m array where n=samples, m=features including intercept column) and target vector `y` (n×1)
- Use the matrix formula above with scipy.linalg functions to compute coefficients
- Return the coefficient vector β as a NumPy array

Use scipy.linalg functions for matrix operations (transpose, matrix multiplication, inverse).

```python
def linear_regression(X, y):
    # BEGIN SOLUTION NO PROMPT
    # Apply the linear regression formula: β = (X^T X)^(-1) X^T y
    # Step 1: X^T (transpose)
    X_transpose = X.T
    
    # Step 2: X^T X
    XTX = X_transpose @ X
    
    # Step 3: (X^T X)^(-1)
    XTX_inverse = scipy.linalg.inv(XTX)
    
    # Step 4: (X^T X)^(-1) X^T
    XTX_inv_XT = XTX_inverse @ X_transpose
    
    # Step 5: (X^T X)^(-1) X^T y
    beta = XTX_inv_XT @ y
    
    return beta
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q4c")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
# BEGIN SOLUTION NO PROMPT
# END SOLUTION
""" # BEGIN PROMPT
"""
# write ai disclosure here:

"""
"""; # END PROMPT
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> classify_charge(1.6e-19) == 'positive'
True
```
```
>>> classify_charge(-1.6e-19) == 'negative'
True
```
```
>>> classify_charge(0) == 'neutral'
True
```


Tests for q1b (points: None)
```
>>> classify_electric_field(500) == 'weak'
True
```
```
>>> classify_electric_field(5000) == 'moderate'
True
```
```
>>> classify_electric_field(50000) == 'strong'
True
```
```
>>> classify_electric_field(200000) == 'extreme'
True
```
```
>>> classify_electric_field(-75000) == 'strong'
True
```
```
>>> classify_electric_field(0) == 'weak'
True
```


Tests for q1c (points: None)
```
>>> force, force_type = calculate_electromagnetic_force(1e-06, 2e-06, 0.1)
>>> abs(force - 1.798) < 0.01 and force_type == 'repulsive'
True
```
```
>>> force, force_type = calculate_electromagnetic_force(1e-06, -1e-06, 0.1)
>>> abs(force - 0.899) < 0.01 and force_type == 'attractive'
True
```
```
>>> force, force_type = calculate_electromagnetic_force(0, 1e-06, 0.1)
>>> force == 0 and force_type == 'no_force'
True
```
```
>>> force, force_type = calculate_electromagnetic_force(-2e-06, -3e-06, 0.05)
>>> abs(force - 21.576) < 0.1 and force_type == 'repulsive'
True
```
```
>>> force, force_type = calculate_electromagnetic_force(0, 0, 0.1)
>>> force == 0 and force_type == 'no_force'
True
```


Tests for q2a (points: None)
```
>>> get_element_info([1, 6, 8]) == ['Hydrogen', 'Carbon', 'Oxygen']
True
```
```
>>> get_element_info([2, 10, 18]) == ['Helium', 'Neon', 'Argon']
True
```
```
>>> get_element_info([104, 200]) == ['Unknown', 'Unknown']
True
```
```
>>> get_element_info([79, 104, 26]) == ['Gold', 'Unknown', 'Iron']
True
```
```
>>> get_element_info([]) == [] and get_element_info([47]) == ['Silver']
True
```


Tests for q2b (points: None)
```
>>> get_atomic_numbers(['Hydrogen', 'Carbon', 'Oxygen']) == [1, 6, 8]
True
```
```
>>> get_atomic_numbers(['Helium', 'Neon', 'Argon']) == [2, 10, 18]
True
```
```
>>> get_atomic_numbers(['Iron', 'Gold', 'Silver']) == [26, 79, 47]
True
```
```
>>> get_atomic_numbers(['Unknownium', 'Fakium']) == [-1, -1]
True
```
```
>>> get_atomic_numbers(['Lithium', 'Unknownium', 'Fluorine']) == [3, -1, 9]
True
```
```
>>> get_atomic_numbers([]) == [] and get_atomic_numbers(['Uranium']) == [92]
True
```
```
>>> get_atomic_numbers(['Lead', 'lawrencium']) == [82, -1]
True
```


Tests for q3a (points: None)
```
>>> abs(euclidean_distance((0, 0), (3, 4)) - 5.0) < 1e-10
True
```
```
>>> abs(euclidean_distance((0, 0, 0), (1, 1, 1)) - math.sqrt(3)) < 1e-10
True
```
```
>>> euclidean_distance((2, 3), (2, 3)) == 0.0
True
```
```
>>> abs(euclidean_distance((1, 2, 3, 4), (5, 6, 7, 8)) - 8.0) < 1e-10
True
```


Tests for q3b (points: None)
```
>>> taxicab_distance((0, 0), (3, 4)) == 7.0
True
```
```
>>> taxicab_distance((1, 2, 3), (4, 6, 8)) == 12.0
True
```
```
>>> taxicab_distance((1, 2, 3, 4, 5), (6, 7, 8, 9, 10)) == 25.0
True
```
```
>>> taxicab_distance((2, 3), (2, 3)) == 0.0
True
```
```
>>> taxicab_distance((-1, -2), (1, 2)) == 6.0
True
```


Tests for q3c (points: None)
```
>>> points = [(0, 0), (1, 1), (3, 4), (2, 2)]
>>> result = filter_points_by_distance((0, 0), 2.5, points, 'euclidean')
>>> result == [(0, 0), (1, 1)]
True
```
```
>>> points = [(0, 0), (1, 1), (3, 4), (2, 2)]
>>> result = filter_points_by_distance((0, 1), 4, points, 'taxicab')
>>> result == [(0, 0), (1, 1), (2, 2)]
True
```
```
>>> points = [(5, 5), (6, 6), (7, 7)]
>>> result = filter_points_by_distance((0, 0), 1, points, 'euclidean')
>>> result == []
True
```
```
>>> points_3d = [(0, 0, 0), (1, 1, 1), (2, 0, 0), (0, 2, 0), (0, 0, 2), (2, 2, 0), (2, 2, 2)]
>>> result = filter_points_by_distance((0, 0, 0), 3, points_3d, 'taxicab')
>>> result == [(0, 0, 0), (1, 1, 1), (2, 0, 0), (0, 2, 0), (0, 0, 2)]
True
```
```
>>> points_5d = [(1, 2, 3, 4, 5), (6, 7, 8, 9, 10), (2, 2, 2, 2, 2), (0, 0, 0, 0, 0)]
>>> result = filter_points_by_distance((0, 0, 0, 0, 0), 10, points_5d, 'euclidean')
>>> result == [(1, 2, 3, 4, 5), (2, 2, 2, 2, 2), (0, 0, 0, 0, 0)]
True
```


Tests for q4a (points: None)
```
>>> import numpy as np
>>> import scipy.linalg
>>> I2 = np.eye(2)
>>> det, trace, norm = analyze_matrix(I2)
>>> bool(abs(det - 1.0) < 1e-10 and abs(trace - 2.0) < 1e-10 and (abs(norm - np.sqrt(2)) < 1e-10))
True
```
```
>>> import numpy as np
>>> import scipy.linalg
>>> A = np.array([[1, 2], [3, 4]])
>>> det, trace, norm = analyze_matrix(A)
>>> bool(abs(det - -2.0) < 1e-10 and abs(trace - 5.0) < 1e-10)
True
```
```
>>> import numpy as np
>>> import scipy.linalg
>>> B = np.array([[2, 0, 0], [0, 3, 0], [0, 0, 4]])
>>> det, trace, norm = analyze_matrix(B)
>>> bool(abs(det - 24.0) < 1e-10 and abs(trace - 9.0) < 1e-10 and (abs(norm - np.sqrt(29)) < 1e-10))
True
```


Tests for q4b (points: None)
```
>>> import numpy as np
>>> import scipy.linalg
>>> A = np.array([[3, 1], [0, 2]])
>>> dominant_eigenval, dominant_eigenvec = find_dominant_eigenvalue(A)
>>> bool(abs(dominant_eigenval - 3.0) < 1e-10 and abs(scipy.linalg.norm(dominant_eigenvec) - 1.0) < 1e-10)
True
```
```
>>> import numpy as np
>>> import scipy.linalg
>>> B = np.array([[5, 0], [0, -2]])
>>> dominant_eigenval, dominant_eigenvec = find_dominant_eigenvalue(B)
>>> bool(abs(dominant_eigenval - 5.0) < 1e-10)
True
```
```
>>> import numpy as np
>>> import scipy.linalg
>>> C = np.array([[4, 0], [0, 1]])
>>> dominant_eigenval, dominant_eigenvec = find_dominant_eigenvalue(C)
>>> bool(abs(scipy.linalg.norm(dominant_eigenvec) - 1.0) < 1e-10)
True
```


Tests for q4c (points: None)
```
>>> import numpy as np
>>> import scipy.linalg
>>> X_simple = np.array([[1, 1], [1, 2], [1, 3], [1, 4]])
>>> y_simple = np.array([3, 5, 7, 9])
>>> beta = linear_regression(X_simple, y_simple)
>>> bool(abs(beta[0] - 1.0) < 1e-10 and abs(beta[1] - 2.0) < 1e-10)
True
```
```
>>> import numpy as np
>>> import scipy.linalg
>>> X_multi = np.array([[1, 1, 1], [1, 2, 1], [1, 3, 2], [1, 4, 2]])
>>> y_multi = np.array([6, 9, 14, 17])
>>> beta = linear_regression(X_multi, y_multi)
>>> bool(abs(beta[0] - 1.0) < 1e-10 and abs(beta[1] - 3.0) < 1e-10 and (abs(beta[2] - 2.0) < 1e-10))
True
```
```
>>> import numpy as np
>>> import scipy.linalg
>>> X_const = np.array([[1, 1], [1, 2], [1, 3]])
>>> y_const = np.array([5, 5, 5])
>>> beta = linear_regression(X_const, y_const)
>>> bool(abs(beta[0] - 5.0) < 1e-10 and abs(beta[1])) < 1e-10
True
```
