```python
# Initialize Otter
import otter
grader = otter.Notebook("ws4.ipynb")
```

```python
# Setup
# You may import math if needed in your solutions.
import math
import numpy as np
```

## **Question 1**: Working with NumPy String Operations

**Your task:**
Write a function `analyze_physics_terms(terms_array)` that processes an array of physics terminology using various `numpy.strings` operations.

**Requirements:**
Use functions from the `numpy.strings` module to perform text analysis on physics terminology.

**Function Specifications:**

`analyze_physics_terms(terms_array)` should:

1. **Convert to uppercase**: Convert all terms to uppercase using `numpy.strings.upper()`
2. **Filter by length**: Find terms that are longer than 5 characters (after uppercasing)
3. **Count characters**: Count how many times the letter 'I' appears in each term (after uppercasing)
4. **Check suffixes**: Determine which terms end with the letter 'S' (after uppercasing)
5. **Extract root words**: Remove common physics suffixes ('ICS', 'ION', 'URE') from terms to get root words

**Parameters:**
- `terms_array`: NumPy array of strings containing physics terms

**Return:**
A dictionary with the following keys:
- `'uppercase_terms'`: numpy array of all terms converted to uppercase
- `'long_terms'`: numpy array of terms longer than 5 characters (after uppercasing)
- `'i_counts'`: numpy array of integers showing how many times 'I' appears in each term
- `'ends_with_s'`: numpy array of booleans indicating which terms end with 'S'
- `'root_words'`: numpy array of root words with suffixes removed (words that don't end with the suffixes should be returned unchanged)

**Hint:** 
- Use `numpy.strings` documentation: https://numpy.org/doc/stable/reference/routines.strings.html
- Functions you'll use: `upper()`, `str_len()`, `count()`, `endswith()`
- Use boolean indexing with the results of string comparisons

**Example:**
```python
terms = np.array(['physics', 'optics', 'energy'])
result = analyze_physics_terms(terms)
# result['uppercase_terms'] would be ['PHYSICS', 'OPTICS', 'ENERGY']
# result['long_terms'] would be ['PHYSICS', 'OPTICS', 'ENERGY'] (all > 5 chars)
# result['i_counts'] would be [1, 1, 0] (count of 'I' in each term)
# result['ends_with_s'] would be [True, True, False] (which end with 'S')
# result['root_words'] would be ['PHYS', 'OPT', 'ENERGY'] (suffixes removed)
```

```python
def analyze_physics_terms(terms):
#write code here
return {
        'uppercase_terms': upper_terms,
        'long_terms': long_terms,
        'i_counts': i_counts,
        'ends_with_s': ends_with_s,
        'root_words': root_words
    }

```

```python
grader.check("q1")
```

## **Question 2**: Creating an Earthquake Class

**Your task:**
Create a class called `Earthquake` that models seismic events and provides methods to analyze earthquake properties.

**Requirements:**
Define a class with an `__init__` method and several analysis methods using basic mathematical operations.

**Class Specifications:**

Your `Earthquake` class should have:

**Constructor: `__init__(self, magnitude, depth_km, distance_km)`**
- `magnitude`: Richter scale magnitude (float)
- `depth_km`: depth of earthquake focus in kilometers (float) 
- `distance_km`: distance from earthquake epicenter in kilometers (float)

**Methods to implement:**

1. **`energy_joules(self)`**
   - Calculate released energy using: Energy = 10^(11.8 + 1.5 * magnitude) Joules
   - Return the energy as a float

2. **`intensity_scale(self)`** 
   - Calculate Modified Mercalli intensity using: I = magnitude + 3.0 * log10(depth_km) - 3.0 * log10(distance_km)
   - Use `math.log10()` for the logarithm
   - Return the intensity as a float

3. **`classification(self)`**
   - Return earthquake classification as a string based on magnitude:
     - magnitude < 3.0: "Minor"
     - 3.0 ≤ magnitude < 5.0: "Light" 
     - 5.0 ≤ magnitude < 7.0: "Moderate"
     - magnitude ≥ 7.0: "Major"

**Example:**
```python
quake = Earthquake(magnitude=6.5, depth_km=10.0, distance_km=50.0)
energy = quake.energy_joules()  # Calculate energy released
intensity = quake.intensity_scale()  # Calculate intensity  
classification = quake.classification()  # Returns "Moderate"
```

Write your `Earthquake` class below:

```python
class Earthquake:
    # Write your Earthquake class here!
    pass
```

```python
grader.check("q2")
```

## **Question 3**: Safe Factorial Function with Error Handling

**Your task:**
Write a function `safe_factorial(n)` that calculates the factorial of a number with comprehensive error handling using try/except/else/finally blocks.

**Requirements:**
The function must validate inputs and handle various error conditions before performing the calculation.

**Function Specifications:**

Your `safe_factorial(n)` function should:

**Required Error Checks:**
1. **Type validation**: Input must be an integer (not float, string, etc.)
2. **Value constraints**: 
   - Number must be non-negative (n ≥ 0)
   - Number must be reasonably small to avoid overflow (n ≤ 170)
3. **Mathematical correctness**: Handle potential computational issues

**Return behavior:**
- **Success**: Return the factorial as an integer (n!)
- **Any error**: Return a descriptive error message as a string in the format: `"Error: [descriptive message]"`

**Structure requirements:**
- Use `try` block for the main calculation
- Use multiple `except` blocks to catch and handle different types of errors
- Include proper input validation within the try block

**Mathematical Background:**
- Factorial: n! = n × (n-1) × (n-2) × ... × 2 × 1
- Special case: 0! = 1
- Only defined for non-negative integers

**Parameters:**
- `n`: The number to calculate factorial for

**Example:**
```python
result1 = safe_factorial(5)      # Returns 120
result2 = safe_factorial(-3)     # Returns "Error: Factorial not defined for negative numbers"
result3 = safe_factorial("5")    # Returns "Error: Input must be an integer"
result4 = safe_factorial(200)    # Returns "Error: Input too large, may cause overflow"
```

Hint: Your except blocks can be of this form:
```python
   except TypeError:
        return "Error: Invalid input type for factorial calculation"
    except ValueError as e:
        return f"Error: Mathematical error - {str(e)}"
    except OverflowError:
        return "Error: Number too large for calculation"
    except MemoryError:
        return "Error: Insufficient memory for calculation"
    except Exception as e:
        return f"Error: Unexpected error - {str(e)}"
```

Write your function below:

```python
def safe_factorial(n):
    # Write your safe_factorial function here!
    pass
```

```python
grader.check("q3")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: None)
```
>>> test_terms = np.array(['physics', 'energy', 'quantum'])
>>> result = analyze_physics_terms(test_terms)
>>> isinstance(result, dict) and all((key in result for key in ['uppercase_terms', 'long_terms', 'i_counts', 'ends_with_s', 'root_words']))
True
```
```
>>> test_terms = np.array(['physics', 'energy'])
>>> result = analyze_physics_terms(test_terms)
>>> expected_uppercase = np.array(['PHYSICS', 'ENERGY'])
>>> np.array_equal(result['uppercase_terms'], expected_uppercase)
True
```
```
>>> test_terms = np.array(['atom', 'quantum', 'proton'])
>>> result = analyze_physics_terms(test_terms)
>>> expected_long = np.array(['QUANTUM', 'PROTON'])
>>> np.array_equal(result['long_terms'], expected_long)
True
```
```
>>> test_terms = np.array(['optics', 'kinetics'])
>>> result = analyze_physics_terms(test_terms)
>>> expected_i_counts = np.array([1, 2])
>>> np.array_equal(result['i_counts'], expected_i_counts)
True
```
```
>>> test_terms = np.array(['physics', 'energy', 'atoms'])
>>> result = analyze_physics_terms(test_terms)
>>> expected_ends_s = np.array([True, False, True])
>>> np.array_equal(result['ends_with_s'], expected_ends_s)
True
```
```
>>> test_terms = np.array(['mechanics', 'thermodynamics', 'gamma'])
>>> result = analyze_physics_terms(test_terms)
>>> expected_roots = np.array(['MECHAN', 'THERMODYNAM', 'GAMMA'])
>>> np.array_equal(result['root_words'], expected_roots)
True
```
```
>>> test_terms = np.array(['photon', 'electronics', 'field'])
>>> result = analyze_physics_terms(test_terms)
>>> has_long_terms = len(result['long_terms']) == 2
>>> has_correct_i_count = result['i_counts'][1] == 1
>>> bool(has_long_terms and has_correct_i_count)
True
```
```
>>> test_terms = np.array(['physics', 'optics'])
>>> result = analyze_physics_terms(test_terms)
>>> expected_upper = np.array(['PHYSICS', 'OPTICS'])
>>> expected_long = np.array(['PHYSICS', 'OPTICS'])
>>> expected_roots = np.array(['PHYS', 'OPT'])
>>> check1 = np.array_equal(result['uppercase_terms'], expected_upper)
>>> check2 = np.array_equal(result['long_terms'], expected_long)
>>> check3 = np.array_equal(result['root_words'], expected_roots)
>>> check1 and check2 and check3
True
```


Tests for q2 (points: None)
```
>>> quake = Earthquake(5.5, 15.0, 100.0)
>>> hasattr(quake, 'magnitude') and hasattr(quake, 'depth_km') and hasattr(quake, 'distance_km')
True
```
```
>>> quake = Earthquake(4.0, 10.0, 50.0)
>>> energy = quake.energy_joules()
>>> expected_energy = 10 ** (11.8 + 1.5 * 4.0)
>>> abs(energy - expected_energy) < 10000000000.0
True
```
```
>>> quake = Earthquake(5.0, 10.0, 100.0)
>>> intensity = quake.intensity_scale()
>>> expected_intensity = 5.0 + 3.0 * math.log10(10.0) - 3.0 * math.log10(100.0)
>>> abs(intensity - expected_intensity) < 1e-10
True
```
```
>>> minor_quake = Earthquake(2.5, 5.0, 20.0)
>>> light_quake = Earthquake(4.0, 10.0, 50.0)
>>> moderate_quake = Earthquake(6.0, 15.0, 100.0)
>>> major_quake = Earthquake(7.5, 20.0, 200.0)
>>> minor_quake.classification() == 'Minor' and light_quake.classification() == 'Light' and (moderate_quake.classification() == 'Moderate') and (major_quake.classification() == 'Major')
True
```
```
>>> boundary_light = Earthquake(3.0, 10.0, 50.0)
>>> boundary_moderate = Earthquake(5.0, 10.0, 50.0)
>>> boundary_major = Earthquake(7.0, 10.0, 50.0)
>>> boundary_light.classification() == 'Light' and boundary_moderate.classification() == 'Moderate' and (boundary_major.classification() == 'Major')
True
```
```
>>> quake = Earthquake(6.5, 12.0, 80.0)
>>> energy_ok = isinstance(quake.energy_joules(), (int, float))
>>> intensity_ok = isinstance(quake.intensity_scale(), (int, float))
>>> classification_ok = isinstance(quake.classification(), str)
>>> magnitude_ok = quake.classification() == 'Moderate'
>>> energy_ok and intensity_ok and classification_ok and magnitude_ok
True
```


Tests for q3 (points: None)
```
>>> result = safe_factorial(5)
>>> isinstance(result, int) and result == 120
True
```
```
>>> result = safe_factorial(0)
>>> isinstance(result, int) and result == 1
True
```
```
>>> result = safe_factorial('5')
>>> isinstance(result, str) and result.startswith('Error:')
True
```
```
>>> result = safe_factorial(-5)
>>> isinstance(result, str) and result.startswith('Error:')
True
```
```
>>> result = safe_factorial(5.5)
>>> isinstance(result, str) and result.startswith('Error:')
True
```
```
>>> result = safe_factorial(200)
>>> isinstance(result, str) and result.startswith('Error:')
True
```
```
>>> result = safe_factorial(4)
>>> isinstance(result, int) and result == 24
True
```
