```python
# Initialize Otter
import otter
grader = otter.Notebook("ws11.ipynb")
```

```python
rng_seed = 70
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import scipy as sp
import pandas as pd
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

# **Question 1**: Solving the 1D Heat Equation with Fourier Series

In this question, you'll implement numerical solutions to the **1D heat equation** using **Fourier series analysis**. The heat equation is a fundamental partial differential equation that describes how temperature (or heat density) evolves over time in a material.

## Background: The 1D Heat Equation

The 1D heat equation is given by:

$$\frac{\partial u}{\partial t} = \alpha \frac{\partial^2 u}{\partial x^2}$$

where:
- $u(x,t)$ is the temperature/heat density at position $x$ and time $t$
- $\alpha$ is the thermal diffusivity (we'll assume $\alpha = 1$ for simplicity)
- The spatial domain is $x \in [0, L]$ where we'll use $L = \pi$ for convenience

## Neumann Boundary Conditions

We impose **Neumann boundary conditions**, which specify that the heat flux (spatial derivative) is zero at the boundaries:

$$\frac{\partial u}{\partial x}\bigg|_{x=0} = 0, \quad \frac{\partial u}{\partial x}\bigg|_{x=L} = 0$$

This means the boundaries are **insulated** - no heat flows in or out of the domain.

## Fourier Series Solution

With Neumann boundary conditions on the domain $[0, \pi]$, the appropriate basis functions are **cosines**:

$$u(x,t) = \sum_{n=0}^{\infty} A_n(t) \cos(nx)$$

The cosine functions $\cos(nx)$ are **eigenfunctions** of the second derivative operator with Neumann boundary conditions:

$$\frac{d^2}{dx^2}\cos(nx) = -n^2\cos(nx)$$

Substituting the Fourier series into the heat equation:

$$\sum_{n=0}^{\infty} \frac{dA_n}{dt} \cos(nx) = -\alpha \sum_{n=0}^{\infty} n^2 A_n(t) \cos(nx)$$

Since the cosines are orthogonal, each coefficient evolves independently:

$$\frac{dA_n}{dt} = -\alpha n^2 A_n(t)$$

This gives us the time evolution:

$$A_n(t) = A_n(0) e^{-\alpha n^2 t}$$

where $A_n(0)$ are the initial Fourier coefficients.

## Complete Solution Form

The complete solution is:

$$u(x,t) = \sum_{n=0}^{M} A_n(0) e^{-\alpha n^2 t} \cos(nx)$$

where:
- $A_n(0)$ are the initial cosine coefficients from the discrete cosine transform of the initial condition
- $M$ is the maximum order of cosines to include (truncation parameter)
- Higher frequency modes ($n > M$) are ignored

## **Part A**: Implement Heat Equation Solver

Implement `solve_heat_equation_1d(u_initial, M, t_max, N, alpha=1.0)` that:

1. **Computes initial Fourier coefficients** using the discrete cosine transform (DCT)
2. **Evolves coefficients in time** using the analytical solution $A_n(t) = A_n(0) e^{-\alpha n^2 t}$
3. **Reconstructs the solution** at each time point using the inverse DCT

**Mathematical Steps:**

1. **DCT of initial condition**: Use `scipy.fft.dct(u_initial, type=2, norm='ortho')` to get $A_n(0)$
2. **Time evolution**: For each time $t$, compute $A_n(t) = A_n(0) e^{-\alpha n^2 t}$ for $n = 0, 1, ..., M-1$
3. **Inverse DCT**: Use `scipy.fft.idct(A_t, type=2, norm='ortho')` to get $u(x,t)$

**Requirements:**
- Use only the first `M` Fourier coefficients (set higher-order coefficients to zero)
- Create time array using `np.linspace(0, t_max, N)`
- Return a 2D array of shape `(N, len(u_initial))` where each row is the solution at one time
- Use `scipy.fft.dct` and `scipy.fft.idct` with `type=2` and `norm='ortho'`

**Parameters:**
- `u_initial`: numpy array, initial heat density values at spatial grid points
- `M`: int, maximum order of cosines to include (truncation parameter)
- `t_max`: float, final time
- `N`: int, number of time points
- `alpha`: float, thermal diffusivity (default 1.0)

**Returns:**
- `solution`: numpy array of shape (N, len(u_initial)), heat density at each time and position

```python
def solve_heat_equation_1d(u_initial, M, t_max, N, alpha=1.0):
    from scipy.fft import dct, idct
    
    # Convert to numpy array
    u_initial = np.array(u_initial)
    
    # Compute initial Fourier coefficients using DCT
    
    # Truncate to M coefficients
    
    
    # For each time point:
    #   1. Evolve coefficients: A_n(t) = A_n(0) * exp(-alpha * n^2 * t)
    #   2. Reconstruct solution using inverse DCT
    
    return solution
```

```python
grader.check("q1a")
```

## **Part B**: Visualize Heat Evolution

Implement `plot_heat_evolution(u_initial, M, t_max, N, x_positions, show_plot=False)` that solves the heat equation and visualizes how the heat density evolves over all time points.

**Requirements:**
- Call `solve_heat_equation_1d` with the provided parameters
- Create a single figure with the heat density curves at ALL time points in the solution
- Use `plt.plot()` for each time curve (not `plt.subplot()`)
- Plot configuration:
  - Each curve should have a different color (let matplotlib auto-assign colors)
  - Line width of 1 (since there will be many curves)
  - Label only every 5th curve as `f't = {time:.2f}'` to avoid legend clutter
  - X-axis label: "Position (x)"
  - Y-axis label: "Heat Density u(x,t)"
  - Title: "1D Heat Equation Evolution"
  - Grid with alpha=0.3
  - Legend
- Only call `plt.show()` if `show_plot=True`
- Return the figure object

**Parameters:**
- `u_initial`: numpy array, initial heat density values at spatial grid points
- `M`: int, maximum order of cosines to include (truncation parameter)
- `t_max`: float, final time
- `N`: int, number of time points
- `x_positions`: numpy array, spatial coordinates for the x-axis
- `show_plot`: bool, default False. If True, display the plot

**Returns:**
- `fig`: matplotlib figure object

```python
def plot_heat_evolution(u_initial, M, t_max, N, x_positions, show_plot=False):
    # Solve the heat equation using Part A function
    solution = solve_heat_equation_1d(u_initial, M, t_max, N)
    
    # Create time array
    t_array = np.linspace(0, t_max, N)
    
    # Create the figure
    fig, ax = plt.subplots(figsize=(12, 8))
    
    # For each time point:
    #   1. Plot the curve with plt.plot()
    #   2. Label every 5th curve to avoid clutter (label should be "t = ...")
    
    # Set labels, title, grid, and legend
    
    # Show plot if requested
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Example: Solve and visualize heat equation for different initial conditions

# Example 1: Step function initial condition
print("Example 1: Step function initial condition")
x = np.linspace(0, np.pi, 100)
u_step = np.where(x < np.pi/2, 1.0, 0.0)

fig1 = plot_heat_evolution(u_step, M=100, t_max=3, N=30, x_positions=x, show_plot=True)
print()

# Example 2: Gaussian initial condition
print("Example 2: Gaussian initial condition")
u_gauss = np.exp(-((x - np.pi/2)**2) / 0.2)

fig2 = plot_heat_evolution(u_gauss, M=15, t_max=1.0, N=25, x_positions=x, show_plot=True)
print()

# Example 3: Sinusoidal initial condition
print("Example 3: Sinusoidal initial condition")
u_sin = np.cos(2*x) + 0.5*np.cos(4*x)

fig3 = plot_heat_evolution(u_sin, M=10, t_max=0.8, N=20, x_positions=x, show_plot=True)
```

```python
grader.check("q1b")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> u_test = np.array([1, 2, 3, 4, 5])
>>> solution = solve_heat_equation_1d(u_test, M=3, t_max=1.0, N=10)
>>> test_result = solution.shape == (10, 5)
>>> test_result
True
```
```
>>> u_test = np.array([1, 0, -1, 0, 1])
>>> solution = solve_heat_equation_1d(u_test, M=5, t_max=1.0, N=20)
>>> test_result = np.allclose(solution[0, :], u_test, atol=1e-10)
>>> test_result
True
```
```
>>> u_test = np.array([0, 1, 0, -1, 0])
>>> solution = solve_heat_equation_1d(u_test, M=5, t_max=2.0, N=100)
>>> initial_var = np.var(solution[0, :])
>>> final_var = np.var(solution[-1, :])
>>> test_result = final_var < initial_var
>>> bool(test_result)
True
```
```
>>> u_test = np.ones(8) * 5.0
>>> solution = solve_heat_equation_1d(u_test, M=4, t_max=1.0, N=10)
>>> test_result = np.allclose(solution, 5.0, atol=1e-10)
>>> test_result
True
```


Tests for q1b (points: None)
```
>>> x_test = np.linspace(0, np.pi, 10)
>>> u_test = np.ones(10)
>>> fig_test = plot_heat_evolution(u_test, M=5, t_max=1.0, N=10, x_positions=x_test)
>>> test_result = isinstance(fig_test, matplotlib.figure.Figure)
>>> plt.close(fig_test)
>>> test_result
True
```
```
>>> x_test = np.linspace(0, np.pi, 8)
>>> u_test = np.array([1, 2, 3, 2, 1, 0, -1, 0])
>>> fig_test = plot_heat_evolution(u_test, M=4, t_max=0.5, N=15, x_positions=x_test)
>>> ax = fig_test.axes[0]
>>> has_multiple_lines = len(ax.lines) == 15
>>> has_xlabel = ax.get_xlabel() == 'Position (x)'
>>> has_ylabel = ax.get_ylabel() == 'Heat Density u(x,t)'
>>> has_title = 'Heat Equation' in ax.get_title()
>>> has_legend = ax.get_legend() is not None
>>> plt.close(fig_test)
>>> test_result = has_multiple_lines and has_xlabel and has_ylabel and has_title and has_legend
>>> test_result
True
```
```
>>> x_test = np.linspace(0, np.pi, 6)
>>> u_test = np.array([0, 1, 2, 1, 0, -1])
>>> fig_test = plot_heat_evolution(u_test, M=3, t_max=1.0, N=20, x_positions=x_test)
>>> ax = fig_test.axes[0]
>>> legend_labels = [text.get_text() for text in ax.get_legend().get_texts()]
>>> has_time_labels = any(('t = ' in label for label in legend_labels))
>>> plt.close(fig_test)
>>> test_result = has_time_labels
>>> test_result
True
```
```
>>> x_test = np.linspace(0, np.pi, 5)
>>> u_test = np.array([1, 0, 1, 0, 1])
>>> fig_test = plot_heat_evolution(u_test, M=5, t_max=0.1, N=10, x_positions=x_test)
>>> ax = fig_test.axes[0]
>>> has_correct_lines = len(ax.lines) == 10
>>> has_grid = ax.grid
>>> plt.close(fig_test)
>>> test_result = has_correct_lines
>>> test_result
True
```
