```python
# Initialize Otter
import otter
grader = otter.Notebook("hw10.ipynb")
```

```python
rng_seed = 60
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import scipy as sp
import pandas as pd
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

# **Question 1**: Numerical Differentiation Methods

In this question, you'll implement and compare different methods for numerically computing derivatives of a function. Numerical differentiation is essential when we don't have an analytical formula for the derivative, or when we only have discrete data points.

### The Function

Throughout this question, we'll work with the function:

$$f(x) = 3\cos(2x) + \sin(3x) + e^{-x^2}$$

We'll compare numerical approximations to the exact derivative.

## **Part A**: Central Finite Difference Method

The **central finite difference** formula is a second-order accurate method for approximating the first derivative. It uses function values on both sides of the point of interest.

### Mathematical Background

The central difference formula for the first derivative at point $x$ is:

$$f'(x) \approx \frac{f(x + h) - f(x - h)}{2h}$$

where $h$ is the step size (spacing between points).

**Why "central"?** This method uses points symmetrically placed around $x$, giving better accuracy than forward or backward differences.

### Your Task

Write a function `central_diff(a, b, N)` that:
1. Evaluates $f(x)$ on a grid of $N$ points from $a$ to $b$
2. Computes the first derivative using central finite differences
3. Returns both the $x$ values and the computed derivative values

**Requirements:**
- Create $x$ array using `np.linspace(a, b, N)`
- Compute step size: $h = \frac{b-a}{N-1}$
- For interior points $i = 1, 2, ..., N-2$: use $f'(x_i) \approx \frac{f(x_{i+1}) - f(x_{i-1})}{2h}$
- For boundary points (first and last):
  - First point: use forward difference $f'(x_0) \approx \frac{f(x_1) - f(x_0)}{h}$
  - Last point: use backward difference $f'(x_{N-1}) \approx \frac{f(x_{N-1}) - f(x_{N-2})}{h}$
- Return arrays: `(x_values, derivative_values)`

**Parameters:**
- `a`: float, left endpoint of interval
- `b`: float, right endpoint of interval
- `N`: int, number of points

**Returns:**
- `x`: numpy array of shape (N,), x-coordinates
- `dfdx`: numpy array of shape (N,), derivative values at each x

```python
def central_diff(a, b, N):
    
    return x, dfdx
```

```python
grader.check("q1a")
```

## **Part B**: 5-Point Centered Finite Difference

The **5-point centered finite difference** is a higher-order method that achieves fourth-order accuracy by using more neighboring points. This significantly reduces the error compared to the 2-point central difference.

### Mathematical Background

The 5-point centered difference formula for the first derivative is:

$$f'(x) \approx \frac{-f(x+2h) + 8f(x+h) - 8f(x-h) + f(x-2h)}{12h} + O(h^4)$$

**Why is this more accurate?**

This formula is derived from a higher-order Taylor series expansion. By combining function values at ±$h$ and ±$2h$, we can cancel out error terms up to $O(h^4)$, giving us:
- 2-point central difference: $O(h^2)$ error
- 5-point central difference: $O(h^4)$ error

For small step sizes $h$, the error $h^4$ is much smaller than $h^2$, making this method significantly more accurate.


### Your Task

Write a function `five_point_diff(a, b, N)` that:
1. Evaluates $f(x)$ on a grid of $N$ points from $a$ to $b$
2. Computes the first derivative using the 5-point centered finite difference
3. Returns both the $x$ values and the computed derivative values

**Requirements:**
- Create $x$ array using `np.linspace(a, b, N)`
- Compute step size: $h = \frac{b-a}{N-1}$
- For interior points $i = 2, 3, ..., N-3$: use the 5-point formula:
  $$f'(x_i) \approx \frac{-f(x_{i+2}) + 8f(x_{i+1}) - 8f(x_{i-1}) + f(x_{i-2})}{12h}$$
- For points near boundaries (indices 0, 1, N-2, N-1):
  - Use 2-point central difference (from Part A) or forward/backward differences
  - These points won't have the high accuracy of interior points
- Return arrays: `(x_values, derivative_values)`

**Parameters:**
- `a`: float, left endpoint of interval
- `b`: float, right endpoint of interval
- `N`: int, number of points (should be at least 5)

**Returns:**
- `x`: numpy array of shape (N,), x-coordinates
- `dfdx`: numpy array of shape (N,), derivative values at each x

```python
def five_point_diff(a, b, N):
    # Define the function
    f = lambda x: 3*np.cos(2*x) + np.sin(3*x) + np.exp(-x**2)
    
    # Create x array and compute step size
    ...
    
    # Evaluate function at all points
    ...
    
    # Initialize derivative array
    dfdx = np.zeros(N)
    
    # 5-point centered difference for interior points
    ...
    
    # Handle boundary points
    ...
    
    return x, dfdx
```

```python
grader.check("q1b")
```

### Part C: Comparing Derivative Approximations (6 points)

Now we'll compare the two numerical differentiation methods with the exact analytical derivative.

**Task:** Write a function `plot_derivatives(a, b, N_list, show_plot=False)` that:
- Takes an interval `[a, b]` and a list of grid sizes `N_list`
- For each `N` in `N_list`, computes derivatives using both methods from parts A and B
- Plots the numerical derivatives alongside the exact derivative, $f'(x)$
- Plots all curves on the same plot, with appropriate (f string) labels
- Returns a matplotlib figure object
- If `show_plot=False`, suppresses the display but still returns the figure

```python
def plot_derivatives(a, b, N_list, show_plot=False):
    """
    Compare numerical derivative approximations with the exact derivative.
    
    Parameters:
    -----------
    a, b : float
        Interval endpoints
    N_list : list of int
        List of grid sizes to compare
    show_plot : bool, optional
        If True, display the plot; if False, suppress display (default: False)
    
    Returns:
    --------
    fig : matplotlib.figure.Figure
        The figure object containing the plot
    """
    fig, ax = plt.subplots(figsize=(10, 8))
    
    x_exact = np.linspace(a, b, 500)
    fprime_exact = -6*np.sin(2*x_exact) + 3*np.cos(3*x_exact) - 2*x_exact*np.exp(-x_exact**2)
    ax.plot(x_exact, fprime_exact, 'k-', label='Exact', linewidth=2)

    for idx, N in enumerate(N_list):
        # Compute derivatives using both methods
        x_cd, fprime_cd = central_diff(a, b, N)
        x_5pt, fprime_5pt = five_point_diff(a, b, N)
        
        # Compute exact derivative
        
        
        # Plot
        ax.plot(x_cd, fprime_cd, 'r-', label=f'Central Diff N = {N}', markersize=4, alpha=0.6)
        ax.plot(x_5pt, fprime_5pt, 'b-', label=f'5-Point N = {N}', markersize=4, alpha=0.6)
        
    ax.set_xlabel('x')
    ax.set_ylabel("f'(x)")
    ax.set_title("Derivative approximations")
    ax.legend()
    ax.grid(True, alpha=0.3)
    
    plt.tight_layout()
    
    if not show_plot:
        plt.close(fig)
    
    return fig
```

```python
# Compare the methods with different grid sizes
fig = plot_derivatives(0, 2*np.pi, [10, 25, 50], show_plot=True)
plt.show()
```

```python
grader.check("q1c")
```

# **Question 2**: The Lorenz Equations and Chaotic Dynamics

The **Lorenz equations** are a system of three coupled ordinary differential equations that were originally derived by Edward Lorenz in 1963 as a simplified model of atmospheric convection. These equations are famous for exhibiting **chaotic behavior** - extreme sensitivity to initial conditions, popularly known as the "butterfly effect."

### The Lorenz System

The Lorenz equations are given by:

$$\frac{dx}{dt} = \sigma(y - x)$$

$$\frac{dy}{dt} = x(\rho - z) - y$$

$$\frac{dz}{dt} = xy - \beta z$$

where:
- $\sigma$ is the Prandtl number (relates viscosity to thermal conductivity)
- $\rho$ is the Rayleigh number (relates buoyancy to viscosity)
- $\beta$ is a geometric factor

**Standard Parameters:** For this problem, we'll use the classic chaotic regime parameters:
- $\sigma = 10$
- $\rho = 28$
- $\beta = 8/3$

With these parameters, the system exhibits the famous "Lorenz attractor" - a butterfly-shaped strange attractor that represents chaotic dynamics.

### Numerical Integration: Explicit Euler Method

To solve this system numerically, we'll use the **Explicit Euler method** (also called Forward Euler). This is the simplest time-stepping scheme for ordinary differential equations.

**The Method:**

Given a differential equation $\frac{d\mathbf{u}}{dt} = \mathbf{f}(t, \mathbf{u})$ and a time step $\Delta t$, the Explicit Euler method updates the solution from time $t$ to $t + \Delta t$ using:

$$\mathbf{u}(t + \Delta t) \approx \mathbf{u}(t) + \Delta t \cdot \mathbf{f}(t, \mathbf{u}(t))$$

For the Lorenz system with state vector $\mathbf{u} = (x, y, z)$:

$$x_{n+1} = x_n + \Delta t \cdot \sigma(y_n - x_n)$$

$$y_{n+1} = y_n + \Delta t \cdot [x_n(\rho - z_n) - y_n]$$

$$z_{n+1} = z_n + \Delta t \cdot [x_n y_n - \beta z_n]$$

**Note:** While Explicit Euler is simple, it requires small time steps for stability. More sophisticated methods (like Runge-Kutta) are generally preferred for chaotic systems, but Euler is excellent for learning the basics!

## **Part A**: Single Euler Step for the Lorenz System

First, we'll implement a single time step of the Explicit Euler method for the Lorenz equations.

### Your Task

Write a function `lorenz_euler_step(x, y, z, dt, sigma=10, rho=28, beta=8/3)` that:
1. Takes the current state $(x, y, z)$ at time $t$
2. Computes the derivatives using the Lorenz equations
3. Updates the state using Explicit Euler method
4. Returns the new state $(x', y', z')$ at time $t + \Delta t$

**Requirements:**
- Compute derivatives:
  - $\frac{dx}{dt} = \sigma(y - x)$
  - $\frac{dy}{dt} = x(\rho - z) - y$
  - $\frac{dz}{dt} = xy - \beta z$
- Update using Euler: 
  - $x' = x + \Delta t \cdot \frac{dx}{dt}$
  - $y' = y + \Delta t \cdot \frac{dy}{dt}$
  - $z' = z + \Delta t \cdot \frac{dz}{dt}$
- Return the updated state as a tuple: `(x_new, y_new, z_new)`

**Parameters:**
- `x, y, z`: float, current state variables
- `dt`: float, time step size
- `sigma, rho, beta`: float, Lorenz equation parameters (with defaults)

**Returns:**
- `(x_new, y_new, z_new)`: tuple of floats, updated state at $t + \Delta t$

```python
def lorenz_euler_step(x, y, z, dt, sigma=10, rho=28, beta=8/3):
    
    return x_new, y_new, z_new
```

```python
grader.check("q2a")
```

## **Part B**: Simulating the Lorenz Attractor

Now we'll use our Euler step function to simulate the full trajectory of the Lorenz system and visualize the famous "butterfly" attractor.

### Your Task

Write a function `simulate_lorenz(x0, y0, z0, dt, N, show_plot=False)` that:
1. Takes initial conditions $(x_0, y_0, z_0)$ and time step $\Delta t$
2. Simulates $N$ time steps using repeated calls to `lorenz_euler_step`
3. Stores the trajectory (all $x$, $y$, $z$ values over time)
4. Plots $z$ vs $x$ to visualize the Lorenz attractor (the "butterfly plot")
5. Returns the figure object

**Requirements:**
- Initialize arrays to store the trajectory:
  - `x_history = np.zeros(N+1)` (similarly for `y_history`, `z_history`)
  - Store initial conditions: `x_history[0] = x0`, etc.
- Loop through $N$ steps:
  - Call `lorenz_euler_step` to get next state
  - Store the new state in history arrays
- Create a plot of $z$ vs $x$ (this gives the classic butterfly view)
  - Use `plt.plot(x_history, z_history, linewidth=0.5)` for a clean trajectory
  - Label axes: "x" and "z"
  - Add title: "Lorenz Attractor"
  - Add grid for readability
- If `show_plot=False`, suppress display but return the figure
- Return the matplotlib figure object

**Parameters:**
- `x0, y0, z0`: float, initial conditions
- `dt`: float, time step size
- `N`: int, number of time steps to simulate
- `show_plot`: bool, optional (default: False)

**Returns:**
- `fig`: matplotlib.figure.Figure, the figure object containing the plot

**Recommended Parameters for the Butterfly:**
- Initial conditions: $(x_0, y_0, z_0) = (1, 1, 1)$
- Time step: $\Delta t = 0.01$
- Number of steps: $N = 10000$ (simulates up to $t = 100$ time units)

```python
def simulate_lorenz(x0, y0, z0, dt, N, show_plot=False):
    
    return fig
```

```python
grader.check("q2b")
```

### Example: The Famous Butterfly Attractor

Let's simulate the Lorenz system with the recommended parameters to see the iconic butterfly-shaped attractor:

```python
# Simulate the Lorenz system
# Initial conditions: (1, 1, 1)
# Time step: 0.01
# Number of steps: 10000 (simulates t = 0 to t = 100)
fig = simulate_lorenz(1, 1, 1, 0.01, 10000, show_plot=True)
plt.show()

print("The butterfly attractor!")
print("Notice how the trajectory spirals around two 'wings' of the butterfly.")
print("This chaotic behavior means tiny changes in initial conditions lead to vastly different trajectories.")
```

### Bonus Exploration: The Butterfly Effect

Try running the simulation with slightly different initial conditions to see how sensitive the system is to small changes. Even a tiny difference (like 1.0 vs 1.001) will lead to completely different trajectories after some time!

```python
# Compare two nearly identical initial conditions
fig, (ax1, ax2) = plt.subplots(1, 2, figsize=(16, 6))

# Simulate with initial condition (1, 1, 1)
fig1 = simulate_lorenz(1.0, 1.0, 1.0, 0.01, 10000, show_plot=False)
line1 = fig1.axes[0].lines[0]
x1, z1 = line1.get_data()

# Simulate with slightly different initial condition (1.001, 1, 1)
fig2 = simulate_lorenz(1.001, 1.0, 1.0, 0.01, 10000, show_plot=False)
line2 = fig2.axes[0].lines[0]
x2, z2 = line2.get_data()

# Plot both trajectories
ax1.plot(x1, z1, linewidth=0.5, color='blue', alpha=0.7, label='IC: (1.0, 1.0, 1.0)')
ax1.plot(x2, z2, linewidth=0.5, color='red', alpha=0.7, label='IC: (1.001, 1.0, 1.0)')
ax1.set_xlabel('x')
ax1.set_ylabel('z')
ax1.set_title('Lorenz Attractor: Two Trajectories')
ax1.legend()
ax1.grid(True, alpha=0.3)

# Plot the difference in trajectories over time
time = np.arange(10001) * 0.01
distance = np.sqrt((x1 - x2)**2 + (z1 - z2)**2)
ax2.plot(time, distance, linewidth=1, color='purple')
ax2.set_xlabel('Time')
ax2.set_ylabel('Distance between trajectories')
ax2.set_title('Divergence of Trajectories (Butterfly Effect)')
ax2.set_yscale('log')
ax2.grid(True, alpha=0.3)

plt.tight_layout()
plt.show()

print("Notice how the trajectories start nearly identical but diverge exponentially!")
print("This is the essence of chaos: sensitive dependence on initial conditions.")
```

```python

```

# **Question 3**: Numerical Integration with the Trapezoid Rule

Numerical integration (also called **quadrature**) is the process of approximating the definite integral of a function when we either don't have a closed-form antiderivative or only have discrete data points. This is essential in scientific computing where data often comes from experiments or simulations.

### The Problem

Given discrete data points $(x_0, y_0), (x_1, y_1), \ldots, (x_n, y_n)$, we want to approximate the integral:

$$\int_{x_0}^{x_n} y(x) \, dx$$

### The Trapezoid Rule

The **Trapezoid Rule** is one of the simplest and most intuitive methods for numerical integration. The key idea is to approximate the area under the curve $y(x)$ by dividing it into trapezoids rather than rectangles.

**Geometric Interpretation:**

Between any two consecutive points $(x_i, y_i)$ and $(x_{i+1}, y_{i+1})$, we approximate the curve with a straight line (linear interpolation). The area under this line segment is a trapezoid with:
- Parallel sides of height $y_i$ and $y_{i+1}$
- Width $h_i = x_{i+1} - x_i$

The area of one trapezoid is:

$$A_i = \frac{1}{2}(y_i + y_{i+1}) \cdot h_i$$

**The Formula:**

For $n$ intervals (meaning $n+1$ data points), the total integral approximation is:

$$\int_{x_0}^{x_n} y(x) \, dx \approx \sum_{i=0}^{n-1} \frac{1}{2}(y_i + y_{i+1})(x_{i+1} - x_i)$$


## Your Task

Write a function `trapezoid_integrate(x, y)` that:
1. Takes arrays of $x$ and $y$ values representing discrete data points
2. Implements the Trapezoid Rule **by hand** (no `scipy.integrate` or similar functions!)
3. Returns the approximate integral of $y(x)$ over the range $[x_0, x_n]$

**Requirements:**
- Input validation:
  - `x` and `y` must be numpy arrays of the same length
  - Need at least 2 points to integrate
- Handle **non-uniform spacing** - don't assume $x$ values are equally spaced
- Algorithm:
  - Loop through consecutive pairs of points
  - For each interval, compute $h_i = x_{i+1} - x_i$
  - Compute trapezoid area: $\frac{1}{2}(y_i + y_{i+1}) \cdot h_i$
  - Sum all trapezoid areas
- Return the total integral as a float

**Parameters:**
- `x`: numpy array, x-coordinates (must be in increasing order)
- `y`: numpy array, y-coordinates (function values at each x)

**Returns:**
- `integral`: float, approximate value of $\int y(x) \, dx$


```python
def trapezoid_integrate(x, y):
    
    return integral
```

```python
grader.check("q3")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> x_test, dfdx_test = central_diff(0, 1, 10)
>>> test_result = isinstance(x_test, np.ndarray) and isinstance(dfdx_test, np.ndarray)
>>> test_result
True
```
```
>>> x_test, dfdx_test = central_diff(0, np.pi, 100)
>>> df_analytical = -6 * np.sin(2 * x_test) + 3 * np.cos(3 * x_test) - 2 * x_test * np.exp(-x_test ** 2)
>>> max_error = np.max(np.abs(dfdx_test[1:-1] - df_analytical[1:-1]))
>>> test_result = max_error < 0.01
>>> bool(test_result)
True
```


Tests for q1b (points: None)
```
>>> x_vals, fprime_vals = five_point_diff(0, 2 * np.pi, 50)
>>> test_result = isinstance(x_vals, np.ndarray) and isinstance(fprime_vals, np.ndarray) and (len(x_vals) == len(fprime_vals))
>>> test_result
True
```
```
>>> x_test = np.pi / 2
>>> x_vals, fprime_vals = five_point_diff(0, np.pi, 500)
>>> idx = np.argmin(np.abs(x_vals - x_test))
>>> analytic_deriv = -6 * np.sin(2 * x_test) + 3 * np.cos(3 * x_test) - 2 * x_test * np.exp(-x_test ** 2)
>>> test_result = np.abs(fprime_vals[idx] - analytic_deriv) < 0.1
>>> bool(test_result)
True
```


Tests for q1c (points: None)
```
>>> fig = plot_derivatives(0, 2 * np.pi, [20, 50], show_plot=False)
>>> test_result = isinstance(fig, matplotlib.figure.Figure)
>>> test_result
True
```
```
>>> fig = plot_derivatives(0, np.pi, [10, 20], show_plot=False)
>>> test_result = len(fig.axes) == 1
>>> bool(test_result)
True
```
```
>>> fig = plot_derivatives(0, np.pi, [25], show_plot=False)
>>> ax = fig.axes[0]
>>> legend_labels = [text.get_text() for text in ax.get_legend().get_texts()]
>>> test_result = not '' in legend_labels and (not '' in legend_labels) and (not '' in legend_labels)
>>> bool(test_result)
True
```


Tests for q2a (points: None)
```
>>> result = lorenz_euler_step(1.0, 1.0, 1.0, 0.01)
>>> test_result = isinstance(result, tuple) and len(result) == 3
>>> test_result
True
```
```
>>> x_new, y_new, z_new = lorenz_euler_step(0, 0, 0, 0.01)
>>> test_result = np.abs(x_new) < 1e-10 and np.abs(y_new) < 1e-10 and (np.abs(z_new) < 1e-10)
>>> bool(test_result)
True
```


Tests for q2b (points: None)
```
>>> fig = simulate_lorenz(1, 1, 1, 0.01, 100, show_plot=False)
>>> test_result = isinstance(fig, matplotlib.figure.Figure)
>>> test_result
True
```
```
>>> fig = simulate_lorenz(1, 1, 1, 0.01, 500, show_plot=False)
>>> ax = fig.axes[0]
>>> test_result = len(ax.lines) == 1
>>> bool(test_result)
True
```
```
>>> fig = simulate_lorenz(1, 1, 1, 0.01, 1000, show_plot=False)
>>> ax = fig.axes[0]
>>> line = ax.lines[0]
>>> x_data, y_data = line.get_data()
>>> test_result = len(x_data) == 1001 and len(y_data) == 1001
>>> bool(test_result)
True
```


Tests for q3 (points: None)
```
>>> x = np.array([0.0, 1.0])
>>> y = np.array([1.0, 1.0])
>>> result = trapezoid_integrate(x, y)
>>> test_result = np.abs(result - 1.0) < 1e-10
>>> bool(test_result)
True
```
```
>>> x = np.linspace(0, 2, 50)
>>> y = x
>>> result = trapezoid_integrate(x, y)
>>> test_result = np.abs(result - 2.0) < 0.01
>>> bool(test_result)
True
```
```
>>> x = np.linspace(0, 1, 100)
>>> y = x ** 2
>>> result = trapezoid_integrate(x, y)
>>> exact = 1.0 / 3.0
>>> test_result = np.abs(result - exact) < 0.001
>>> bool(test_result)
True
```
```
>>> x = np.array([0.0, 0.1, 0.5, 1.0])
>>> y = np.array([0.0, 0.01, 0.25, 1.0])
>>> result = trapezoid_integrate(x, y)
>>> test_result = np.abs(result - 1.0 / 3.0) < 0.05
>>> bool(test_result)
True
```
