```python
# Initialize Otter
import otter
grader = otter.Notebook("ws1.ipynb")
```

# Workshop 1
due Tues, Sept 9 at 9pm

```python
# Setup - Import math for this worksheet
import math
```

## **Question 1**: Orbital Mechanics and Built-in Math Operations

**Given Constants:**
- Earth's orbital radius: r = 1.496 × 10¹¹ m (1 AU)
- Gravitational constant: G = 6.674 × 10⁻¹¹ m³/(kg⋅s²)
- Solar mass: M = 1.989 × 10³⁰ kg
- Earth's mass: m = 5.972 × 10²⁴ kg

### Part (a): Orbital Circumference

Calculate and return the circumference of Earth's orbit assuming zero eccentricity. Use built in python functions for any necessary constants

```python

def orbital_circumference(radius):
    # BEGIN SOLUTION NO PROMPT
    return 2 * math.pi * radius
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT

```

```python
grader.check("q1a")
```

### Part (b): Gravitational Force

Calculate and return the gravitational force between the Sun and Earth using Newton's law: F = GMm/r², ensure this function always returns a positive force for possible negative input values.

```python
def gravitational_force(G, M, m, r):
    # BEGIN SOLUTION NO PROMPT
    force = G * M * m / (r**2)
    return abs(force)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT

```

```python
grader.check("q1b")
```

### Part (c): Orbital Velocity

This function should **use the results from part (b)** to calculate and return the centripetal force, then derive the velocity. (You need to relate tangential velocity to centripetal force)

```python
def orbital_velocity(G, M, m, r):
    # BEGIN SOLUTION NO PROMPT
    F = gravitational_force(G, M, m, r)
    velocity = math.sqrt((F * r) / m)
    return velocity
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q1c")
```

## **Question 2**: Physics Data with Dictionaries

**Note**: Your solutions should use dictionary operations only - no loops or if/else statements are required.

### Part (a): Simple Dictionary

**Example dictionary structure:**
```python
experiment = {
    "temperature": 295.0,  # Kelvin
    "pressure": 101325.0,  # Pascal
    "voltage": 12.0,       # Volts
    "current": 2.5,        # Amperes
    "time": 3600.0         # seconds
}
```

Complete the function `calculate_power_energy(experiment_dict)` that returns a tuple containing:
1. Power (P = V × I) in Watts
2. Energy (E = P × t) in Joules

```python
def calculate_power_energy(experiment_dict):
    # BEGIN SOLUTION NO PROMPT
    voltage = experiment_dict["voltage"]
    current = experiment_dict["current"]
    time = experiment_dict["time"]
    
    power = voltage * current
    energy = power * time
    
    return (power, energy)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    # Return the tuple (power, energy)
    """; # END PROMPT
```

```python
grader.check("q2a")
```

### Part (b): Nested Dictionary - Multi-Sensor Physics Data

You're given a nested dictionary containing data from multiple physics sensors.

**Example nested dictionary structure:**
```python
sensor_data = {
    "location": "Lab_Room_101",
    "sensors": {
        "thermometer": {"value": 298.15, "unit": "K", "accuracy": 0.1},
        "barometer": {"value": 101325, "unit": "Pa", "accuracy": 10},
        "accelerometer": {"value": 9.81, "unit": "m/s^2", "accuracy": 0.01}
    },
    "timestamp": "2025-09-05T14:30:00"
}
```

Complete the function `analyze_sensor_data(sensor_data)` that returns a tuple containing:
1. The temperature value in Celsius (0 degrees C = 273.15 K)
2. The total number of sensors

```python
def analyze_sensor_data(sensor_data):
    # BEGIN SOLUTION NO PROMPT
    # Get temperature in Kelvin and convert to Celsius
    temp_kelvin = sensor_data["sensors"]["thermometer"]["value"]
    temp_celsius = temp_kelvin - 273.15
    
    # Count number of sensors
    num_sensors = len(sensor_data["sensors"])
    
    return (temp_celsius, num_sensors)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    # Return the tuple (temp_celsius, num_sensors)
    """; # END PROMPT
```

```python
grader.check("q2b")
```

## **Question 3**: Safety Conditions with Boolean Logic

### Laboratory Safety Check

A laser physics experiment requires multiple safety conditions to be met before it can start. Write a function that checks all safety conditions and returns `True` only if **ALL** conditions are safe.

**Safety Requirements:**
1. Temperature must be between 18°C and 25°C (inclusive)
2. Humidity must be less than 60%
3. Laser power must not exceed 50 Watts
4. Emergency stop button must be functional (True)
5. Protective equipment must be present (True)

Complete the function `is_experiment_safe(temperature, humidity, laser_power, emergency_stop, protective_equipment)` that returns `True` if all safety conditions are met, `False` otherwise. Note that `temperature`, `humidity`, and `laser_power` are numbers while `emergency_stop` and `protective_equiptment` are themselves boolean values.

```python
def is_experiment_safe(temperature, humidity, laser_power, emergency_stop, protective_equipment):
    # BEGIN SOLUTION NO PROMPT
    # Check all safety conditions
    temp_ok = 18 <= temperature <= 25
    humidity_ok = humidity < 60
    power_ok = laser_power <= 50
    
    # All conditions must be True
    return temp_ok and humidity_ok and power_ok and emergency_stop and protective_equipment
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q3")
```

## **Question 4**: Special Characters and F-String Formatting

Create a function `format_lab_header(experiment_name, date, researcher)` that takes those three strings and returns a properly formatted header string for a physics lab report. The header should include:

1. A title line: PHYSICS LABORATORY REPORT
2. A line of equals signs (=) underneath the title (25 of them)
3. The experiment name with quotation marks around it
4. The researcher's name with a tab character before it
5. The date on a new line

**Example output string:**
```
PHYSICS LABORATORY REPORT
=========================
Experiment: "Photoelectric Effect"
	Researcher: Dr. Smith
Date: 2025-09-05
```

```python
def format_lab_header(experiment_name, date, researcher):
    # BEGIN SOLUTION NO PROMPT
    header = "PHYSICS LABORATORY REPORT\n"
    header += "=========================\n"
    header += f"Experiment: \"{experiment_name}\"\n"
    header += f"\tResearcher: {researcher}\n"
    header += f"Date: {date}"
    return header
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    # Use special characters
    """; # END PROMPT
```

```python
grader.check("q4")
```

## **Question 5**: Collections in Physics - Lists, Tuples, and Sets

### Particle Detector Analysis

You're analyzing data from a particle physics experiment. The detector records different types of particles, and you need to process this information using various Python collection types.

**Given data structures:**
- `detected_particles` (list): A list of particle names detected in chronological order
- `reference_particles` (tuple): A fixed reference list of known particle types
- `confirmed_detections` (set): A set of particles that have been confirmed by multiple detectors

Complete the function `analyze_particle_data(detected_particles, reference_particles, confirmed_detections)` that returns a tuple containing:

1. **Unique particles** (set): Set containing just detected particles
2. **Missing particles** (set): Set of reference particles that weren't detected
3. **Confirmed and detected** (set): Find particles that are both in detected_particles AND confirmed_detections
4. **Reference check** (bool): Check if the string "muon" is in the reference_particles tuple
5. **Detection count** (int): Count how many times "electron" appears in the detected_particles list

**Example:**
```python
detected = ["electron", "muon", "electron", "photon", "neutrino"]
reference = ("electron", "muon", "photon", "proton", "neutron")
confirmed = {"electron", "photon", "pion"}

# Expected result:
# unique: {"electron", "muon", "photon", "neutrino"}
# missing: {"proton", "neutron"}  
# confirmed_and_detected: {"electron", "photon"}
# has_muon: True
# electron_count: 2
```

```python
def analyze_particle_data(detected_particles, reference_particles, confirmed_detections):
    # BEGIN SOLUTION NO PROMPT
    # 1. Convert detected particles list to set to find unique particles
    unique_particles = set(detected_particles)
    
    # 2. Find missing particles (in reference but not detected)
    reference_set = set(reference_particles)
    detected_set = set(detected_particles)
    missing_particles = reference_set - detected_set
    
    # 3. Find particles that are both detected and confirmed (intersection)
    confirmed_and_detected = detected_set & confirmed_detections
    
    # 4. Check if "muon" is in reference particles tuple
    has_muon = "muon" in reference_particles
    
    # 5. Count occurrences of "electron" in detected particles list
    electron_count = detected_particles.count("electron")
    
    return (unique_particles, missing_particles, confirmed_and_detected, has_muon, electron_count)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    # Return the tuple (unique_particles, missing_particles, confirmed_and_detected, has_muon, electron_count)
    """; # END PROMPT
```

```python
grader.check("q5")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the raw cell below.

```python
# BEGIN SOLUTION NO PROMPT
# END SOLUTION
""" # BEGIN PROMPT
"""
# write ai disclosure here:

"""
"""; # END PROMPT
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> earth_radius = 149600000000.0
>>> result = orbital_circumference(earth_radius)
>>> expected = 2 * math.pi * earth_radius
>>> abs(result - expected) < 100000.0
True
```
```
>>> result = orbital_circumference(1.0)
>>> expected = 2 * math.pi
>>> abs(result - expected) < 1e-10
True
```


Tests for q1b (points: None)
```
>>> G = 6.674e-11
>>> M_sun = 1.989e+30
>>> m_earth = 5.972e+24
>>> r = 149600000000.0
>>> result = gravitational_force(G, M_sun, m_earth, r)
>>> expected = G * M_sun * m_earth / r ** 2
>>> abs(result - expected) < expected * 1e-10
True
```
```
>>> result1 = gravitational_force(1.0, 10.0, 5.0, 2.0)
>>> result2 = gravitational_force(1.0, -10.0, 5.0, 2.0)
>>> result3 = gravitational_force(1.0, 10.0, -5.0, 2.0)
>>> result1 > 0 and result2 > 0 and (result3 > 0) and (result1 == result2 == result3)
True
```


Tests for q1c (points: None)
```
>>> G = 6.674e-11
>>> M_sun = 1.989e+30
>>> m_earth = 5.972e+24
>>> r = 149600000000.0
>>> result = orbital_velocity(G, M_sun, m_earth, r)
>>> expected = math.sqrt(abs(G * M_sun * m_earth / r ** 2) * r / m_earth)
>>> abs(result - expected) < 1e-10
True
```
```
>>> G, M, m, r = (1.0, 4.0, 2.0, 1.0)
>>> result = orbital_velocity(G, M, m, r)
>>> expected = 2.0
>>> abs(result - expected) < 1e-10
True
```


Tests for q2a (points: None)
```
>>> experiment = {'temperature': 295.0, 'pressure': 101325.0, 'voltage': 12.0, 'current': 2.5, 'time': 3600.0}
>>> power, energy = calculate_power_energy(experiment)
>>> expected_power = 12.0 * 2.5
>>> expected_energy = 30.0 * 3600.0
>>> abs(power - expected_power) < 1e-10 and abs(energy - expected_energy) < 1e-10
True
```
```
>>> experiment2 = {'temperature': 300.0, 'pressure': 98000.0, 'voltage': 5.0, 'current': 1.2, 'time': 1800.0}
>>> power, energy = calculate_power_energy(experiment2)
>>> expected_power = 5.0 * 1.2
>>> expected_energy = 6.0 * 1800.0
>>> abs(power - expected_power) < 1e-10 and abs(energy - expected_energy) < 1e-10
True
```


Tests for q2b (points: None)
```
>>> sensor_data = {'location': 'Lab_Room_101', 'sensors': {'thermometer': {'value': 298.15, 'unit': 'K', 'accuracy': 0.1}, 'barometer': {'value': 101325, 'unit': 'Pa', 'accuracy': 10}, 'accelerometer': {'value': 9.81, 'unit': 'm/s^2', 'accuracy': 0.01}}, 'timestamp': '2025-09-05T14:30:00'}
>>> temp_c, num_sensors = analyze_sensor_data(sensor_data)
>>> expected_temp = 298.15 - 273.15
>>> expected_num = 3
>>> abs(temp_c - expected_temp) < 1e-10 and num_sensors == expected_num
True
```
```
>>> sensor_data2 = {'location': 'Lab_Room_202', 'sensors': {'thermometer': {'value': 273.15, 'unit': 'K', 'accuracy': 0.05}, 'hygrometer': {'value': 45.2, 'unit': '%', 'accuracy': 1.0}}, 'timestamp': '2025-09-05T15:00:00'}
>>> temp_c, num_sensors = analyze_sensor_data(sensor_data2)
>>> expected_temp = 273.15 - 273.15
>>> expected_num = 2
>>> abs(temp_c - expected_temp) < 1e-10 and num_sensors == expected_num
True
```


Tests for q3 (points: None)
```
>>> result = is_experiment_safe(22.0, 45.0, 30.0, True, True)
>>> result == True
True
```
```
>>> result = is_experiment_safe(26.0, 45.0, 30.0, True, True)
>>> result == False
True
```
```
>>> result = is_experiment_safe(22.0, 65.0, 30.0, True, True)
>>> result == False
True
```
```
>>> result = is_experiment_safe(22.0, 45.0, 55.0, True, True)
>>> result == False
True
```
```
>>> result1 = is_experiment_safe(22.0, 45.0, 30.0, False, True)
>>> result2 = is_experiment_safe(22.0, 45.0, 30.0, True, False)
>>> result3 = is_experiment_safe(18.0, 59.9, 50.0, True, True)
>>> result1 == False and result2 == False and (result3 == True)
True
```


Tests for q4 (points: None)
```
>>> result = format_lab_header('Photoelectric Effect', '2025-09-05', 'Dr. Smith')
>>> expected_lines = ['PHYSICS LABORATORY REPORT', '=========================', 'Experiment: "Photoelectric Effect"', '\tResearcher: Dr. Smith', 'Date: 2025-09-05']
>>> expected = '\n'.join(expected_lines)
>>> result == expected
True
```
```
>>> result = format_lab_header('Harmonic Oscillator', '2025-09-06', 'Prof. Johnson')
>>> '"Harmonic Oscillator"' in result and '\t' in result and (result.count('\n') == 4)
True
```


Tests for q5 (points: None)
```
>>> detected = ['electron', 'muon', 'electron', 'photon', 'neutrino']
>>> reference = ('electron', 'muon', 'photon', 'proton', 'neutron')
>>> confirmed = {'electron', 'photon', 'pion'}
>>> unique, missing, confirmed_and_detected, has_muon, electron_count = analyze_particle_data(detected, reference, confirmed)
>>> unique == {'electron', 'muon', 'photon', 'neutrino'} and missing == {'proton', 'neutron'} and (confirmed_and_detected == {'electron', 'photon'}) and (has_muon == True) and (electron_count == 2)
True
```
```
>>> detected = ['tau', 'higgs']
>>> reference = ('electron', 'muon', 'photon')
>>> confirmed = {'neutrino', 'quark'}
>>> unique, missing, confirmed_and_detected, has_muon, electron_count = analyze_particle_data(detected, reference, confirmed)
>>> unique == {'tau', 'higgs'} and missing == {'electron', 'muon', 'photon'} and (confirmed_and_detected == set()) and (has_muon == True) and (electron_count == 0)
True
```
```
>>> detected = ['photon', 'photon', 'electron', 'photon']
>>> reference = ('photon', 'electron', 'positron')
>>> confirmed = {'photon', 'electron', 'muon'}
>>> unique, missing, confirmed_and_detected, has_muon, electron_count = analyze_particle_data(detected, reference, confirmed)
>>> isinstance(unique, set) and isinstance(missing, set) and isinstance(confirmed_and_detected, set) and ('positron' in missing) and (len(unique) == 2) and (electron_count == 1)
True
```
