```python
# Initialize Otter
import otter
grader = otter.Notebook("hw9.ipynb")
```

```python
rng_seed = 60
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import scipy as sp
import pandas as pd
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

## **Question 1**: Power Method for Finding Maximum Eigenvalue

The **power method** is an iterative algorithm for finding the largest eigenvalue (in absolute value) of a matrix. It's particularly useful for large sparse matrices where computing all eigenvalues would be computationally expensive.

### Background: Eigenvalues and Eigenvectors

For a square matrix $\mathbf{A}$, an **eigenvector** $\mathbf{v}$ and its corresponding **eigenvalue** $\lambda$ satisfy:

$$\mathbf{A}\mathbf{v} = \lambda \mathbf{v}$$

Every square matrix has eigenvalues $\lambda_1, \lambda_2, ..., \lambda_n$ (some may be complex or repeated). We order them by absolute value: $|\lambda_1| \geq |\lambda_2| \geq ... \geq |\lambda_n|$.

The **dominant eigenvalue** is $\lambda_1$ (the one with largest absolute value).

### The Power Method Algorithm

The power method is based on a simple observation: if we repeatedly multiply a vector by matrix $\mathbf{A}$, the result will increasingly point in the direction of the dominant eigenvector.

**Algorithm:**

1. **Initialize**: Start with a random unit vector $\mathbf{x}^{(0)}$ (where $||\mathbf{x}^{(0)}|| = 1$)

2. **Iterate** for $k = 1, 2, ..., N$:
   - Multiply by $\mathbf{A}$: $\mathbf{y}^{(k)} = \mathbf{A} \mathbf{x}^{(k-1)}$
   - Normalize: $\mathbf{x}^{(k)} = \frac{\mathbf{y}^{(k)}}{||\mathbf{y}^{(k)}||}$

3. **Extract eigenvalue**: After $N$ iterations, $\mathbf{x}^{(N)}$ approximates the dominant eigenvector. The corresponding eigenvalue can be computed using the **Rayleigh quotient**:

$$\lambda_1 \approx \frac{\mathbf{x}^T \mathbf{A} \mathbf{x}}{\mathbf{x}^T \mathbf{x}} = \mathbf{x}^T \mathbf{A} \mathbf{x}$$

(The denominator equals 1 since $\mathbf{x}$ is a unit vector)

### Why Does This Work?

Any vector $\mathbf{x}^{(0)}$ can be written as a linear combination of eigenvectors:
$$\mathbf{x}^{(0)} = c_1 \mathbf{v}_1 + c_2 \mathbf{v}_2 + ... + c_n \mathbf{v}_n$$

After $k$ multiplications by $\mathbf{A}$:
$$\mathbf{A}^k \mathbf{x}^{(0)} = c_1 \lambda_1^k \mathbf{v}_1 + c_2 \lambda_2^k \mathbf{v}_2 + ... + c_n \lambda_n^k \mathbf{v}_n$$

Since $|\lambda_1| > |\lambda_i|$ for $i > 1$, the term $\lambda_1^k$ grows fastest, and the vector increasingly points toward $\mathbf{v}_1$.

### Your Task

Write a function `power_method(A, N)` that implements the power method to find the dominant eigenvalue of a square matrix.

**Requirements:**
- Input `A` is a square numpy array (any size $n \times n$)
- Input `N` is the number of iterations to perform
- Initialize a random unit vector using `np.random.rand(n)` then normalize it
- Perform $N$ iterations of:
  - Multiply: $\mathbf{y} = \mathbf{A} \mathbf{x}$
  - Normalize: $\mathbf{x} = \mathbf{y} / ||\mathbf{y}||$
- After $N$ iterations, compute the eigenvalue using: $\lambda = \mathbf{x}^T \mathbf{A} \mathbf{x}$
- Return the estimated dominant eigenvalue

**Parameters:**
- `A`: numpy array, square matrix of shape (n, n)
- `N`: int, number of iterations to perform

**Returns:**
- `eigenvalue`: float, the estimated dominant eigenvalue

**Hints:**
- Use `np.random.rand(n)` to create a random vector
- Use `np.linalg.norm(x)` to compute the norm $||x||$
- Use `np.dot(x, y)` or `x @ y` for matrix-vector multiplication
- The Rayleigh quotient is: `x.T @ A @ x` or `np.dot(x, np.dot(A, x))`
- For a unit vector $\mathbf{x}$, we have $\mathbf{x}^T \mathbf{A} \mathbf{x} = \mathbf{x} \cdot (\mathbf{A} \mathbf{x})$

```python
def power_method(A, N):
    n = A.shape[0]
    
    # Initialize random unit vector
    x = np.random.rand(n)
    x = ...  # normalize
    
    # Power iteration
    for i in range(N):
        # Multiply by A
        y = ...
        
        # Normalize
        x = ...
    
    # Compute eigenvalue using Rayleigh quotient
    eigenvalue = ...
    
    return eigenvalue
```

```python
grader.check("q1")
```

## **Question 2**: Nearest Neighbor Algorithm in 2D

The **nearest neighbor algorithm** is a fundamental method in machine learning and data analysis. Given a query point and a dataset, it finds the data point that is closest to the query point. This has applications in classification, clustering, anomaly detection, and more.

### Background: Distance Metrics

To find the nearest neighbor, we need to measure the distance between points. In 2D, the most common metric is the **Euclidean distance**:

For two points $\mathbf{p} = (p_x, p_y)$ and $\mathbf{q} = (q_x, q_y)$, the Euclidean distance is:

$$d(\mathbf{p}, \mathbf{q}) = \sqrt{(p_x - q_x)^2 + (p_y - q_y)^2}$$

### Part a: Finding the Nearest Neighbor

Write a function `find_nearest_neighbor(data, query_point)` that finds the nearest neighbor to a query point in a 2D dataset.

**Algorithm:**
1. For each point in the dataset, compute the Euclidean distance to the query point
2. Find the index of the point with the minimum distance
3. Return that index

**Requirements:**
- Input `data` is a numpy array of shape $(N, 2)$ where each row is a 2D point
- Input `query_point` is a numpy array of shape $(2,)$ representing the query coordinates
- Compute Euclidean distance: $d_i = \sqrt{(x_i - q_x)^2 + (y_i - q_y)^2}$
- Return the integer index $i$ of the nearest neighbor (the point with minimum distance)

**Parameters:**
- `data`: numpy array of shape (N, 2), dataset of 2D points
- `query_point`: numpy array of shape (2,), the query point coordinates

**Returns:**
- `index`: int, the index of the nearest neighbor in the data array

**Hints:**
- Use `np.linalg.norm()` to compute Euclidean distance, or compute it manually
- You can compute distances to all points at once using array operations
- Use `np.argmin()` to find the index of the minimum distance
- For vectorized distance computation: `distances = np.linalg.norm(data - query_point, axis=1)`

```python
def find_nearest_neighbor(data, query_point):
    # Compute distances from query_point to all data points
    distances = ...
    
    # Find the index of minimum distance
    index = ...
    
    return index
```

```python
grader.check("q2a")
```

### Part b: Visualizing Nearest Neighbor Regions

Now we'll visualize the nearest neighbor algorithm by creating a **Voronoi diagram** - a partition of the plane into regions, where each region contains all points closest to a particular data point.

**Concept:**
- Create a dense grid of query points covering the data region
- For each grid point, find its nearest neighbor from the dataset
- Color each grid point according to which data point is its nearest neighbor
- This creates colored regions showing the "zones of influence" for each data point

Write a function `plot_nearest_neighbor_regions(data, show_plot=False)` that visualizes these regions.

**Algorithm:**
1. Determine the bounds of the data (min and max x and y coordinates)
2. Create a dense grid of points covering this region (e.g., 100×100 grid)
3. For each grid point, use `find_nearest_neighbor()` to determine which data point is closest
4. Create a plot where grid points are colored by their nearest neighbor
5. Overlay the original data points on top

**Requirements:**
- Input `data` is a numpy array of shape $(N, 2)$
- Determine data bounds with some padding (e.g., extend 10% beyond min/max)
- Create a meshgrid with at least 300 points in each direction
- Use `find_nearest_neighbor()` from part (a) for each grid point
- Plot the colored regions using `plt.scatter()` or `plt.contourf()`
- Overlay the original data points with different markers
- Include appropriate labels, title, and legend
- Return the matplotlib figure object

**Parameters:**
- `data`: numpy array of shape (N, 2), dataset of 2D points
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object

**Plotting Guidelines:**
- Use a colormap (e.g., 'tab10', 'Set3', or 'rainbow') to distinguish regions
- Plot grid points with small alpha (e.g., 0.3) to show regions as colored areas
- Plot original data points with larger markers and black edges
- Set appropriate axis labels: "x" and "y"
- Title: "Nearest Neighbor Regions (Voronoi Diagram)"
- Make sure the plot works for different data sizes (e.g., 3, 5, 10, or more points)

**Hints:**
- Use `np.meshgrid()` to create the grid
- Flatten the grid to get all query points: `grid_points = np.c_[xx.ravel(), yy.ravel()]`
- Use a loop or list comprehension to find nearest neighbors for all grid points
- Reshape the results back to grid shape for plotting
- For coloring regions, `plt.scatter()` with small `s` and `alpha` works well

Note: This may not run instantly but should run within ~10 seconds.

```python
def plot_nearest_neighbor_regions(data, show_plot=False):
    # Determine data bounds with padding
    x_min, y_min = data.min(axis=0)
    x_max, y_max = data.max(axis=0)
    
    # Add padding
    ...
    
    # Create grid
    n_grid = 100
    xx, yy = np.meshgrid(...)
    
    # Flatten grid
    grid_points = ...
    
    # Find nearest neighbor for each grid point
    nearest_indices = ...
    
    # Create plot
    fig, ax = plt.subplots(figsize=(10, 8))
    
    # Plot colored regions
    ...
    
    # Plot original data points
    ...
    
    # Add labels and formatting
    ...
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Create a random dataset with 10 points for plotting
example_data = np.random.uniform(low = 0, high = 1, size = (10,2))

# Plot the nearest neighbor regions
fig = plot_nearest_neighbor_regions(example_data, show_plot=True)
```

```python
grader.check("q2b")
```

## **Question 3**: Taylor Series Approximation

A **Taylor series** is a powerful mathematical tool that approximates an analytic function using a polynomial. By including more terms, we can achieve increasingly accurate approximations near a chosen center point.

### Background: Taylor Series

For an analytic function $f(x)$, the Taylor series centered at $x_0$ is:

$$f(x) \approx \sum_{n=0}^{N} \frac{f^{(n)}(x_0)}{n!}(x - x_0)^n$$

where $f^{(n)}(x_0)$ denotes the $n$-th derivative of $f$ evaluated at $x_0$.

Expanding this sum:
$$f(x) \approx f(x_0) + f'(x_0)(x-x_0) + \frac{f''(x_0)}{2!}(x-x_0)^2 + \frac{f'''(x_0)}{3!}(x-x_0)^3 + ... + \frac{f^{(N)}(x_0)}{N!}(x-x_0)^N$$

The Taylor series is most accurate near $x_0$ and becomes less accurate as we move away from the center point.

### The Function: $f(x) = e^x + \sin(x)$

For this question, we'll work with the function:
$$f(x) = e^x + \sin(x)$$

This function is analytic everywhere, making it perfect for Taylor series approximation. The derivatives follow patterns:
- For $e^x$: all derivatives are $e^x$
- For $\sin(x)$: derivatives cycle through $\sin(x), \cos(x), -\sin(x), -\cos(x), ...$

### Part a: Computing the Taylor Series

Write a function `taylor_approximation(N, x, x_0)` that computes the Taylor series approximation of $f(x) = e^x + \sin(x)$ up to order $N$, centered at $x_0$, and evaluates it at point $x$.

**Algorithm:**
1. Compute the derivatives of $f(x) = e^x + \sin(x)$ up to order $N$ at point $x_0$
2. For each term $n = 0, 1, 2, ..., N$:
   - Compute $\frac{f^{(n)}(x_0)}{n!}(x - x_0)^n$
3. Sum all terms to get the approximation

**Requirements:**
- Input `N` is the highest order term to include (integer $\geq 0$)
- Input `x` is the point at which to evaluate the approximation (float)
- Input `x_0` is the center point of the Taylor series (float)
- Compute derivatives analytically:
  - $\frac{d^n}{dx^n}(e^x) = e^x$ for all $n$
  - $\frac{d^n}{dx^n}(\sin(x))$ cycles: $\sin(x), \cos(x), -\sin(x), -\cos(x), ...$
- Return the sum of all terms from $n=0$ to $n=N$

**Parameters:**
- `N`: int, highest order term (e.g., N=3 means include terms up to $(x-x_0)^3$)
- `x`: float, point at which to evaluate the approximation
- `x_0`: float, center point of the Taylor series

**Returns:**
- `approximation`: float, the Taylor series approximation $f(x) \approx \sum_{n=0}^{N} \frac{f^{(n)}(x_0)}{n!}(x - x_0)^n$

**Hints:**
- Use `np.exp()` for exponential and `np.sin()`, `np.cos()` for trigonometric functions
- Use `np.math.factorial(n)` or `scipy.special.factorial(n)` for $n!$
- For sin derivatives: use `n % 4` to determine which function (0: sin, 1: cos, 2: -sin, 3: -cos)
- Remember that $0! = 1$ and $(x-x_0)^0 = 1$
- Build the sum term by term in a loop

```python
def taylor_approximation(N, x, x_0):
    
    return approximation
```

```python
grader.check("q3a")
```

### Part b: Visualizing Taylor Series Convergence

Now we'll visualize how Taylor series approximations improve as we include more terms. This will help us understand how the approximation becomes more accurate near the center point.

**Concept:**
- Plot the true function $f(x) = e^x + \sin(x)$ over a range centered at $x_0$
- Overlay Taylor series approximations of increasing order: $0, 1, 2, ..., N$
- Each approximation should be plotted as a different line
- This shows how adding more terms makes the approximation converge to the true function

Write a function `plot_taylor_series(N, x_0, show_plot=False)` that creates this visualization.

**Algorithm:**
1. Define a range of x values centered at $x_0$ (e.g., from $x_0 - 5$ to $x_0 + 5$)
2. Compute the true function values: $f(x) = e^x + \sin(x)$ for all x in the range
3. For each order $n = 0, 1, 2, ..., N$:
   - Compute the Taylor approximation at each x value using `taylor_approximation(n, x, x_0)`
   - Plot these values
4. Create a legend showing which line corresponds to which order
5. Return the figure object

**Requirements:**
- Input `N` is the maximum order to plot (integer $\geq 0$)
- Input `x_0` is the center point of the Taylor series (float)
- Create a range of at least 200 x values from $x_0 - 5$ to $x_0 + 5$
- Plot the true function as a thick black line with label "True function"
- Plot each Taylor approximation (orders 0 through N) with different colors
- Label each approximation as "Order 0", "Order 1", ..., "Order N"
- Include axis labels: "x" and "f(x)"
- Include title: "Taylor Series Approximation of $e^x + \sin(x)$ centered at $x_0 = $ {x_0:.1f}"
- Include a legend
- Include a vertical dashed line at $x_0$ to show the center point
- Return the matplotlib figure object

**Parameters:**
- `N`: int, maximum order of Taylor series to plot
- `x_0`: float, center point of the Taylor series
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object containing the plot

**Hints:**
- Use `np.linspace(...)` to create the x range
- Use a loop to plot each order's approximation
- You can use a list comprehension to compute approximations: `[taylor_approximation(n, xi, x_0) for xi in x_range]`
- Use `plt.axvline(x_0, linestyle='--', color='gray', alpha=0.5)` to add the vertical line at center
- Use different line styles or colors for different orders (matplotlib will auto-cycle colors)

```python
def plot_taylor_series(N, x_0, show_plot=False):

    ...

    # Create the plot
    fig, ax = plt.subplots(figsize=(10, 6))
    
    ...

    if show_plot:
        plt.show()
    
    return fig
```

```python
# Plot Taylor series centered at x_0 = 0
fig = plot_taylor_series(N=5, x_0=0, show_plot=True)
```

```python
grader.check("q3b")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: None)
```
>>> np.random.seed(rng_seed)
>>> A_test = np.array([[2, 1], [1, 2]])
>>> result = power_method(A_test, 20)
>>> test_result = isinstance(result, (int, float, np.floating))
>>> test_result
True
```
```
>>> np.random.seed(rng_seed)
>>> A = np.array([[3, 1], [1, 3]])
>>> estimated = power_method(A, 50)
>>> true_max = np.max(np.abs(np.linalg.eigvals(A)))
>>> test_result = abs(estimated - true_max) < 0.01
>>> bool(test_result)
True
```
```
>>> np.random.seed(rng_seed)
>>> A = np.array([[4, 1, 0], [1, 4, 1], [0, 1, 4]])
>>> estimated = power_method(A, 50)
>>> true_max = np.max(np.abs(np.linalg.eigvals(A)))
>>> test_result = abs(estimated - true_max) < 0.01
>>> bool(test_result)
True
```
```
>>> np.random.seed(rng_seed)
>>> A = np.diag([5, 3, 1])
>>> estimated = power_method(A, 30)
>>> true_max = 5.0
>>> test_result = abs(estimated - true_max) < 0.01
>>> bool(test_result)
True
```


Tests for q2a (points: None)
```
>>> data_test = np.array([[0, 0], [1, 1], [2, 2]])
>>> query_test = np.array([0.5, 0.5])
>>> result = find_nearest_neighbor(data_test, query_test)
>>> test_result = isinstance(result, (int, np.integer))
>>> test_result
True
```
```
>>> data = np.array([[0, 0], [1, 1], [2, 2]])
>>> query = np.array([0.1, 0.1])
>>> result = find_nearest_neighbor(data, query)
>>> test_result = result == 0
>>> bool(test_result)
True
```


Tests for q2b (points: None)
```
>>> test_data_q2b = np.array([[0, 0], [1, 1], [2, 0]])
>>> fig_test = plot_nearest_neighbor_regions(test_data_q2b, show_plot=False)
>>> is_figure = isinstance(fig_test, plt.Figure)
>>> plt.close(fig_test)
>>> is_figure
True
```
```
>>> test_data_small = np.array([[0, 0], [3, 3], [0, 3]])
>>> fig_small = plot_nearest_neighbor_regions(test_data_small, show_plot=False)
>>> has_figure = fig_small is not None
>>> has_axes = len(fig_small.axes) > 0
>>> plt.close(fig_small)
>>> (has_figure, has_axes)
(True, True)
```


Tests for q3a (points: None)
```
>>> result = taylor_approximation(3, 1.0, 0.0)
>>> isinstance(result, (int, float, np.floating))
True
```
```
>>> x_0_test = 0.0
>>> expected = np.exp(x_0_test) + np.sin(x_0_test)
>>> result = taylor_approximation(0, x_0_test, x_0_test)
>>> bool(abs(result - expected) < 1e-10)
True
```
```
>>> x_test = 0.5
>>> x_0_test = 0.0
>>> true_value = np.exp(x_test) + np.sin(x_test)
>>> approx = taylor_approximation(10, x_test, x_0_test)
>>> bool(abs(approx - true_value) < 1e-06)
True
```


Tests for q3b (points: None)
```
>>> fig_test = plot_taylor_series(3, 0.0, show_plot=False)
>>> is_figure = isinstance(fig_test, plt.Figure)
>>> plt.close(fig_test)
>>> is_figure
True
```
```
>>> N_test = 2
>>> fig_test = plot_taylor_series(N_test, 0.0, show_plot=False)
>>> ax = fig_test.axes[0]
>>> num_lines = len(ax.get_lines())
>>> has_correct_lines = num_lines >= N_test + 2
>>> plt.close(fig_test)
>>> has_correct_lines
True
```
