```python
# Initialize Otter
import otter
grader = otter.Notebook("ws9.ipynb")
```

```python
rng_seed = 70
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import scipy as sp
import pandas as pd
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

# **Question 1**: Root Finding Methods for f(x) = x³ - 3

In this question, you'll implement and compare three different methods for finding the root (zero) of the function:

$$f(x) = x^3 - 3$$

The exact root is $x = \sqrt[3]{3} \approx 1.4422$, but we'll use numerical methods to approximate it.

## **Part A**: Newton-Raphson Method

The **Newton-Raphson method** (also called Newton's method) is an iterative root-finding algorithm that uses calculus to converge quickly to a root.

### Mathematical Background

Starting from an initial guess $x_0$, the method generates a sequence of approximations using the iteration formula:

$$x_{n+1} = x_n - \frac{f(x_n)}{f'(x_n)}$$

**Geometric Interpretation**: At each step, we:
1. Draw the tangent line to $f(x)$ at the current point $(x_n, f(x_n))$
2. Find where this tangent line crosses the x-axis
3. Use that x-intercept as our next approximation $x_{n+1}$

### Your Task

Implement `newton_raphson(x0, N)` that:
1. Takes an initial guess `x0` and number of iterations `N`
2. Applies the Newton-Raphson iteration formula N times
3. Returns the approximation after N iterations

**Requirements:**
- Perform exactly N iterations (not checking for convergence)
- Return the final approximation as a float

**Parameters:**
- `x0`: float, initial guess for the root
- `N`: int, number of iterations to perform

**Returns:**
- `root`: float, approximation of the root after N iterations

**Note**
- Your function should return the initial guess if `N = 0`

```python
def newton_raphson(x0, N):
    # your code here
```

```python
grader.check("q1a")
```

## **Part B**: Bisection Method

The **bisection method** is a robust root-finding algorithm based on the **Intermediate Value Theorem**. It's slower than Newton-Raphson but guaranteed to converge for continuous functions.

### Mathematical Background

**Intermediate Value Theorem**: If a continuous function $f(x)$ has opposite signs at two points $a$ and $b$ (i.e., $f(a) \cdot f(b) < 0$), then $f(x)$ must cross zero somewhere in the interval $[a, b]$.

**Algorithm**: Starting with an interval $[a, b]$ where $f(a)$ and $f(b)$ have opposite signs:

1. Compute the midpoint: $m = \frac{a + b}{2}$
2. Evaluate $f(m)$
3. Determine which half-interval contains the root:
   - If $f(a) \cdot f(m) < 0$: root is in $[a, m]$ → set $b = m$
   - If $f(m) \cdot f(b) < 0$: root is in $[m, b]$ → set $a = m$
   - If $f(m) = 0$: we found the exact root!
4. Repeat N times

After each iteration, the interval size is halved, so after $N$ iterations, the uncertainty in the root is $\frac{b-a}{2^N}$.

### Your Task

Implement `bisection_method(a, b, N)` that:
1. Takes an initial interval `[a, b]` and number of iterations `N`
2. Verifies that $f(a)$ and $f(b)$ have opposite signs (raise `ValueError` if not)
3. Applies the bisection algorithm N times
4. Returns the midpoint of the final interval

**Requirements:**
- Define the function $f(x) = x^3 - 3$ inside your function
- Check that `f(a) * f(b) < 0`; if not, raise `ValueError("f(a) and f(b) must have opposite signs")`
- Perform exactly N iterations of bisection
- Return the midpoint of the final interval: $(a + b) / 2$

**Parameters:**
- `a`: float, left endpoint of initial interval
- `b`: float, right endpoint of initial interval
- `N`: int, number of iterations to perform

**Returns:**
- `root`: float, approximation of the root (midpoint of final interval)

**Raises:**
- `ValueError`: if f(a) and f(b) do not have opposite signs

```python
def bisection_method(a, b, N):
    # your code here
```

```python
grader.check("q1b")
```

## **Part C**: Using SciPy's `root_scalar`

In practice, we don't need to implement root-finding algorithms from scratch. The **SciPy library** provides optimized, well-tested implementations that handle edge cases and numerical issues automatically.

### Background: `scipy.optimize.root_scalar`

The `scipy.optimize.root_scalar` function provides a unified interface to several root-finding algorithms. It's more sophisticated than our implementations because it:
- Automatically handles convergence criteria
- Uses adaptive algorithms that adjust based on the function's behavior
- Provides detailed information about the solution (convergence status, iterations used, etc.)
- Supports multiple methods (bisection, Newton, Brent's method, etc.)

[root_scalar documentation](https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.root_scalar.html)
### Your Task

Implement `find_root_scipy()` that uses `scipy.optimize.root_scalar` to find the root of $f(x) = x^3 - 3$ using **Brent's method**.

**Requirements:**
- Use the bracket `[0, 2]`
- Return only the root value (`.root` attribute)

**Parameters:**
- None

**Returns:**
- `root`: float, the root found by scipy

```python
def find_root_scipy():
    from scipy.optimize import root_scalar
    
    # Define f(x) = x^3 - 3
    
    # Use root_scalar with method='brentq' and bracket=[0, 2]
    
    # Return the root
```

```python
grader.check("q1c")
```

# **Question 2**: Lagrange Polynomial Interpolation

In this question, you'll implement **Lagrange polynomial interpolation**, a method for constructing a polynomial that passes exactly through a given set of data points.

## Background: Lagrange Interpolation

Given $n$ data points $(x_0, y_0), (x_1, y_1), ..., (x_{n-1}, y_{n-1})$ with distinct $x$-values, we want to find a polynomial $P(x)$ of degree at most $n-1$ such that:

$$P(x_i) = y_i \quad \text{for all } i = 0, 1, ..., n-1$$

That is, the polynomial passes exactly through all the data points.

### The Lagrange Form

The **Lagrange polynomial** is constructed as a weighted sum of **basis polynomials**:

$$P(x) = \sum_{i=0}^{n-1} y_i \cdot L_i(x)$$

where $L_i(x)$ is the **i-th Lagrange basis polynomial**:

$$L_i(x) = \prod_{\substack{j=0 \\ j \neq i}}^{n-1} \frac{x - x_j}{x_i - x_j}$$

### Understanding the Basis Polynomials

Each basis polynomial $L_i(x)$ has a special property:

$$L_i(x_j) = \begin{cases} 1 & \text{if } i = j \\ 0 & \text{if } i \neq j \end{cases}$$

This is called the **Kronecker delta property**. Because of this property:

$$P(x_k) = \sum_{i=0}^{n-1} y_i \cdot L_i(x_k) = y_k \cdot 1 + \sum_{i \neq k} y_i \cdot 0 = y_k$$

So the polynomial automatically passes through all data points!

### Example

For 3 points $(x_0, y_0), (x_1, y_1), (x_2, y_2)$:

$$L_0(x) = \frac{(x - x_1)(x - x_2)}{(x_0 - x_1)(x_0 - x_2)}$$

$$L_1(x) = \frac{(x - x_0)(x - x_2)}{(x_1 - x_0)(x_1 - x_2)}$$

$$L_2(x) = \frac{(x - x_0)(x - x_1)}{(x_2 - x_0)(x_2 - x_1)}$$

$$P(x) = y_0 L_0(x) + y_1 L_1(x) + y_2 L_2(x)$$

### Key Properties

1. **Exact fit**: The polynomial passes exactly through all data points
2. **Degree**: For $n$ points, the polynomial has degree at most $n-1$
3. **Uniqueness**: There is exactly one polynomial of degree ≤ $n-1$ passing through $n$ points with distinct x-values

## Your Task

Implement `lagrange_interpolation(x_data, y_data, show_plot=False)` that:
1. Takes arrays of x and y data points
2. Constructs the Lagrange interpolating polynomial
3. Plots the data points and the interpolating polynomial
4. Returns the figure object

**Requirements:**
- Check that `x_data` and `y_data` have the same length and at least 2 points
- Create an array of x-values for plotting using `np.linspace()` with at least 200 points between `min(x_data)` and `max(x_data)`
- For each plotting x-value, evaluate the Lagrange polynomial by:
  - Computing all basis polynomials $L_i(x)$
  - Computing $P(x) = \sum_{i=0}^{n-1} y_i \cdot L_i(x)$
- Plot configuration:
  - Scatter plot of data with `'o'` markers, `s=80`, `color='red'`, `label='Data Points'`, `zorder=3`
  - Line plot of interpolating polynomial with `linewidth=2`, `color='blue'`, `label='Lagrange Interpolation'`
  - X-axis label: "x"
  - Y-axis label: "y"
  - Title: "Lagrange Polynomial Interpolation"
  - Grid with `alpha=0.3`
  - Legend
- Only call `plt.show()` if `show_plot=True`
- Return the figure object

**Parameters:**
- `x_data`: numpy array or list, x-coordinates of data points
- `y_data`: numpy array or list, y-coordinates of data points
- `show_plot`: bool, default False. If True, display the plot

**Returns:**
- `fig`: matplotlib figure object

```python
def lagrange_interpolation(x_data, y_data, show_plot=False):

    ... 
    
    # Create the plot
    fig, ax = plt.subplots(figsize=(10, 6))
    
    ...
    
    # Show plot if requested
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Example: Lagrange interpolation through different sets of points

# Example 1: Interpolate through points on y = x^2
print("Example 1: Interpolating through points on y = x²")
x1 = np.array([0, 1, 2, 3, 4])
y1 = x1**2
fig1 = lagrange_interpolation(x1, y1, show_plot=True)
print()

# Example 2: Interpolate through points on y = sin(x)
print("Example 2: Interpolating through points on y = sin(x)")
x2 = np.linspace(0,30,10)
y2 = np.sin(x2)
fig2 = lagrange_interpolation(x2, y2, show_plot=True)
print()

# Example 3: Random points
print("Example 3: Interpolating through random points")
np.random.seed(rng_seed)
x3 = np.array([0, 1, 2, 3, 4, 5])
y3 = np.random.randn(6) * 2 + 3
fig3 = lagrange_interpolation(x3, y3, show_plot=True)
```

```python
grader.check("q2")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> result = newton_raphson(2.0, 5)
>>> test_result = isinstance(result, (float, np.floating))
>>> test_result
True
```
```
>>> result = newton_raphson(2.0, 0)
>>> test_result = abs(result - 2.0) < 1e-10
>>> test_result
True
```
```
>>> result = newton_raphson(2.0, 10)
>>> true_root = 3 ** (1 / 3)
>>> test_result = abs(result - true_root) < 1e-08
>>> test_result
True
```
```
>>> result = newton_raphson(1.0, 10)
>>> true_root = 3 ** (1 / 3)
>>> test_result = abs(result - true_root) < 1e-08
>>> test_result
True
```


Tests for q1b (points: None)
```
>>> result = bisection_method(0, 2, 10)
>>> test_result = isinstance(result, (float, np.floating))
>>> test_result
True
```
```
>>> try:
...     bisection_method(2, 3, 10)
...     test_result = False
... except ValueError:
...     test_result = True
>>> test_result
True
```
```
>>> result = bisection_method(0, 2, 20)
>>> true_root = 3 ** (1 / 3)
>>> test_result = abs(result - true_root) < 1e-05
>>> test_result
True
```
```
>>> result = bisection_method(1, 2, 20)
>>> true_root = 3 ** (1 / 3)
>>> test_result = abs(result - true_root) < 1e-05
>>> test_result
True
```


Tests for q1c (points: None)
```
>>> result = find_root_scipy()
>>> test_result = isinstance(result, (float, np.floating))
>>> test_result
True
```
```
>>> result = find_root_scipy()
>>> true_root = 3 ** (1 / 3)
>>> test_result = abs(result - true_root) < 1e-10
>>> test_result
True
```


Tests for q2 (points: None)
```
>>> x_test = [0, 1, 2]
>>> y_test = [0, 1, 4]
>>> fig_test = lagrange_interpolation(x_test, y_test)
>>> test_result = isinstance(fig_test, matplotlib.figure.Figure)
>>> plt.close(fig_test)
>>> test_result
True
```
```
>>> x_test = np.array([0, 1, 2])
>>> y_test = np.array([0, 1, 4])
>>> fig_test = lagrange_interpolation(x_test, y_test)
>>> def eval_lagrange(x, x_data, y_data):
...     n = len(x_data)
...     result = 0
...     for i in range(n):
...         L_i = 1
...         for j in range(n):
...             if i != j:
...                 L_i *= (x - x_data[j]) / (x_data[i] - x_data[j])
...         result += y_data[i] * L_i
...     return result
>>> errors = [abs(eval_lagrange(x_test[i], x_test, y_test) - y_test[i]) for i in range(len(x_test))]
>>> test_result = all((error < 1e-10 for error in errors))
>>> plt.close(fig_test)
>>> test_result
True
```
```
>>> x_test = np.array([0, 1])
>>> y_test = np.array([2, 5])
>>> fig_test = lagrange_interpolation(x_test, y_test)
>>> def eval_lagrange(x, x_data, y_data):
...     n = len(x_data)
...     result = 0
...     for i in range(n):
...         L_i = 1
...         for j in range(n):
...             if i != j:
...                 L_i *= (x - x_data[j]) / (x_data[i] - x_data[j])
...         result += y_data[i] * L_i
...     return result
>>> y_mid = eval_lagrange(0.5, x_test, y_test)
>>> test_result = abs(y_mid - 3.5) < 1e-10
>>> plt.close(fig_test)
>>> bool(test_result)
True
```
```
>>> x_test = np.array([0, 1, 2, 3])
>>> y_test = np.array([0, 1, 8, 27])
>>> fig_test = lagrange_interpolation(x_test, y_test)
>>> ax = fig_test.axes[0]
>>> has_scatter = len(ax.collections) > 0
>>> has_line = len(ax.lines) > 0
>>> has_xlabel = ax.get_xlabel() == 'x'
>>> has_ylabel = ax.get_ylabel() == 'y'
>>> has_title = 'Lagrange' in ax.get_title()
>>> has_legend = ax.get_legend() is not None
>>> plt.close(fig_test)
>>> test_result = has_scatter and has_line and has_xlabel and has_ylabel and has_title and has_legend
>>> test_result
True
```
```
>>> x_test = np.array([-1, 0, 1])
>>> y_test = np.array([1, 0, 1])
>>> fig_test = lagrange_interpolation(x_test, y_test)
>>> def eval_lagrange(x, x_data, y_data):
...     n = len(x_data)
...     result = 0
...     for i in range(n):
...         L_i = 1
...         for j in range(n):
...             if i != j:
...                 L_i *= (x - x_data[j]) / (x_data[i] - x_data[j])
...         result += y_data[i] * L_i
...     return result
>>> y_at_2 = eval_lagrange(2, x_test, y_test)
>>> test_result = abs(y_at_2 - 4.0) < 1e-10
>>> plt.close(fig_test)
>>> bool(test_result)
True
```
