```python
# Initialize Otter
import otter
grader = otter.Notebook("hw3.ipynb")
```

# Homework 3

```python
# Setup
# You may import math if needed in your solutions.
import math
import numpy as np
```

## **Question 1**: NumPy Array Creation

NumPy arrays are fundamental for scientific computing and physics simulations. This question explores different methods of creating numpy arrays for physics applications.

### Q1a: Arrays from Python Objects
A physics lab collects experimental data as Python lists. Write a function `create_physics_arrays(positions, velocities, forces)` that converts the physics data into numpy arrays and returns them as a tuple.

The function should:
- Convert the `positions` list (in meters) to a numpy array
- Convert the `velocities` list (in m/s) to a numpy array  
- Convert the `forces` list (in Newtons) to a numpy array
- Return a tuple: `(position_array, velocity_array, force_array)`

Example: If given positions=[1.0, 2.5, 3.2], velocities=[0.5, 1.2, 0.8], forces=[10.0, 15.5, 12.3], return the corresponding numpy arrays.

```python
def create_physics_arrays(positions, velocities, forces):
    # BEGIN SOLUTION NO PROMPT
    position_array = np.array(positions)
    velocity_array = np.array(velocities)
    force_array = np.array(forces)
    return (position_array, velocity_array, force_array)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (position_array, velocity_array, force_array)
    """; # END PROMPT
```

```python
grader.check("q1a")
```

### Q1b: Built-in Array Creation Functions
In quantum mechanics and linear transformations, we often need special matrices. Write a function `create_physics_matrices(n)` that creates fundamental matrices used in physics simulations.

[Numpy array creation documentation](https://numpy.org/doc/2.3/reference/routines.array-creation.html)

The function should create and return a tuple of three n×n matrices (as numpy arrays):
1. **Zero matrix**: Represents initial conditions where all values are zero.
2. **Identity matrix**: Represents the identity transformation.
3. **Ones matrix**: Matrix where all values are 1.

Return a tuple: `(zero_matrix, identity_matrix, ones_matrix)`

Example: For n=3, return (3×3 zero matrix, 3×3 identity matrix, 3×3 ones matrix)

```python
def create_physics_matrices(n):
    # BEGIN SOLUTION NO PROMPT
    zero_matrix = np.zeros((n, n))
    identity_matrix = np.eye(n)
    ones_matrix = np.ones((n, n))
    return (zero_matrix, identity_matrix, ones_matrix)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (zero_matrix, identity_matrix, ones_matrix)
    """; # END PROMPT
```

```python
grader.check("q1b")
```

### Q1c: Array-like Functions
In computational physics, we often need to create arrays with the same shape as existing data but with different values. Write a function `simulate_field_interactions(electric_field, magnetic_field)` that creates related fields with the same dimensions.

[Numpy array creation documentation](https://numpy.org/doc/2.3/reference/routines.array-creation.html)


Given existing `electric_field` and `magnetic_field` numpy arrays, the function should:
1. Create a **zero potential field** with the same shape as `electric_field`
2. Create a **unit charge density** with the same shape as `electric_field` (1 for each entry)
3. Create a **constant force field** with the same shape as `magnetic_field`, filled with the value 9.8 for all entries (representing gravity)

Return a tuple: `(potential_field, charge_density, force_field)` (all three should be numpy arrays)

Note: The original arrays can have any shape (1D, 2D, 3D, etc.) and the new arrays should match exactly. The arrays returned should be new arrays, not views.

Hint: use the *_like family of numpy functions

```python
def simulate_field_interactions(electric_field, magnetic_field):
    # BEGIN SOLUTION NO PROMPT
    potential_field = np.zeros_like(electric_field)
    charge_density = np.ones_like(electric_field)
    force_field = np.full_like(magnetic_field, 9.8)
    return (potential_field, charge_density, force_field)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (potential_field, charge_density, force_field)
    """; # END PROMPT
```

```python
grader.check("q1c")
```

## **Question 2**: NumPy Array Generation and Mathematical Operations

NumPy provides powerful functions for creating arrays with specific patterns and supports vectorized mathematical operations. This question explores creating arrays and performing calculations efficiently.

### Q2a: Linear Motion Analysis with `linspace`
A projectile follows a parabolic trajectory. Write a function `analyze_projectile_motion(v0, angle, t_max, num_points=50)` that calculates the position and velocity components over time.

The function should:
1. Create a time array from 0 to `t_max` with `num_points` evenly spaced values using `np.linspace`
2. Calculate horizontal position: `x(t) = v0 * cos(angle) * t`
3. Calculate vertical position: `y(t) = v0 * sin(angle) * t - 0.5 * 9.8 * t²`
4. Calculate horizontal velocity: `vx(t) = v0 * cos(angle)` (constant array)
5. Calculate vertical velocity: `vy(t) = v0 * sin(angle) - 9.8 * t`

Parameters:
- `v0`: initial velocity magnitude (m/s)
- `angle`: launch angle in radians
- `t_max`: maximum time (seconds)
- `num_points`: number of time points (default 50)

Return a tuple: `(time_array, x_positions, y_positions, x_velocities, y_velocities)` (all numpy arrays)

Hint: You don't need any for loops for this question, just vectorized elementwise numpy operations

```python
def analyze_projectile_motion(v0, angle, t_max, num_points=50):
    # BEGIN SOLUTION NO PROMPT
    # Create time array using linspace
    time_array = np.linspace(0, t_max, num_points)
    
    # Calculate position components using array operations
    x_positions = v0 * np.cos(angle) * time_array
    y_positions = v0 * np.sin(angle) * time_array - 0.5 * 9.8 * time_array**2
    
    # Calculate velocity components using array operations
    x_velocities = np.full_like(time_array, v0 * np.cos(angle))  # constant horizontal velocity
    y_velocities = v0 * np.sin(angle) - 9.8 * time_array
    
    return (time_array, x_positions, y_positions, x_velocities, y_velocities)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (time_array, x_positions, y_positions, x_velocities, y_velocities)
    """; # END PROMPT
```

```python
grader.check("q2a")
```

### Q2b: Exponential Decay Analysis with `logspace`
Radioactive decay follows an exponential law. Write a function `analyze_exponential_decay(N0, half_life, decades=5, num_points=20)` that models radioactive decay over logarithmic time scales.

The function should:
1. Create a time array spanning `decades` orders of magnitude using `np.logspace`, from 1 year to 10^`decades` years
2. Calculate the decay constant: `lambda = ln(2) / half_life`
3. Calculate the remaining atoms: `N(t) = N0 * exp(-lambda * t)`
4. Calculate the decay rate: `R(t) = lambda * N(t)` (decays per year)
5. Calculate the activity ratio: `A(t) = N(t) / N0` (fraction remaining)

Parameters:
- `N0`: initial number of atoms
- `half_life`: half-life in years
- `decades`: number of decades to span (default 5, so 1 to 10^5 years)
- `num_points`: number of logarithmically spaced time points (default 20)

Return a tuple: `(time_array, atoms_remaining, decay_rates, activity_ratios)` (all numpy arrays)

Hint: np.exp() will exponientiate an array elementwise, efficiently

```python
def analyze_exponential_decay(N0, half_life, decades=5, num_points=20):
    # BEGIN SOLUTION NO PROMPT
    # Create logarithmic time array using logspace
    time_array = np.logspace(0, decades, num_points)  # 10^0 to 10^decades years
    
    # Calculate decay constant
    decay_constant = np.log(2) / half_life
    
    # Calculate quantities using array operations
    atoms_remaining = N0 * np.exp(-decay_constant * time_array)
    decay_rates = decay_constant * atoms_remaining
    activity_ratios = atoms_remaining / N0
    
    return (time_array, atoms_remaining, decay_rates, activity_ratios)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (time_array, atoms_remaining, decay_rates, activity_ratios)
    """; # END PROMPT
```

```python
grader.check("q2b")
```

### Q2c: Wave Interference with `arange` and Array Operations
Two waves with different frequencies create interference patterns. Write a function `calculate_wave_interference(f1, f2, amplitude, duration, sample_rate=100)` that calculates the combined wave and analyzes its properties.

The function should:
1. Create a time array using `np.arange` from 0 to `duration` with step size `1/sample_rate`
2. Calculate wave 1: `y1(t) = amplitude * sin(2π * f1 * t)`
3. Calculate wave 2: `y2(t) = amplitude * sin(2π * f2 * t)`
4. Calculate the combined wave: `y_total(t) = y1(t) + y2(t)`
5. Calculate the instantaneous power: `P(t) = y_total(t)²`
6. Calculate the RMS amplitude: `A_rms = sqrt(mean(y_total²))`

Parameters:
- `f1`, `f2`: frequencies of the two waves (Hz)
- `amplitude`: amplitude of each individual wave (same amplitude)
- `duration`: total time duration (seconds)
- `sample_rate`: samples per second (default 100 Hz)

Return a tuple: `(time_array, wave1, wave2, combined_wave, power, rms_amplitude)` 
Note: `rms_amplitude` should be a single float value, others are arrays.

Note: looping over arrays is unnecessary, look up numpy functions that perform operations you need to keep it efficient.

```python
def calculate_wave_interference(f1, f2, amplitude, duration, sample_rate=100):
    # BEGIN SOLUTION NO PROMPT
    # Create time array using arange
    time_step = 1.0 / sample_rate
    time_array = np.arange(0, duration, time_step)
    
    # Calculate individual waves using array operations
    wave1 = amplitude * np.sin(2 * np.pi * f1 * time_array)
    wave2 = amplitude * np.sin(2 * np.pi * f2 * time_array)
    
    # Calculate combined wave and derived quantities
    combined_wave = wave1 + wave2
    power = combined_wave**2
    rms_amplitude = np.sqrt(np.mean(combined_wave**2))
    
    return (time_array, wave1, wave2, combined_wave, power, rms_amplitude)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (time_array, wave1, wave2, combined_wave, power, rms_amplitude)
    """; # END PROMPT
```

```python
grader.check("q2c")
```

## **Question 3**: NumPy Array Indexing and Slicing 

### Q3a: Time Series Data Analysis with 1D Indexing
A motion sensor records position data over time. Write a function `analyze_motion_data(positions, sample_rate)` that extracts key information from the position array using indexing and slicing.

The function should extract:
1. **Initial position**: The first element of the array
2. **Final position**: The last element of the array
3. **Middle section**: The middle third of the data (from 1/3 to 2/3 of total length)
4. **Every 10th sample**: Every 10th data point starting from index 0
5. **Last 5 seconds of data**: Extract data from the last 5 seconds (use `sample_rate` to determine how many samples)

Parameters:
- `positions`: 1D numpy array of position measurements (meters)
- `sample_rate`: sampling frequency (samples per second)

Return a tuple: `(initial_pos, final_pos, middle_section, every_10th, last_5_seconds)`

Note: For the middle section, use `int(n/3)` to `int(2*n/3)` where n is array length to get starting and ending indices, include the element at `int(n/3)` and exclude the element at `int(2*n/3)` (this is default for slicing). For last 5 seconds, calculate the number of samples as `5 * sample_rate` and slice from the end (similarly start at `int(5*sample_rate)` from the end).

```python
def analyze_motion_data(positions, sample_rate):
    # BEGIN SOLUTION NO PROMPT
    # Extract initial and final positions using indexing
    initial_pos = positions[0]
    final_pos = positions[-1]
    
    # Extract middle third using slicing
    n = len(positions)
    start_idx = int(n / 3)
    end_idx = int(2 * n / 3)
    middle_section = positions[start_idx:end_idx]
    
    # Extract every 10th sample using step slicing
    every_10th = positions[::10]
    
    # Extract last 5 seconds of data
    samples_in_5_sec = int(5 * sample_rate)
    last_5_seconds = positions[-samples_in_5_sec:]
    
    return (initial_pos, final_pos, middle_section, every_10th, last_5_seconds)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (initial_pos, final_pos, middle_section, every_10th, last_5_seconds)
    """; # END PROMPT
```

```python
grader.check("q3a")
```

### Q3b: 2D Temperature Field Analysis
A thermal imaging camera captures temperature measurements in a 2D grid. Write a function `analyze_temperature_field(temp_field)` that extracts specific regions and measurements from the 2D temperature array.

The function should extract:
1. **Corner temperatures**: The four corner values [top-left, top-right, bottom-left, bottom-right] (as an array)
2. **Center region**: A 3×3 subarray centered on the middle of the field
3. **Top row**: The entire first row of temperature measurements  
4. **Right column**: The entire last column of temperature measurements
5. **Diagonal temperatures**: The main diagonal elements (from top-left to bottom-right)

All of the outputs in the tuple should be numpy arrays

Parameters:
- `temp_field`: 2D numpy array of temperature measurements (rows = y-positions, columns = x-positions)

Return a tuple: `(corners, center_region, top_row, right_column, diagonal)`

Note: For the center region, assume the input array has odd side lengths.

```python
def analyze_temperature_field(temp_field):
    # BEGIN SOLUTION NO PROMPT
    # Extract corner temperatures using 2D indexing
    corners = np.array([
        temp_field[0, 0],    # top-left
        temp_field[0, -1],   # top-right
        temp_field[-1, 0],   # bottom-left
        temp_field[-1, -1]   # bottom-right
    ])
    
    # Extract center 3x3 region
    rows, cols = temp_field.shape
    center_row, center_col = rows // 2, cols // 2
    center_region = temp_field[center_row-1:center_row+2, center_col-1:center_col+2]
    
    # Extract top row and right column
    top_row = temp_field[0, :]
    right_column = temp_field[:, -1]
    
    # Extract main diagonal
    diagonal = np.diag(temp_field)
    
    return (corners, center_region, top_row, right_column, diagonal)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (corners, center_region, top_row, right_column, diagonal)
    """; # END PROMPT
```

```python
grader.check("q3b")
```

## **Question 4**: NumPy Boolean Masks and Conditional Operations

### Q4a: Particle Detection and Filtering
A particle detector measures the energy and velocity of particles passing through it. Write a function `filter_particle_data(energies, velocities, min_energy, max_velocity)` that uses boolean masks to filter and analyze the particle data.

The function should:
1. Create a **high-energy mask**: Boolean array where `energies >= min_energy`
2. Create a **slow-particle mask**: Boolean array where `velocities <= max_velocity` 
3. Create a **valid-particle mask**: Combine both conditions (high energy AND slow velocity)
4. Filter the **valid energies**: Use the valid-particle mask to extract energies that meet both criteria
5. Filter the **valid velocities**: Use the valid-particle mask to extract velocities that meet both criteria
6. Calculate the **detection efficiency**: Fraction of particles that meet both criteria

Parameters:
- `energies`: 1D numpy array of particle energies (GeV)
- `velocities`: 1D numpy array of particle velocities (km/s)
- `min_energy`: minimum energy threshold (GeV)
- `max_velocity`: maximum velocity threshold (km/s)

Return a tuple: `(high_energy_mask, slow_particle_mask, valid_particle_mask, valid_energies, valid_velocities, detection_efficiency)`

Note: `detection_efficiency` should be a float between 0 and 1.

```python
def filter_particle_data(energies, velocities, min_energy, max_velocity):
    # BEGIN SOLUTION NO PROMPT
    # Create boolean masks using comparison operators
    high_energy_mask = energies >= min_energy
    slow_particle_mask = velocities <= max_velocity
    
    # Combine conditions using logical AND
    valid_particle_mask = high_energy_mask & slow_particle_mask
    
    # Filter arrays using the boolean mask
    valid_energies = energies[valid_particle_mask]
    valid_velocities = velocities[valid_particle_mask]
    
    # Calculate detection efficiency
    detection_efficiency = np.sum(valid_particle_mask) / len(energies)
    
    return (high_energy_mask, slow_particle_mask, valid_particle_mask, 
            valid_energies, valid_velocities, detection_efficiency)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (high_energy_mask, slow_particle_mask, valid_particle_mask, 
            valid_energies, valid_velocities, detection_efficiency)
    """; # END PROMPT
```

```python
grader.check("q4a")
```

### Q4b: Experimental Data Quality Control
A physics experiment collects temperature and pressure measurements, but some readings may be invalid due to sensor errors. Write a function `analyze_sensor_data(temperatures, pressures, temp_tolerance=5.0, pressure_range=(0.8, 1.2))` that uses boolean operations to identify and replace invalid measurements.

The function should:
1. **Identify temperature outliers**: Create a boolean mask for temperatures that deviate more than `temp_tolerance` from the median temperature
2. **Identify pressure outliers**: Create a boolean mask for pressures outside the valid range `[pressure_range[0], pressure_range[1]]`
3. **Create valid data mask**: Boolean array where BOTH temperature and pressure are valid (not outliers)
4. **Clean temperatures**: Replace invalid temperatures with the median temperature of valid readings
5. **Clean pressures**: Replace invalid pressures with 1.0 (atmospheric pressure)
6. **Calculate data quality**: Fraction of measurements that were originally valid

Parameters:
- `temperatures`: 1D numpy array of temperature measurements (°C)
- `pressures`: 1D numpy array of pressure measurements (atm)
- `temp_tolerance`: maximum allowed deviation from median temperature (default 5.0°C)
- `pressure_range`: tuple of (min_pressure, max_pressure) for valid range (default (0.8, 1.2))

Return a tuple: `(temp_outlier_mask, pressure_outlier_mask, valid_data_mask, clean_temperatures, clean_pressures, data_quality)`

Hint: Use `np.where()` to conditionally replace values, and `np.median()` to find the median.

Hint: with boolean arrays, `~` is not, `&` is and, and `|` is or

```python
def analyze_sensor_data(temperatures, pressures, temp_tolerance=5.0, pressure_range=(0.8, 1.2)):
    # BEGIN SOLUTION NO PROMPT
    # Calculate median temperature for outlier detection
    temp_median = np.median(temperatures)
    
    # Create boolean masks for outliers
    temp_outlier_mask = np.abs(temperatures - temp_median) > temp_tolerance
    pressure_outlier_mask = (pressures < pressure_range[0]) | (pressures > pressure_range[1])
    
    # Create valid data mask (NOT outliers for both)
    valid_data_mask = ~temp_outlier_mask & ~pressure_outlier_mask
    
    # Clean the data using np.where
    # For temperatures: replace outliers with median of valid temperatures
    valid_temp_median = np.median(temperatures[~temp_outlier_mask])
    clean_temperatures = np.where(temp_outlier_mask, valid_temp_median, temperatures)
    
    # For pressures: replace outliers with 1.0 atm
    clean_pressures = np.where(pressure_outlier_mask, 1.0, pressures)
    
    # Calculate data quality
    data_quality = np.sum(valid_data_mask) / len(temperatures)
    
    return (temp_outlier_mask, pressure_outlier_mask, valid_data_mask,
            clean_temperatures, clean_pressures, data_quality)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (temp_outlier_mask, pressure_outlier_mask, valid_data_mask,
            clean_temperatures, clean_pressures, data_quality)
    """; # END PROMPT
```

```python
grader.check("q4b")
```

## **Question 5**: Electric Field Calculations with Broadcasting

In electrostatics, we need to calculate the electric field at multiple observation points due to multiple point charges. Write a function `calculate_electric_field(charges, charge_positions, observation_points)` that uses broadcasting to efficiently compute the electric field vectors.

The function should:
1. **Calculate displacement vectors**: For each observation point, find the displacement vector from each charge position (use broadcasting to avoid loops)
2. **Calculate distances**: Find the distance from each charge to each observation point 
3. **Calculate unit vectors**: Normalize the displacement vectors to get unit direction vectors
4. **Calculate field magnitudes**: Use Coulomb's law: `E = k * |q| / r²` where k = 8.99e9 N⋅m²/C²
5. **Calculate field vectors**: Multiply field magnitudes by unit vectors and sum contributions from all charges
6. **Calculate total field magnitude**: Find the magnitude of the total electric field at each observation point

Parameters:
- `charges`: 1D array of point charges (Coulombs) - shape (N,)
- `charge_positions`: 2D array of charge positions (meters) - shape (N, 2) for x,y coordinates  
- `observation_points`: 2D array of observation points (meters) - shape (M, 2) for x,y coordinates

Return a tuple: `(displacement_vectors, distances, field_vectors, total_field_magnitudes)`
- `displacement_vectors`: shape (M, N, 2) - displacement from each charge to each observation point
- `distances`: shape (M, N) - distance from each charge to each observation point  
- `field_vectors`: shape (M, 2) - total electric field vector at each observation point
- `total_field_magnitudes`: shape (M,) - magnitude of total field at each observation point

Note: Use broadcasting to avoid explicit loops. The displacement calculation should use `observation_points[:, np.newaxis, :] - charge_positions[np.newaxis, :, :]`.

```python
def calculate_electric_field(charges, charge_positions, observation_points):
    # BEGIN SOLUTION NO PROMPT
    k = 8.99e9  # Coulomb's constant in N⋅m²/C²
    
    # Use broadcasting to calculate displacement vectors: (M, N, 2)
    # observation_points is (M, 2), charge_positions is (N, 2)
    displacement_vectors = observation_points[:, np.newaxis, :] - charge_positions[np.newaxis, :, :]
    
    # Calculate distances: (M, N)
    distances = np.linalg.norm(displacement_vectors, axis=2)
    
    # Calculate unit vectors (avoiding division by zero)
    unit_vectors = displacement_vectors / distances[:, :, np.newaxis]
    
    # Calculate field magnitudes using Coulomb's law: (M, N)
    field_magnitudes = k * np.abs(charges)[np.newaxis, :] / (distances**2)
    
    # Account for charge sign in field direction
    charge_signs = np.sign(charges)[np.newaxis, :, np.newaxis]
    
    # Calculate individual field vectors: (M, N, 2)
    individual_fields = field_magnitudes[:, :, np.newaxis] * unit_vectors * charge_signs
    
    # Sum contributions from all charges: (M, 2)
    field_vectors = np.sum(individual_fields, axis=1)
    
    # Calculate total field magnitudes: (M,)
    total_field_magnitudes = np.linalg.norm(field_vectors, axis=1)
    
    return (displacement_vectors, distances, field_vectors, total_field_magnitudes)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (displacement_vectors, distances, field_vectors, total_field_magnitudes)
    """; # END PROMPT
```

```python
grader.check("q5")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
# BEGIN SOLUTION NO PROMPT
# END SOLUTION
""" # BEGIN PROMPT
"""
# write ai disclosure here:

"""
"""; # END PROMPT
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> pos, vel, force = create_physics_arrays([1.0, 2.0], [0.5, 1.0], [10.0, 20.0])
>>> isinstance(pos, np.ndarray) and isinstance(vel, np.ndarray) and isinstance(force, np.ndarray)
True
```
```
>>> pos, vel, force = create_physics_arrays([1.5, 2.5, 3.5], [0.1, 0.2, 0.3], [5.0, 10.0, 15.0])
>>> np.allclose(pos, [1.5, 2.5, 3.5]) and np.allclose(vel, [0.1, 0.2, 0.3]) and np.allclose(force, [5.0, 10.0, 15.0])
True
```
```
>>> pos, vel, force = create_physics_arrays([3.14], [2.71], [9.8])
>>> bool(pos[0] == 3.14 and vel[0] == 2.71 and (force[0] == 9.8))
True
```
```
>>> pos, vel, force = create_physics_arrays([], [], [])
>>> len(pos) == 0 and len(vel) == 0 and (len(force) == 0) and isinstance(pos, np.ndarray)
True
```
```
>>> pos, vel, force = create_physics_arrays([1, 2.5, 3], [0.5, 1, 1.5], [10, 15.5, 20])
>>> bool(pos.dtype == np.float64 and vel.dtype == np.float64 and (force.dtype == np.float64) and np.allclose(pos, [1.0, 2.5, 3.0]))
True
```


Tests for q1b (points: None)
```
>>> zeros, identity, ones = create_physics_matrices(2)
>>> zeros.shape == (2, 2) and identity.shape == (2, 2) and (ones.shape == (2, 2))
True
```
```
>>> zeros, identity, ones = create_physics_matrices(3)
>>> np.allclose(zeros, np.zeros((3, 3))) and np.allclose(identity, np.eye(3)) and np.allclose(ones, np.ones((3, 3)))
True
```
```
>>> zeros, identity, ones = create_physics_matrices(4)
>>> bool(np.all(np.diag(identity) == 1) and np.sum(identity) == 4)
True
```
```
>>> zeros, identity, ones = create_physics_matrices(1)
>>> bool(zeros[0, 0] == 0.0 and identity[0, 0] == 1.0 and (ones[0, 0] == 1.0))
True
```
```
>>> zeros, identity, ones = create_physics_matrices(5)
>>> bool(zeros.dtype == np.float64 and np.count_nonzero(zeros) == 0 and (np.sum(ones) == 25))
True
```


Tests for q1c (points: None)
```
>>> e_field = np.array([1.0, 2.0, 3.0])
>>> b_field = np.array([4.0, 5.0])
>>> potential, charge, force = simulate_field_interactions(e_field, b_field)
>>> potential.shape == (3,) and charge.shape == (3,) and (force.shape == (2,))
True
```
```
>>> e_field = np.array([[1.0, 2.0], [3.0, 4.0]])
>>> b_field = np.array([5.0, 6.0, 7.0])
>>> potential, charge, force = simulate_field_interactions(e_field, b_field)
>>> np.allclose(potential, 0.0) and np.allclose(charge, 1.0) and np.allclose(force, 9.8)
True
```
```
>>> e_field_3d = np.ones((2, 3, 4))
>>> b_field_3d = np.zeros((5, 2))
>>> potential, charge, force = simulate_field_interactions(e_field_3d, b_field_3d)
>>> potential.shape == (2, 3, 4) and charge.shape == (2, 3, 4) and (force.shape == (5, 2)) and np.allclose(force, 9.8)
True
```
```
>>> e_field = np.array([1, 2, 3], dtype=np.int32)
>>> b_field = np.array([1.0, 2.0], dtype=np.float64)
>>> potential, charge, force = simulate_field_interactions(e_field, b_field)
>>> potential.dtype == e_field.dtype and charge.dtype == e_field.dtype and (force.dtype == b_field.dtype) and np.allclose(force, 9.8)
True
```
```
>>> e_field_scalar = np.array(5.0)
>>> b_field_scalar = np.array(2.0)
>>> potential, charge, force = simulate_field_interactions(e_field_scalar, b_field_scalar)
>>> potential.shape == () and charge.shape == () and (force.shape == ()) and (float(force) == 9.8)
True
```
```
>>> e_field = np.array([1.0, 2.0])
>>> b_field = np.array([3.0, 4.0])
>>> potential, charge, force = simulate_field_interactions(e_field, b_field)
>>> potential[0] = 999
>>> bool(e_field[0] == 1.0 and potential[0] == 999)
True
```


Tests for q2a (points: None)
```
>>> t, x, y, vx, vy = analyze_projectile_motion(10, np.pi / 4, 2.0, num_points=10)
>>> len(t) == 10 and len(x) == 10 and (len(y) == 10) and (len(vx) == 10) and (len(vy) == 10)
True
```
```
>>> t, x, y, vx, vy = analyze_projectile_motion(20, np.pi / 6, 3.0)
>>> bool(abs(x[0] - 0.0) < 1e-10 and abs(y[0] - 0.0) < 1e-10 and (abs(vx[0] - 20 * np.cos(np.pi / 6)) < 1e-10))
True
```
```
>>> t, x, y, vx, vy = analyze_projectile_motion(15, np.pi / 3, 4.0, num_points=25)
>>> bool(np.allclose(vx, vx[0]) and abs(t[0] - 0.0) < 1e-10 and (abs(t[-1] - 4.0) < 1e-10))
True
```
```
>>> t, x, y, vx, vy = analyze_projectile_motion(10, np.pi / 2, 2.0, num_points=100)
>>> peak_index = np.argmin(np.abs(vy))
>>> bool(abs(vy[peak_index]) < 0.5)
True
```
```
>>> t, x, y, vx, vy = analyze_projectile_motion(19.6, np.pi / 2, 4.0)
>>> max_height = np.max(y)
>>> bool(abs(max_height - 19.6) < 1.0)
True
```


Tests for q2b (points: None)
```
>>> t, N, R, A = analyze_exponential_decay(1000, 100, decades=3, num_points=10)
>>> len(t) == 10 and len(N) == 10 and (len(R) == 10) and (len(A) == 10)
True
```
```
>>> t, N, R, A = analyze_exponential_decay(1000, 50)
>>> bool(abs(A[0] - np.exp(-np.log(2) / 50)) < 1e-10 and abs(N[0] - 1000 * np.exp(-np.log(2) / 50)) < 1e-10)
True
```
```
>>> t, N, R, A = analyze_exponential_decay(500, 25, decades=2, num_points=5)
>>> decay_lambda = np.log(2) / 25
>>> np.allclose(R, decay_lambda * N) and np.allclose(A, N / 500)
True
```
```
>>> t, N, R, A = analyze_exponential_decay(1000, 50, decades=3, num_points=4)
>>> t[0] == 1.0 and t[-1] == 1000.0 and (len(t) == 4)
True
```
```
>>> t, N, R, A = analyze_exponential_decay(1000, 100, decades=2, num_points=100)
>>> half_life_idx = np.argmin(np.abs(t - 100))
>>> bool(abs(N[half_life_idx] - 500) < 50)
True
```


Tests for q2c (points: None)
```
>>> t, y1, y2, y_total, P, A_rms = calculate_wave_interference(1, 2, 1.0, 1.0, sample_rate=50)
>>> len(t) == 50 and len(y1) == 50 and (len(y2) == 50) and (len(y_total) == 50) and (len(P) == 50)
True
```
```
>>> t, y1, y2, y_total, P, A_rms = calculate_wave_interference(0, 0, 2.0, 2.0)
>>> np.allclose(y1, 0.0) and np.allclose(y2, 0.0) and np.allclose(y_total, 0.0)
True
```
```
>>> t, y1, y2, y_total, P, A_rms = calculate_wave_interference(1, 1, 1.0, 4.0)
>>> np.allclose(P, y_total ** 2) and isinstance(A_rms, (int, float, np.number))
True
```
```
>>> t, y1, y2, y_total, P, A_rms = calculate_wave_interference(5, 10, 1.5, 0.5, sample_rate=200)
>>> bool(abs(t[1] - t[0] - 1.0 / 200) < 1e-10 and t[-1] < 0.5)
True
```
```
>>> t, y1, y2, y_total, P, A_rms = calculate_wave_interference(1, 1, 1.0, 2.0)
>>> np.allclose(y_total, 2.0 * y1)
True
```
```
>>> t, y1, y2, y_total, P, A_rms = calculate_wave_interference(0, 0, 3.0, 1.0)
>>> expected_rms = 0.0
>>> bool(abs(A_rms - expected_rms) < 1e-10)
True
```


Tests for q3a (points: None)
```
>>> pos_data = np.array([0.0, 1.0, 2.0, 3.0, 4.0, 5.0, 6.0, 7.0, 8.0, 9.0])
>>> initial, final, middle, every_10th, last_5sec = analyze_motion_data(pos_data, 2)
>>> bool(initial == 0.0 and final == 9.0)
True
```
```
>>> pos_data = np.arange(12)
>>> initial, final, middle, every_10th, last_5sec = analyze_motion_data(pos_data, 1)
>>> np.array_equal(middle, np.array([4, 5, 6, 7]))
True
```
```
>>> pos_data = np.arange(20)
>>> initial, final, middle, every_10th, last_5sec = analyze_motion_data(pos_data, 2)
>>> np.array_equal(every_10th, np.array([0, 10])) and len(last_5sec) == 10
True
```
```
>>> pos_data = np.array([1.0, 2.0, 3.0])
>>> initial, final, middle, every_10th, last_5sec = analyze_motion_data(pos_data, 10)
>>> bool(initial == 1.0 and final == 3.0 and (len(middle) == 1) and (middle[0] == 2.0))
True
```
```
>>> pos_data = np.arange(25)
>>> initial, final, middle, every_10th, last_5sec = analyze_motion_data(pos_data, 3.5)
>>> bool(len(every_10th) == 3 and every_10th[0] == 0 and (every_10th[1] == 10) and (every_10th[2] == 20))
True
```


Tests for q3b (points: None)
```
>>> temp_data = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])
>>> corners, center, top, right, diag = analyze_temperature_field(temp_data)
>>> np.array_equal(corners, [1, 3, 7, 9])
True
```
```
>>> temp_data = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])
>>> corners, center, top, right, diag = analyze_temperature_field(temp_data)
>>> np.array_equal(top, [1, 2, 3]) and np.array_equal(right, [3, 6, 9])
True
```
```
>>> temp_data = np.arange(25).reshape(5, 5)
>>> corners, center, top, right, diag = analyze_temperature_field(temp_data)
>>> center.shape == (3, 3) and np.array_equal(diag, [0, 6, 12, 18, 24])
True
```
```
>>> temp_data = np.ones((7, 7)) * 10
>>> temp_data[2:5, 2:5] = 99
>>> corners, center, top, right, diag = analyze_temperature_field(temp_data)
>>> np.allclose(center, 99) and center.shape == (3, 3)
True
```


Tests for q4a (points: None)
```
>>> energies = np.array([1.0, 5.0, 3.0, 8.0, 2.0])
>>> velocities = np.array([100, 200, 150, 250, 120])
>>> he_mask, sp_mask, vp_mask, ve, vv, eff = filter_particle_data(energies, velocities, 3.0, 200)
>>> np.array_equal(he_mask, [False, True, True, True, False]) and np.array_equal(sp_mask, [True, True, True, False, True])
True
```
```
>>> energies = np.array([1.0, 5.0, 3.0, 8.0, 2.0])
>>> velocities = np.array([100, 200, 150, 250, 120])
>>> he_mask, sp_mask, vp_mask, ve, vv, eff = filter_particle_data(energies, velocities, 3.0, 200)
>>> np.array_equal(vp_mask, [False, True, True, False, False]) and np.array_equal(ve, [5.0, 3.0]) and np.array_equal(vv, [200, 150])
True
```
```
>>> energies = np.array([10.0, 2.0, 8.0, 1.0])
>>> velocities = np.array([50, 100, 60, 120])
>>> he_mask, sp_mask, vp_mask, ve, vv, eff = filter_particle_data(energies, velocities, 5.0, 100)
>>> bool(abs(eff - 0.5) < 1e-10)
True
```
```
>>> energies = np.array([10.0, 15.0, 12.0])
>>> velocities = np.array([50, 80, 60])
>>> he_mask, sp_mask, vp_mask, ve, vv, eff = filter_particle_data(energies, velocities, 5.0, 100)
>>> eff == 1.0 and len(ve) == 3 and np.array_equal(ve, energies)
True
```
```
>>> energies = np.array([1.0, 2.0, 3.0])
>>> velocities = np.array([200, 300, 250])
>>> he_mask, sp_mask, vp_mask, ve, vv, eff = filter_particle_data(energies, velocities, 10.0, 100)
>>> eff == 0.0 and len(ve) == 0 and (len(vv) == 0)
True
```


Tests for q4b (points: None)
```
>>> temps = np.array([20.0, 22.0, 25.0, 30.0, 21.0])
>>> pressures = np.array([1.0, 0.5, 1.1, 1.5, 0.9])
>>> t_out, p_out, valid, clean_t, clean_p, quality = analyze_sensor_data(temps, pressures)
>>> np.array_equal(t_out, [False, False, False, True, False]) and np.array_equal(p_out, [False, True, False, True, False])
True
```
```
>>> temps = np.array([10.0, 12.0, 15.0, 25.0])
>>> pressures = np.array([1.0, 1.1, 0.7, 0.9])
>>> t_out, p_out, valid, clean_t, clean_p, quality = analyze_sensor_data(temps, pressures, temp_tolerance=5.0)
>>> bool(np.array_equal(valid, [True, True, False, False]) and abs(quality - 0.5) < 1e-10)
True
```
```
>>> temps = np.array([50.0, 52.0, 49.0, 60.0])
>>> pressures = np.array([2.0, 2.1, 1.9, 2.5])
>>> t_out, p_out, valid, clean_t, clean_p, quality = analyze_sensor_data(temps, pressures, temp_tolerance=2.0, pressure_range=(1.8, 2.2))
>>> bool(np.array_equal(valid, [True, True, True, False]) and abs(quality - 0.75) < 1e-10)
True
```


Tests for q5 (points: None)
```
>>> charges = np.array([1e-09, -2e-09])
>>> charge_pos = np.array([[0, 0], [1, 0]])
>>> obs_points = np.array([[0.5, 0], [0.5, 1]])
>>> disp, dist, field, mag = calculate_electric_field(charges, charge_pos, obs_points)
>>> disp.shape == (2, 2, 2) and dist.shape == (2, 2) and (field.shape == (2, 2)) and (mag.shape == (2,))
True
```
```
>>> charges = np.array([1e-09])
>>> charge_pos = np.array([[0, 0]])
>>> obs_points = np.array([[3, 4]])
>>> disp, dist, field, mag = calculate_electric_field(charges, charge_pos, obs_points)
>>> bool(np.allclose(disp[0, 0], [3, 4]) and abs(dist[0, 0] - 5.0) < 1e-10)
True
```
```
>>> charges = np.array([1e-09])
>>> charge_pos = np.array([[0, 0]])
>>> obs_points = np.array([[1, 0]])
>>> disp, dist, field, mag = calculate_electric_field(charges, charge_pos, obs_points)
>>> k = 8990000000.0
>>> expected_field_mag = k * 1e-09 / 1.0 ** 2
>>> bool(abs(mag[0] - expected_field_mag) < 1e-06 and field[0, 0] > 0 and (abs(field[0, 1]) < 1e-10))
True
```
```
>>> charges = np.array([1e-09, -1e-09])
>>> charge_pos = np.array([[0, 0], [2, 0]])
>>> obs_points = np.array([[1, 0]])
>>> disp, dist, field, mag = calculate_electric_field(charges, charge_pos, obs_points)
>>> bool(mag[0] > 0)
True
```
```
>>> charges = np.array([-2e-09])
>>> charge_pos = np.array([[0, 0]])
>>> obs_points = np.array([[1, 0], [-1, 0]])
>>> disp, dist, field, mag = calculate_electric_field(charges, charge_pos, obs_points)
>>> field[0, 0] < 0 and field[1, 0] > 0 and np.allclose(mag, mag[0])
True
```
