```python
# Initialize Otter
import otter
grader = otter.Notebook("ws3.ipynb")
```

# Workshop 3

```python
# Setup
import math
import numpy as np
```

## **Question 1**: NumPy Array Creation for Physics Applications

### Q1a: Temperature Measurements Array

Write a function `create_temperature_array(temp_list)` that:
- Takes a Python list of temperature values
- Returns a NumPy array made from the list

Example: `create_temperature_array([20.5, 25.3, 30.1])` should return a NumPy array with those values.

```python
import numpy as np #so numpy is visible as np to otter grader

def create_temperature_array(temp_list):
    # BEGIN SOLUTION NO PROMPT
    return np.array(temp_list)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    """; # END PROMPT
```

```python
grader.check("q1a")
```

### Q1b: Time and Frequency Arrays for Wave Analysis
For analyzing electromagnetic waves, we need to create arrays representing time intervals and frequency ranges.

[Numpy array creation documentation](https://numpy.org/doc/stable/reference/routines.array-creation.html)

Write a function `create_wave_arrays(t_start, t_end, num_points, f_min, f_max, num_freq)` that:
- Creates a time array from `t_start` to `t_end` with `num_points` evenly spaced values
- Creates a frequency array from `f_min` to `f_max` with `num_freq` logarithmically spaced values
- Returns a tuple `(time_array, frequency_array)`

Note: For the loagrithmic spacing, you'll need to decide whether to use `np.log10(f_min)` and `np.log10(f_max)` or `f_min` and `f_max` as the start and stop values in the numpy loagarithm spacing array creation function.

Example: `create_wave_arrays(0, 1, 5, 1, 1000, 4)` should return time array [0, 0.25, 0.5, 0.75, 1] and frequency array [1, 10, 100, 1000].

```python
import numpy as np #so numpy is visible as np to otter grader

def create_wave_arrays(t_start, t_end, num_points, f_min, f_max, num_freq):
    # BEGIN SOLUTION NO PROMPT
    time_array = np.linspace(t_start, t_end, num_points)
    frequency_array = np.logspace(np.log10(f_min), np.log10(f_max), num_freq)
    return (time_array, frequency_array)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (time_array, frequency_array)
    """; # END PROMPT
```

```python
grader.check("q1b")
```

### Q1c: Initializing Measurement Arrays for Data Collection

[Numpy array creation documentation](https://numpy.org/doc/stable/reference/routines.array-creation.html)

Write a function `initialize_experiment_arrays(reference_array, initial_voltage)` that:
- Takes a `reference_array` (existing NumPy array) and an `initial_voltage` value
- Creates a zeros array with the same shape as `reference_array`
- Creates an array filled with `initial_voltage` values with the same shape as `reference_array`
- Creates a separate 1D zeros array of length 10
- Returns a tuple `(zeros_array, voltage_array, data_buffer)` where `data_buffer` is the 1D zeros array

Example: If `reference_array` has shape (3,) and `initial_voltage` is 5.0, the function should return three arrays:
- A zeros array of shape (3,)
- A voltage array of shape (3,) filled with 5.0
- A 1D zeros array of length 10

Note: For the `voltage_array`, the input may be all ints, and the initial voltage may not be an int. The usual numpy function for this purpose by default inherits the datatype of the array whose shape you're copying, but this can result in floats being rounded to ints and giving unexpected results, for this reason, make sure the `voltage_array` has float datatype.

```python
import numpy as np #so numpy is visible as np to otter grader

def initialize_experiment_arrays(reference_array, initial_voltage):
    # BEGIN SOLUTION NO PROMPT
    zeros_array = np.zeros_like(reference_array)
    voltage_array = np.full_like(reference_array, initial_voltage, dtype=float)
    data_buffer = np.zeros(10)
    return (zeros_array, voltage_array, data_buffer)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (zeros_array, voltage_array, data_buffer)
    """; # END PROMPT
```

```python
grader.check("q1c")
```

## **Question 2**: NumPy Array Indexing and Slicing

### Q2a: Stellar Magnitude Analysis
Astronomers have collected magnitude measurements (brightness) of stars in different wavelength bands. The data is stored as a 1D array where each element represents the magnitude of a star.

Given a stellar magnitude array, write a function `analyze_stellar_magnitudes(magnitudes, bright_threshold)` that:
- Takes a 1D NumPy array `magnitudes` and a float `bright_threshold`
- Extracts the first 3 magnitude values using slicing
- Extracts the last 2 magnitude values using negative indexing/slicing
- Finds all stars brighter than the threshold (magnitude < threshold, since lower magnitude = brighter star)
- Returns a tuple `(first_three, last_two, bright_stars)` 

Example: For magnitudes `[2.1, 4.5, 1.8, 6.2, 3.7, 5.1]` and threshold `4.0`, the function should return:
- First three: `[2.1, 4.5, 1.8]`
- Last two: `[3.7, 5.1]` 
- Bright stars: `[2.1, 1.8, 3.7]` (all values < 4.0)

```python
import numpy as np #so numpy is visible as np to otter grader

def analyze_stellar_magnitudes(magnitudes, bright_threshold):
    # BEGIN SOLUTION NO PROMPT
    first_three = magnitudes[:3]
    last_two = magnitudes[-2:]
    bright_stars = magnitudes[magnitudes < bright_threshold]
    return (first_three, last_two, bright_stars)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (first_three, last_two, bright_stars)
    """; # END PROMPT
```

```python
grader.check("q2a")
```

### Q2b: Galaxy Image Processing
Astronomers use 2D arrays to represent digital images of galaxies. Each element represents the brightness intensity at that pixel location.

Write a function `process_galaxy_image(image_data)` that takes a 2D NumPy array representing a galaxy image and performs the following operations:
- Extract the central 2x2 region (assuming the image is at least 4x4, and even side length)
- Extract the entire first row of the image
- Extract the entire last column of the image  
- Extract every other row (rows 0, 2, 4, etc.) using step slicing
- Extract every other column (columns 0, 2, 4, etc.) using step slicing
- Returns a tuple `(central_region, first_row, last_column, every_other_row, every_other_column)`

For a 4x4 image, the central 2x2 region would be the elements at rows 1-2 and columns 1-2.

Example: For a 4x4 image:
```
[[10, 20, 30, 40],
 [50, 60, 70, 80], 
 [90, 100, 110, 120],
 [130, 140, 150, 160]]
```
The central 2x2 region should be `[[60, 70], [100, 110]]`

Note: For an NxN image array, the shape of `every_other_row` should be `(N/2,N)` and the shape of `every_other_column` should be `(N, N/2)`

```python
import numpy as np #so numpy is visible as np to otter grader

def process_galaxy_image(image_data):
    # BEGIN SOLUTION NO PROMPT
    rows, cols = image_data.shape
    
    # Central 2x2 region (assuming image is at least 4x4)
    center_row_start = rows // 2 - 1
    center_col_start = cols // 2 - 1
    central_region = image_data[center_row_start:center_row_start+2, center_col_start:center_col_start+2]
    
    # First row
    first_row = image_data[0, :]
    
    # Last column  
    last_column = image_data[:, -1]
    
    # Every other row (0, 2, 4, ...)
    every_other_row = image_data[::2, :]
    
    # Every other column (0, 2, 4, ...)
    every_other_column = image_data[:, ::2]
    
    return (central_region, first_row, last_column, every_other_row, every_other_column)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (central_region, first_row, last_column, every_other_row, every_other_column)
    """; # END PROMPT
```

```python
grader.check("q2b")
```

## **Question 3**: Algorithm Time Complexity

For each of the following algorithm examples, analyze the time complexity and assign the correct Big O notation to the specified variable.

**Available time complexity options:**
- `"O(1)"` - Constant time
- `"O(log n)"` - Logarithmic time  
- `"O(n)"` - Linear time
- `"O(n log n)"` - Linearithmic time
- `"O(n^2)"` - Quadratic time
- `"O(n^3)"` - Cubic time
- `"O(2^n)"` - Exponential time

**Important:** Your answer must be one of the exact strings listed above.

### Q3a: Matrix Element Access
Consider this function that accesses a specific element in a matrix:

```python
def get_particle_energy(energy_matrix, particle_id):
    """
    Gets the energy of a specific particle from a pre-computed energy matrix.
    energy_matrix: 2D numpy array of size n x n
    particle_id: integer index of the particle
    """
    row = particle_id // len(energy_matrix)
    col = particle_id % len(energy_matrix)
    return energy_matrix[row][col]
```

Analyze the time complexity of this function with respect to the size of the energy matrix (n × n). 

Set the variable `q3a_complexity` to the correct time complexity string.

```python
# BEGIN SOLUTION NO PROMPT
q3a_complexity = "O(1)"
# END SOLUTION
""" # BEGIN PROMPT
# Analyze the time complexity and set q3a_complexity to the correct string
q3a_complexity = # Your answer here
"""; # END PROMPT
```

```python
grader.check("q3a")
```

### Q3b: Linear Search for Maximum Temperature
Consider this function that finds the maximum temperature in a list of measurements:

```python
def find_max_temperature(temperature_readings):
    """
    Finds the maximum temperature from a list of n temperature readings.
    temperature_readings: list of n floating point temperature values
    """
    max_temp = temperature_readings[0]
    for i in range(1, len(temperature_readings)):
        if temperature_readings[i] > max_temp:
            max_temp = temperature_readings[i]
    return max_temp
```

Analyze the time complexity of this function with respect to the number of temperature readings (n).

Set the variable `q3b_complexity` to the correct time complexity string.

```python
# BEGIN SOLUTION NO PROMPT
q3b_complexity = "O(n)"
# END SOLUTION
""" # BEGIN PROMPT
# Analyze the time complexity and set q3b_complexity to the correct string
q3b_complexity = # Your answer here
"""; # END PROMPT
```

```python
grader.check("q3b")
```

### Q3c: Nested Loop Matrix Multiplication
Consider this function that performs matrix multiplication for two matrices:

```python
def multiply_matrices(matrix_A, matrix_B):
    """
    Multiplies two n x n matrices using the standard algorithm.
    matrix_A: n x n matrix 
    matrix_B: n x n matrix
    Returns: n x n result matrix
    """
    n = len(matrix_A)
    result = [[0 for _ in range(n)] for _ in range(n)]
    
    for i in range(n):
        for j in range(n):
            for k in range(n):
                result[i][j] += matrix_A[i][k] * matrix_B[k][j]
    
    return result
```

Analyze the time complexity of this function with respect to the size of the matrices (n × n).

Set the variable `q3c_complexity` to the correct time complexity string.

```python
# BEGIN SOLUTION NO PROMPT
q3c_complexity = "O(n^3)"
# END SOLUTION
""" # BEGIN PROMPT
# Analyze the time complexity and set q3c_complexity to the correct string
q3c_complexity = # Your answer here
"""; # END PROMPT
```

```python
grader.check("q3c")
```

## **Question 4**: NumPy Broadcasting and Boolean Masking

In a particle physics experiment, researchers have collected energy measurements from different particle detectors. The data needs to be analyzed using NumPy broadcasting and boolean masking techniques.

[NumPy Broadcasting Documentation](https://numpy.org/doc/stable/user/basics.broadcasting.html)

Write a function `analyze_particle_energies(energies, backgrounds, threshold_energy, calibration_factors)` that:

**Parameters:**
- `energies`: 2D array of shape (n_detectors, n_measurements) containing raw energy readings
- `backgrounds`: 1D array of shape (n_detectors,) containing background energy levels for each detector  
- `threshold_energy`: scalar value representing the minimum energy to consider as a valid particle detection
- `calibration_factors`: 1D array of shape (n_detectors,) containing calibration factors for each detector

**Operations to perform:**
1. **Background Subtraction with Broadcasting**: Subtract the background energy from each detector's measurements using NumPy broadcasting
2. **Calibration with Broadcasting**: Multiply each detector's background-subtracted energies by its calibration factor using broadcasting
3. **Boolean Masking**: Create a boolean mask to identify measurements above the threshold energy (after subtracting background and applying calibration)
4. **Filtering**: Use the boolean mask to extract only the valid particle detections (above threshold)
5. **Statistics**: Calculate the mean energy of valid detections for each detector

**Returns:**
- A tuple `(corrected_energies, valid_detections, detector_means)` where:
  - `corrected_energies`: 2D array of background-subtracted and calibrated energies (same shape as input)
  - `valid_detections`: 1D array containing only the energy values above threshold (flattened)
  - `detector_means`: 1D array of mean energies for each detector (only counting valid detections) (if there are no valid detections for a detector, set the corresponding value to `np.nan`)

**Example:**
```python
energies = np.array([[100, 150, 80], [120, 200, 90]])  # 2 detectors, 3 measurements each
backgrounds = np.array([50, 60])  # background for each detector
threshold = 70
calibration = np.array([1.1, 1.2])  # calibration factor for each detector

# Expected process:
# 1. Subtract backgrounds: [[50, 100, 30], [60, 140, 30]]
# 2. Apply calibration: [[55, 110, 33], [72, 168, 36]]
# 3. Find valid (>70): [[False, True, False], [True, True, False]]
# 4. Extract valid values: [110, 72, 168]
# 5. Calculate means per detector: [110, 120] (detector 0: 110/1, detector 1: (72+168)/2)
```

Hint: You may want to use `np.newaxis`

```python
import numpy as np #so numpy is visible as np to otter grader

def analyze_particle_energies(energies, backgrounds, threshold_energy, calibration_factors):
    # BEGIN SOLUTION NO PROMPT
    # Step 1: Background subtraction using broadcasting
    # backgrounds is (n_detectors,) and energies is (n_detectors, n_measurements)
    # NumPy automatically broadcasts backgrounds to match energies shape
    background_subtracted = energies - backgrounds[:, np.newaxis]
    
    # Step 2: Apply calibration factors using broadcasting
    corrected_energies = background_subtracted * calibration_factors[:, np.newaxis]
    
    # Step 3: Create boolean mask for values above threshold
    valid_mask = corrected_energies > threshold_energy
    
    # Step 4: Extract valid detections (flatten and filter)
    valid_detections = corrected_energies[valid_mask]
    
    # Step 5: Calculate mean for each detector (only counting valid detections)
    detector_means = []
    for i in range(len(calibration_factors)):
        detector_valid = corrected_energies[i][valid_mask[i]]
        if len(detector_valid) > 0:
            detector_means.append(np.mean(detector_valid))
        else:
            detector_means.append(np.nan)  # No valid detections for this detector
    
    detector_means = np.array(detector_means)
    
    return (corrected_energies, valid_detections, detector_means)
    # END SOLUTION
    """ # BEGIN PROMPT
    # Write your code here!
    return (corrected_energies, valid_detections, detector_means)
    """; # END PROMPT
```

```python
grader.check("q4")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
# BEGIN SOLUTION NO PROMPT
# END SOLUTION
""" # BEGIN PROMPT
"""
# write ai disclosure here:

"""
"""; # END PROMPT
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> temp_array = create_temperature_array([20.5, 25.3, 30.1])
>>> isinstance(temp_array, np.ndarray) and np.allclose(temp_array, [20.5, 25.3, 30.1])
True
```
```
>>> empty_array = create_temperature_array([])
>>> isinstance(empty_array, np.ndarray) and len(empty_array) == 0
True
```
```
>>> single_array = create_temperature_array([23.7])
>>> isinstance(single_array, np.ndarray) and np.allclose(single_array, [23.7])
True
```


Tests for q1b (points: None)
```
>>> time_arr, freq_arr = create_wave_arrays(0, 1, 5, 1, 1000, 4)
>>> np.allclose(time_arr, [0, 0.25, 0.5, 0.75, 1.0])
True
```
```
>>> time_arr, freq_arr = create_wave_arrays(0, 1, 5, 1, 1000, 4)
>>> np.allclose(freq_arr, [1, 10, 100, 1000])
True
```
```
>>> time_arr2, freq_arr2 = create_wave_arrays(-1, 1, 3, 100, 10000, 3)
>>> len(time_arr2) == 3 and len(freq_arr2) == 3 and np.allclose(time_arr2, [-1, 0, 1])
True
```


Tests for q1c (points: None)
```
>>> ref_1d = np.array([1, 2, 3])
>>> zeros_arr, voltage_arr, buffer = initialize_experiment_arrays(ref_1d, 5.0)
>>> zeros_arr.shape == (3,) and np.allclose(zeros_arr, [0, 0, 0]) and np.allclose(voltage_arr, [5.0, 5.0, 5.0])
True
```
```
>>> ref_2d = np.array([[1, 2], [3, 4]])
>>> zeros_arr, voltage_arr, buffer = initialize_experiment_arrays(ref_2d, -2.5)
>>> zeros_arr.shape == (2, 2) and np.allclose(voltage_arr, [[-2.5, -2.5], [-2.5, -2.5]])
True
```
```
>>> ref_any = np.array([1, 2, 3, 4, 5])
>>> zeros_arr, voltage_arr, buffer = initialize_experiment_arrays(ref_any, 0)
>>> buffer.shape == (10,) and np.allclose(buffer, np.zeros(10))
True
```


Tests for q2a (points: None)
```
>>> mags = np.array([2.1, 4.5, 1.8, 6.2, 3.7, 5.1])
>>> first, last, bright = analyze_stellar_magnitudes(mags, 4.0)
>>> np.allclose(first, [2.1, 4.5, 1.8]) and np.allclose(last, [3.7, 5.1])
True
```
```
>>> mags = np.array([2.1, 4.5, 1.8, 6.2, 3.7, 5.1])
>>> first, last, bright = analyze_stellar_magnitudes(mags, 4.0)
>>> np.allclose(bright, [2.1, 1.8, 3.7])
True
```
```
>>> mags2 = np.array([1.5, 3.2, 5.7, 2.8, 4.1])
>>> first, last, bright = analyze_stellar_magnitudes(mags2, 3.0)
>>> len(bright) == 2 and np.allclose(bright, [1.5, 2.8])
True
```


Tests for q2b (points: None)
```
>>> image_4x4 = np.array([[10, 20, 30, 40], [50, 60, 70, 80], [90, 100, 110, 120], [130, 140, 150, 160]])
>>> central, first_row, last_col, every_row, every_col = process_galaxy_image(image_4x4)
>>> np.array_equal(central, [[60, 70], [100, 110]])
True
```
```
>>> image_4x4 = np.array([[10, 20, 30, 40], [50, 60, 70, 80], [90, 100, 110, 120], [130, 140, 150, 160]])
>>> central, first_row, last_col, every_row, every_col = process_galaxy_image(image_4x4)
>>> np.array_equal(first_row, [10, 20, 30, 40]) and np.array_equal(last_col, [40, 80, 120, 160])
True
```
```
>>> image_4x4 = np.array([[10, 20, 30, 40], [50, 60, 70, 80], [90, 100, 110, 120], [130, 140, 150, 160]])
>>> central, first_row, last_col, every_row, every_col = process_galaxy_image(image_4x4)
>>> expected_every_row = [[10, 20, 30, 40], [90, 100, 110, 120]]
>>> expected_every_col = [[10, 30], [50, 70], [90, 110], [130, 150]]
>>> np.array_equal(every_row, expected_every_row) and np.array_equal(every_col, expected_every_col)
True
```
```
>>> image_6x6 = np.arange(36).reshape(6, 6)
>>> central, first_row, last_col, every_row, every_col = process_galaxy_image(image_6x6)
>>> central.shape == (2, 2) and every_row.shape == (3, 6) and (every_col.shape == (6, 3))
True
```


Tests for q3a (points: None)
```
>>> q3a_complexity == 'O(1)'
True
```


Tests for q3b (points: None)
```
>>> q3b_complexity == 'O(n)'
True
```


Tests for q3c (points: None)
```
>>> q3c_complexity == 'O(n^3)'
True
```


Tests for q4 (points: None)
```
>>> energies = np.array([[10, 15, 8, 12], [20, 18, 6, 14], [16, 22, 19, 11]])
>>> backgrounds = np.array([2, 3, 1])
>>> threshold_energy = 10
>>> calibration_factors = np.array([1.0, 1.2, 0.9])
>>> corrected, valid, means = analyze_particle_energies(energies, backgrounds, threshold_energy, calibration_factors)
>>> expected_corrected = np.array([[8.0, 13.0, 6.0, 10.0], [20.4, 18.0, 3.6, 13.2], [13.5, 18.9, 16.2, 9.0]])
>>> np.allclose(corrected, expected_corrected)
True
```
```
>>> energies = np.array([[10, 15, 8, 12], [20, 18, 6, 14], [16, 22, 19, 11]])
>>> backgrounds = np.array([2, 3, 1])
>>> threshold_energy = 10
>>> calibration_factors = np.array([1.0, 1.2, 0.9])
>>> corrected, valid, means = analyze_particle_energies(energies, backgrounds, threshold_energy, calibration_factors)
>>> expected_valid = np.array([13.0, 20.4, 18.0, 13.2, 13.5, 18.9, 16.2])
>>> len(valid) == 7 and np.allclose(np.sort(valid), np.sort(expected_valid))
True
```
```
>>> energies = np.array([[10, 15, 8, 12], [20, 18, 6, 14], [16, 22, 19, 11]])
>>> backgrounds = np.array([2, 3, 1])
>>> threshold_energy = 10
>>> calibration_factors = np.array([1.0, 1.2, 0.9])
>>> corrected, valid, means = analyze_particle_energies(energies, backgrounds, threshold_energy, calibration_factors)
>>> expected_means = np.array([13.0, 17.2, 16.2])
>>> np.allclose(means, expected_means)
True
```
```
>>> energies_edge = np.array([[5, 3, 2, 4], [15, 18, 12, 20], [8, 6, 7, 5]])
>>> backgrounds_edge = np.array([1, 2, 3])
>>> threshold_edge = 12
>>> calibration_edge = np.array([1.5, 1.0, 2.0])
>>> corrected_edge, valid_edge, means_edge = analyze_particle_energies(energies_edge, backgrounds_edge, threshold_edge, calibration_edge)
>>> expected_corrected_edge = np.array([[6.0, 3.0, 1.5, 4.5], [13.0, 16.0, 10.0, 18.0], [10.0, 6.0, 8.0, 4.0]])
>>> bool(np.allclose(corrected_edge, expected_corrected_edge) and len(valid_edge) == 3 and np.isnan(means_edge[0]) and np.isnan(means_edge[2]))
True
```
```
>>> energies_broadcast = np.array([[100, 200], [300, 400]])
>>> backgrounds_broadcast = np.array([10, 20])
>>> threshold_broadcast = 150
>>> calibration_broadcast = np.array([0.5, 2.0])
>>> corrected_broadcast, valid_broadcast, means_broadcast = analyze_particle_energies(energies_broadcast, backgrounds_broadcast, threshold_broadcast, calibration_broadcast)
>>> expected_corrected_broadcast = np.array([[45.0, 95.0], [560.0, 760.0]])
>>> expected_valid_broadcast = np.array([560.0, 760.0])
>>> np.allclose(corrected_broadcast, expected_corrected_broadcast) and len(valid_broadcast) == 2 and np.allclose(valid_broadcast, expected_valid_broadcast)
True
```
