```python
# Initialize Otter
import otter
grader = otter.Notebook("hw7.ipynb")
```

```python
rng_seed = 73
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import pandas as pd
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

## **Question 1**: Central Limit Theorem Demonstration

### Part a: Uniform Distribution Sampling

Write a function `sample_uniform_mean(a, b, n_samples)` that samples from a uniform distribution and returns the mean of the samples.

[numpy uniform dist. docs](https://numpy.org/doc/stable/reference/random/generated/numpy.random.uniform.html)

**Requirements:**
- Use `np.random` methods to generate samples
- Return the mean of the samples

**Parameters:**
- `a`: float, lower bound of uniform distribution
- `b`: float, upper bound of uniform distribution  
- `n_samples`: int, number of samples to draw

**Returns:**
- `sample_mean`: float, mean of the samples

```python
def sample_uniform_mean(a, b, n_samples):
    # Write your code here!
    
    return sample_mean
```

```python
grader.check("q1a")
```

### Part b: Gamma Distribution Sampling

Write a function `sample_gamma_mean(shape, scale, n_samples)` that samples from a gamma distribution and returns the mean of the samples.

[numpy gamma dist. docs](https://numpy.org/doc/stable/reference/random/generated/numpy.random.gamma.html)

**Requirements:**
- Use `np.random` to generate samples
- Return the mean of the samples

**Parameters:**
- `shape`: float, shape parameter of gamma distribution
- `scale`: float, scale parameter of gamma distribution
- `n_samples`: int, number of samples to draw

**Returns:**
- `sample_mean`: float, mean of the samples

```python
def sample_gamma_mean(shape, scale, n_samples):
    # Write your code here!
    return sample_mean
```

```python
grader.check("q1b")
```

### Part c: Central Limit Theorem Visualization

Write a function `plot_sample_means(use_uniform, N, n_samples=100, bins=30, show_plot=False)` that demonstrates the Central Limit Theorem by plotting a histogram of sample means.

**Requirements:**
- If `use_uniform=True`, use `sample_uniform_mean(0, 10, n_samples)` for each trial
- If `use_uniform=False`, use `sample_gamma_mean(2, 2, n_samples)` for each trial
- Perform N trials to collect N sample means
- Create a histogram of the sample means with specified number of bins
- Set x-axis label to "Sample Mean"
- Set y-axis label to "Density"  
- Set title to include "Central Limit Theorem" and distribution information
- Add a grid to the plot
- The function should return the matplotlib figure object

**Parameters:**
- `use_uniform`: bool, if True use uniform distribution, else use gamma
- `N`: int, number of trials (number of sample means to collect)
- `n_samples`: int, default 100, number of samples per trial
- `bins`: int, default 30, number of histogram bins
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object

```python
def plot_sample_means(use_uniform, N, n_samples=100, bins=30, show_plot=False):
    #compute an array of sample means
    
    ...

    # Create histogram plot
    fig, ax = plt.subplots(figsize=(10, 6))


    ...
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Test the function and demonstrate CLT convergence with different N values
np.random.seed(rng_seed)  # Set seed for reproducible results

# Test individual functions first
print("Testing individual functions:")
uniform_mean = sample_uniform_mean(0, 10, 1000)
gamma_mean = sample_gamma_mean(2, 2, 1000)
print(f"Sample mean from Uniform(0,10): {uniform_mean:.3f} (theoretical: 5.0)")
print(f"Sample mean from Gamma(2,2): {gamma_mean:.3f} (theoretical: 4.0)")

# Demonstrate CLT convergence with different N values
print("\nDemonstrating Central Limit Theorem convergence:")

# Small N - should show some deviation from normal
fig1 = plot_sample_means(use_uniform=True, N=50, n_samples=100, show_plot=True)
print("Uniform distribution with N=50 trials")

# Medium N - better convergence
fig2 = plot_sample_means(use_uniform=True, N=200, n_samples=100, show_plot=True) 
print("Uniform distribution with N=200 trials")

# Large N - clear normal distribution
fig3 = plot_sample_means(use_uniform=True, N=1000, n_samples=100, show_plot=True)
print("Uniform distribution with N=1000 trials")

# Compare gamma distribution
fig4 = plot_sample_means(use_uniform=False, N=500, n_samples=100, show_plot=True)
print("Gamma distribution with N=500 trials")

# Show effect of sample size per trial
fig5 = plot_sample_means(use_uniform=True, N=500, n_samples=25, show_plot=True)
print("Uniform distribution with N=500 trials, n_samples=25 (wider distribution)")

fig6 = plot_sample_means(use_uniform=True, N=500, n_samples=400, show_plot=True)
print("Uniform distribution with N=500 trials, n_samples=400 (narrower distribution)")
```

```python
grader.check("q1c")
```

## **Question 2**: Error Propagation

When we have a function $y = g(x_1, x_2, \ldots, x_n)$ where each parameter $x_i$ has an associated uncertainty $\sigma_i$, the propagated uncertainty in $y$ is given by:

$$\sigma_y = \sqrt{\sum_{i=1}^{n} \left(\frac{\partial g}{\partial x_i}\right)^2 \sigma_i^2}$$

This formula assumes that the uncertainties in the input parameters are uncorrelated (independent). Each partial derivative $\frac{\partial g}{\partial x_i}$ is evaluated at the measured values of all parameters.

In this question, you will implement error propagation for two different physics formulas by:
1. Computing the partial derivatives analytically
2. Implementing functions that return both the calculated value and its propagated uncertainty as a tuple `(y, sigma_y)`

### Part a: Projectile Motion Range

The range of a projectile launched at angle $\theta$ with initial velocity $v_0$ is given by:

$$R = \frac{v_0^2 \sin(2\theta)}{g}$$

where $g = 9.81$ m/s² is the acceleration due to gravity.

Write a function `projectile_range_with_error(v0, sigma_v0, theta, sigma_theta, g=9.81)` that calculates both the range and its propagated uncertainty.

**Requirements:**
- Use the error propagation formula with partial derivatives
- `theta` should be in radians
- Return a tuple `(R, sigma_R)` where `R` is the range and `sigma_R` is the propagated uncertainty

**Hint:** You'll need to compute:
$\frac{\partial R}{\partial v_0}$, $\frac{\partial R}{\partial \theta}$
and then implement the analytic expression you derive into your code

**Parameters:**
- `v0`: float, initial velocity in m/s
- `sigma_v0`: float, uncertainty in initial velocity  
- `theta`: float, launch angle in radians
- `sigma_theta`: float, uncertainty in launch angle in radians
- `g`: float, default 9.81, acceleration due to gravity in m/s²

**Returns:**
- `(R, sigma_R)`: tuple of floats, range and its uncertainty in meters

```python
def projectile_range_with_error(v0, sigma_v0, theta, sigma_theta, g=9.81):
    # Calculate the range R
    R = ...
    
    # Calculate partial derivatives
    dR_dv0 = ...
    dR_dtheta = ...
    
    # Apply error propagation formula
    sigma_R = ...
    
    return (R, sigma_R)
```

```python
grader.check("q2a")
```

### Part b: Physical Pendulum Period

The period of a physical pendulum (a rigid body oscillating about a pivot point) is given by:

$$T = 2\pi \sqrt{\frac{I}{mgd}}$$

where:
- $I$ is the moment of inertia about the pivot point
- $m$ is the mass of the pendulum
- $g$ is the acceleration due to gravity  
- $d$ is the distance from the pivot to the center of mass

Write a function `pendulum_period_with_error(I, sigma_I, m, sigma_m, d, sigma_d, g=9.81)` that calculates both the period and its propagated uncertainty.

**Requirements:**
- Use the error propagation formula with partial derivatives
- Return a tuple `(T, sigma_T)` where `T` is the period and `sigma_T` is the propagated uncertainty

**Hint:** You'll need to compute:
$\frac{\partial T}{\partial I}$, $\frac{\partial T}{\partial m}$, $\frac{\partial T}{\partial d}$

**Parameters:**
- `I`: float, moment of inertia in kg⋅m²
- `sigma_I`: float, uncertainty in moment of inertia
- `m`: float, mass in kg
- `sigma_m`: float, uncertainty in mass
- `d`: float, distance to center of mass in meters
- `sigma_d`: float, uncertainty in distance
- `g`: float, default 9.81, acceleration due to gravity in m/s²

**Returns:**
- `(T, sigma_T)`: tuple of floats, period and its uncertainty in seconds

```python
def pendulum_period_with_error(I, sigma_I, m, sigma_m, d, sigma_d, g=9.81):
    # Calculate the period T
    T = ...
    
    # Calculate partial derivatives
    dT_dI = ...
    dT_dm = ...
    dT_dd = ...
    
    # Apply error propagation formula
    sigma_T = ...
    
    return (T, sigma_T)
```

```python
grader.check("q2b")
```

## **Question 3**: Covariance and Correlation Analysis

Covariance and correlation are fundamental measures of how two variables change together. For a dataset with multiple variables, we can compute:

**Covariance Matrix**: $\text{Cov}(X_i, X_j) = \frac{1}{n-1} \sum_{k=1}^{n} (X_{i,k} - \bar{X_i})(X_{j,k} - \bar{X_j})$

**Correlation Matrix**: $\text{Corr}(X_i, X_j) = \frac{\text{Cov}(X_i, X_j)}{\sigma_{X_i} \sigma_{X_j}}$

Where correlation values range from -1 (perfect negative correlation) to +1 (perfect positive correlation), with 0 indicating no linear correlation.

In this question, you will analyze a physics dataset containing 7 variables that represent related physical quantities. Some variables should show strong correlations based on physical laws, while others should be relatively independent.

### Part a: Computing Covariance and Correlation Matrices

Write a function `analyze_covariance(filename)` that reads a CSV data file and computes both the covariance and correlation matrices.

**Requirements:**
- Read the CSV file using pandas (`pd.read_csv`)
- Compute the covariance matrix using numpy (sample covariance with n-1 denominator) (you can also do this by hand, but numpy is much easier/faster)
- Compute the correlation matrix using numpy
- Return both matrices as numpy arrays

**Parameters:**
- `filename`: str, csv filename (`physics_data.csv`)

**Returns:**
- `(cov_matrix, corr_matrix)`: tuple of numpy arrays
  - `cov_matrix`: covariance matrix (7x7)
  - `corr_matrix`: correlation matrix (7x7)

**Note:** The correlation matrix should have 1.0 on the diagonal and values between -1 and 1 elsewhere.

```python
def analyze_covariance(filename):
    # Read the CSV file
    data = ...
    
    # Convert to numpy array
    data_array = ...
    
    # Compute covariance matrix
    cov_matrix = ...
    
    # Compute correlation matrix
    corr_matrix = ...
    
    return cov_matrix, corr_matrix
```

```python
grader.check("q3a")
```

### Part b: Visualizing Covariance and Correlation Matrices

Write a function `plot_covariance_matrices(filename, show_plot=False)` that creates heatmap visualizations of both the covariance and correlation matrices.

**Requirements:**
- Read the data and compute matrices using your `analyze_covariance` function
- Create a figure with two subplots (side by side)
- Left subplot: covariance matrix heatmap
- Right subplot: correlation matrix heatmap (colorbar range should be `[-1,1]`)
- Use appropriate colormaps and include colorbars
- Set titles and axis labels with variable names
- Return the matplotlib figure object

**Parameters:**
- `filename`: str, path to the CSV file containing the data
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object containing both heatmaps

**Note:** For better visualization, consider using different colormaps for covariance vs correlation matrices.

Think about what the correlation matrix says about these data, there's no deliverable on this, but what insight can be gained by examining such a plot?

```python
def plot_covariance_matrices(filename, show_plot=False):
    # Get matrices from part a
    ...
    
    # Read data to get variable names
    data = pd.read_csv(filename)
    variable_names = ...
    
    # Create figure with subplots
    fig, (ax1, ax2) = plt.subplots(1, 2, figsize=(15, 6))
    
    # Create covariance heatmap
    im1 = ax1.imshow(...)
    ax1.set_title('Covariance Matrix')
    # ... add labels and colorbar
    
    # Create correlation heatmap  
    im2 = ax2.imshow(...)
    ax2.set_title('Correlation Matrix')
    # ... add labels and colorbar
    
    plt.tight_layout()
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
#showing the plot
plot_covariance_matrices("physics_data.csv", show_plot=True)
```

```python
grader.check("q3b")
```

## **Question 4**: Poisson Likelihood Estimation

The Poisson distribution is fundamental in physics for describing the probability of observing a discrete number of events (like photon counts, radioactive decays, or particle detections) in a fixed interval when events occur independently at a constant average rate.

The **Poisson probability mass function** for observing exactly $k$ events when the expected number is $\lambda$ is:

$$P(k|\lambda) = \frac{\lambda^k e^{-\lambda}}{k!}$$

When we have a dataset of $n$ independent observations $\{k_1, k_2, ..., k_n\}$, the **likelihood function** represents how probable it is to observe this specific dataset given a particular value of $\lambda$:

$$L(\lambda) = \prod_{i=1}^{n} P(k_i|\lambda) = \prod_{i=1}^{n} \frac{\lambda^{k_i} e^{-\lambda}}{k_i!}$$

For computational convenience, we often work with the **log-likelihood**:

$$\ln L(\lambda) = \sum_{i=1}^{n} \ln P(k_i|\lambda) = \sum_{i=1}^{n} \left[ k_i \ln(\lambda) - \lambda - \ln(k_i!) \right]$$

The **maximum likelihood estimate (MLE)** is the value of $\lambda$ that maximizes the likelihood function. For the Poisson distribution, the MLE has the analytical solution: $\hat{\lambda} = \bar{k} = \frac{1}{n}\sum_{i=1}^{n} k_i$ (the sample mean).

### Part a: Computing the Poisson Likelihood

Write a function `poisson_likelihood(data, lambda_val)` that computes the likelihood of observing the given data for a specific value of λ.

**Mathematical Steps:**
1. For each data point $k_i$, compute $P(k_i|\lambda) = \frac{\lambda^{k_i} e^{-\lambda}}{k_i!}$
2. Multiply all individual probabilities: $L(\lambda) = \prod_{i=1}^{n} P(k_i|\lambda)$

**Requirements:**
- Use the Poisson probability mass function to compute individual probabilities
- Return the product of all probabilities (the likelihood)
- Handle potential numerical underflow by using appropriate functions
- You may use `scipy.stats.poisson.pmf()` or implement the formula directly

**Parameters:**
- `data`: array-like, observed count data (integers)
- `lambda_val`: float, the Poisson parameter λ

**Returns:**
- `likelihood`: float, the likelihood L(λ) of observing the data given λ

**Hint:** For large datasets, likelihoods can become extremely small. Consider using `np.prod()` or working with log-likelihood and then exponentiating if needed.

```python
def poisson_likelihood(data, lambda_val):
    from scipy.stats import poisson
    
    # Convert data to numpy array
    data = np.array(data)
    
    # Compute individual probabilities
    probabilities = ...
    
    # Compute likelihood as product of probabilities
    likelihood = ...
    
    return likelihood
```

```python
grader.check("q4a")
```

### Part b: Likelihood Curve and Maximum Likelihood Estimation

Write a function `plot_likelihood_curve(data, lambda_range=None, n_points=100, show_plot=False)` that plots the likelihood function as a function of λ and determines the maximum likelihood estimate with confidence bounds.

**Mathematical Steps:**

1. **Create λ range**: If not provided, use a range around the sample mean
2. **Calculate likelihood curve**: For each λ value, compute L(λ) using your function from part (a)
3. **Find MLE**: The λ value that maximizes the likelihood
4. **Determine confidence bounds**: 
   - Find the maximum likelihood: $L_{\max} = L(\hat{\lambda})$
   - The 1σ confidence interval corresponds to $L(\lambda) = L_{\max} \times e^{-0.5}$ (where log-likelihood drops by 0.5)
   - Find the λ values where the likelihood curve crosses this threshold

**Requirements:**
- Use your `poisson_likelihood` function from part (a) to compute likelihoods
- Create a smooth curve by evaluating many λ values
- Plot the likelihood vs λ with appropriate labels and formatting
- Mark the MLE and confidence bounds on the plot
- Return the figure object and the three key values

**Parameters:**
- `data`: array-like, observed count data
- `lambda_range`: tuple (min, max), default None. If None, use (0.5*mean, 2*mean)
- `n_points`: int, default 100, number of λ values to evaluate
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `(fig, mle, lower_bound, upper_bound)`: tuple containing:
  - `fig`: matplotlib figure object
  - `mle`: float, maximum likelihood estimate of λ
  - `lower_bound`: float, 1σ lower confidence bound
  - `upper_bound`: float, 1σ upper confidence bound

**Note:** For computational stability with small likelihoods, you may want to work with log-likelihood internally and convert back for plotting and analysis.

```python
def plot_likelihood_curve(data, lambda_range=None, n_points=100, show_plot=False):
    data = np.array(data)
    sample_mean = np.mean(data)
    
    # Set lambda range if not provided
    if lambda_range is None:
        lambda_range = ...
    
    # Create array of lambda values
    lambda_vals = ...
    
    # Calculate likelihood for each lambda value
    likelihoods = ...
    
    # Find MLE
    mle = ...
    
    # Find confidence bounds (likelihood drops by e^(-0.5))
    threshold = ...
    lower_bound = ...
    upper_bound = ...
    
    # Create plot
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Plot likelihood curve and mark MLE and bounds
    ...
    
    if show_plot:
        plt.show()
    
    return fig, mle, lower_bound, upper_bound
```

```python
# Generate sample Poisson data for testing
np.random.seed(rng_seed)

# Create a dataset with known Poisson parameter (lambda = 4.2)
true_lambda = 4.2
n_observations = 50
poisson_data = np.random.poisson(true_lambda, n_observations)

print(f"Generated {n_observations} observations from Poisson(λ={true_lambda})")
print(f"Sample data: {poisson_data[:10]}...")  # Show first 10 values
print(f"Sample mean: {np.mean(poisson_data):.3f}")
print(f"Sample standard deviation: {np.std(poisson_data):.3f}")
print(f"Theoretical mean: {true_lambda}")
print(f"Theoretical std: {np.sqrt(true_lambda):.3f}")

# Test the likelihood function with a few values
print(f"\nTesting likelihood function:")
test_lambdas = [3.0, 4.0, 5.0]
for lam in test_lambdas:
    likelihood = poisson_likelihood(poisson_data, lam)
    print(f"L(λ={lam}) = {likelihood:.2e}")

# Plot the likelihood curve
fig, mle, lower, upper = plot_likelihood_curve(poisson_data, show_plot=True)

print(f"\nResults:")
print(f"Maximum Likelihood Estimate: {mle:.3f}")
print(f"1σ Confidence Interval: [{lower:.3f}, {upper:.3f}]")
print(f"True value λ={true_lambda} is {'within' if lower <= true_lambda <= upper else 'outside'} the confidence interval")
```

```python
grader.check("q4b")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> np.random.seed(rng_seed)
>>> result = sample_uniform_mean(0, 10, 1000)
>>> is_number = isinstance(result, (int, float, np.floating))
>>> in_range = 0 <= result <= 10
>>> reasonable_mean = 3 <= result <= 7
>>> (is_number, bool(in_range), bool(reasonable_mean))
(True, True, True)
```


Tests for q1b (points: None)
```
>>> np.random.seed(rng_seed)
>>> result = sample_gamma_mean(2, 2, 1000)
>>> is_number = isinstance(result, (int, float, np.floating))
>>> positive = result > 0
>>> reasonable_mean = 2 <= result <= 6
>>> (is_number, bool(positive), bool(reasonable_mean))
(True, True, True)
```


Tests for q1c (points: None)
```
>>> np.random.seed(42)
>>> fig = plot_sample_means(use_uniform=True, N=50, n_samples=100, show_plot=False)
>>> import matplotlib.figure
>>> test_result = isinstance(fig, matplotlib.figure.Figure)
>>> plt.close(fig)
>>> test_result
True
```
```
>>> np.random.seed(rng_seed)
>>> fig = plot_sample_means(use_uniform=True, N=100, n_samples=50, show_plot=False)
>>> ax = fig.axes[0]
>>> has_histogram = len(ax.patches) > 0
>>> has_xlabel = ax.get_xlabel() == 'Sample Mean'
>>> has_ylabel = ax.get_ylabel() == 'Density'
>>> title = ax.get_title()
>>> has_clt_in_title = 'Central Limit Theorem' in title
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> plt.close(fig)
>>> (has_histogram, has_xlabel, has_ylabel, has_clt_in_title, has_grid)
(True, True, True, True, True)
```


Tests for q2a (points: None)
```
>>> result = projectile_range_with_error(50, 2, np.pi / 4, 0.05)
>>> is_tuple = isinstance(result, tuple)
>>> has_two_elements = len(result) == 2
>>> both_numbers = all((isinstance(x, (int, float, np.floating)) for x in result))
>>> R, sigma_R = result
>>> positive_values = R > 0 and sigma_R > 0
>>> (is_tuple, has_two_elements, both_numbers, bool(positive_values))
(True, True, True, True)
```
```
>>> np.random.seed(rng_seed)
>>> R_45, _ = projectile_range_with_error(40, 1, np.pi / 4, 0.01)
>>> R_30, _ = projectile_range_with_error(40, 1, np.pi / 6, 0.01)
>>> R_60, _ = projectile_range_with_error(40, 1, np.pi / 3, 0.01)
>>> max_at_45 = R_45 > R_30 and R_45 > R_60
>>> reasonable_range = 150 < R_45 < 170
>>> (bool(max_at_45), bool(reasonable_range))
(True, True)
```


Tests for q2b (points: None)
```
>>> result = pendulum_period_with_error(0.1, 0.005, 2.0, 0.05, 0.5, 0.01)
>>> is_tuple = isinstance(result, tuple)
>>> has_two_elements = len(result) == 2
>>> both_numbers = all((isinstance(x, (int, float, np.floating)) for x in result))
>>> T, sigma_T = result
>>> positive_values = T > 0 and sigma_T > 0
>>> (is_tuple, has_two_elements, both_numbers, bool(positive_values))
(True, True, True, True)
```
```
>>> T1, _ = pendulum_period_with_error(0.1, 0.001, 1.0, 0.01, 0.3, 0.005)
>>> T2, _ = pendulum_period_with_error(0.2, 0.001, 1.0, 0.01, 0.3, 0.005)
>>> T3, _ = pendulum_period_with_error(0.1, 0.001, 2.0, 0.01, 0.3, 0.005)
>>> T4, _ = pendulum_period_with_error(0.1, 0.001, 1.0, 0.01, 0.6, 0.005)
>>> larger_I_increases_T = T2 > T1
>>> larger_m_decreases_T = T3 < T1
>>> larger_d_decreases_T = T4 < T1
>>> reasonable_periods = all((0.5 < T < 5.0 for T in [T1, T2, T3, T4]))
>>> (bool(larger_I_increases_T), bool(larger_m_decreases_T), bool(larger_d_decreases_T), reasonable_periods)
(True, True, True, True)
```


Tests for q3a (points: None)
```
>>> result = analyze_covariance('physics_data.csv')
>>> is_tuple = isinstance(result, tuple)
>>> has_two_elements = len(result) == 2
>>> cov_matrix, corr_matrix = result
>>> both_numpy = isinstance(cov_matrix, np.ndarray) and isinstance(corr_matrix, np.ndarray)
>>> correct_shape = cov_matrix.shape == (7, 7) and corr_matrix.shape == (7, 7)
>>> (is_tuple, has_two_elements, both_numpy, correct_shape)
(True, True, True, True)
```
```
>>> cov_matrix, corr_matrix = analyze_covariance('physics_data.csv')
>>> diagonal_ones = np.allclose(np.diag(corr_matrix), 1.0)
>>> valid_range = np.all((-1 <= corr_matrix) & (corr_matrix <= 1))
>>> is_symmetric = np.allclose(corr_matrix, corr_matrix.T)
>>> has_correlations = np.any(np.abs(corr_matrix - np.eye(7)) > 0.3)
>>> (diagonal_ones, bool(valid_range), is_symmetric, bool(has_correlations))
(True, True, True, True)
```


Tests for q3b (points: None)
```
>>> fig = plot_covariance_matrices('physics_data.csv', show_plot=False)
>>> is_figure = isinstance(fig, matplotlib.figure.Figure)
>>> has_two_subplots = len(fig.axes) == 4
>>> plt.close(fig)
>>> (is_figure, has_two_subplots)
(True, True)
```
```
>>> fig = plot_covariance_matrices('physics_data.csv', show_plot=False)
>>> ax1, ax2 = (fig.axes[0], fig.axes[1])
>>> has_cov_title = 'Covariance' in ax1.get_title()
>>> has_corr_title = 'Correlation' in ax2.get_title()
>>> has_cov_image = len(ax1.images) > 0
>>> has_corr_image = len(ax2.images) > 0
>>> has_labels = ax1.get_xlabel() != '' and ax2.get_xlabel() != ''
>>> plt.close(fig)
>>> (has_cov_title, has_corr_title, has_cov_image, has_corr_image, has_labels)
(True, True, True, True, True)
```


Tests for q4a (points: None)
```
>>> np.random.seed(rng_seed)
>>> test_data = [2, 3, 1, 4, 2]
>>> result = poisson_likelihood(test_data, 2.5)
>>> is_number = isinstance(result, (int, float, np.floating))
>>> is_positive = result > 0
>>> reasonable_magnitude = 0 < result < 1
>>> (is_number, bool(is_positive), bool(reasonable_magnitude))
(True, True, True)
```
```
>>> np.random.seed(rng_seed)
>>> test_data = [3, 2, 4, 3, 2, 3, 4, 1, 3, 5]
>>> likelihood_3 = poisson_likelihood(test_data, 3.0)
>>> likelihood_1 = poisson_likelihood(test_data, 1.0)
>>> likelihood_6 = poisson_likelihood(test_data, 6.0)
>>> better_than_1 = likelihood_3 > likelihood_1
>>> better_than_6 = likelihood_3 > likelihood_6
>>> bool(better_than_1 and better_than_6)
True
```


Tests for q4b (points: None)
```
>>> np.random.seed(rng_seed)
>>> test_data = np.random.poisson(3.5, 30)
>>> result = plot_likelihood_curve(test_data, show_plot=False)
>>> is_tuple = isinstance(result, tuple)
>>> has_four_elements = len(result) == 4
>>> fig, mle, lower, upper = result
>>> is_figure = isinstance(fig, matplotlib.figure.Figure)
>>> all_numbers = all((isinstance(x, (int, float, np.floating)) for x in [mle, lower, upper]))
>>> plt.close(fig)
>>> (is_tuple, has_four_elements, is_figure, all_numbers)
(True, True, True, True)
```
```
>>> np.random.seed(rng_seed)
>>> test_data = np.random.poisson(4.0, 50)
>>> fig, mle, lower, upper = result = plot_likelihood_curve(test_data, show_plot=False)
>>> sample_mean = np.mean(test_data)
>>> mle_close_to_mean = abs(mle - sample_mean) < 0.5
>>> bounds_bracket_mle = lower < mle < upper
>>> plt.close(fig)
>>> (bool(mle_close_to_mean), bool(bounds_bracket_mle))
(True, True)
```
