```python
# Initialize Otter
import otter
grader = otter.Notebook("hw6.ipynb")
```

```python
rng_seed = 63
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
import scipy as sp
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

## **Question 1**: Parameter Estimation and Visualization

Write a function `analyze_normal_dataset(data, bins=30, show_plot=False)` that analyzes a dataset assumed to be drawn from a normal distribution.

Your function should:
1. Compute the sample mean and sample variance from the input data
2. Create a histogram of the data samples
3. Overlay the theoretical normal distribution PDF using the computed mean and variance
4. Display the computed statistics on the plot

**Requirements:**
- Compute sample mean using `np.mean(data)`
- Compute sample variance using `np.var(data, ddof=1)` (use Bessel's correction with ddof=1)
- Create histogram using `ax.hist()` with:
  - `bins=bins` parameter
  - `density=True` to normalize the histogram
  - `alpha=0.7` for transparency
  - `color='lightblue'`
  - `edgecolor='black'`
  - `label='Data Histogram'`
- For the PDF overlay:
  - Create x-values spanning the data range: `x = np.linspace(data.min() - 2*sample_std, data.max() + 2*sample_std, 1000)`
  - Calculate PDF using `scipy.stats.norm.pdf(x, loc=sample_mean, scale=sample_std)`
  - Plot as a red line with `linewidth=2` and `label='Fitted Normal PDF'`
- Add vertical lines for mean and mean±std:
  - Mean line: `ax.axvline(sample_mean, color='green', linestyle='-', linewidth=2, label=f'Mean = {sample_mean:.3f}')`
  - Mean±std lines: `ax.axvline(sample_mean ± sample_std, color='orange', linestyle='--', linewidth=1.5)`
- Set labels and formatting:
  - X-axis: "Value"  
  - Y-axis: "Density"
  - Title: f"Normal Distribution Analysis\\nMean = {sample_mean:.3f}, Std = {sample_std:.3f}, Variance = {sample_var:.3f}"
  - Add legend with `ax.legend()`
  - Add grid with `alpha=0.3`
- Return both the figure object and a tuple of the computed statistics: `(sample_mean, sample_var, sample_std)`

**Parameters:**
- `data`: numpy array, the dataset to analyze (assumed normally distributed)
- `bins`: int, default 30, number of bins for the histogram
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object
- `stats`: tuple (sample_mean, sample_variance, sample_std)

```python
def analyze_normal_dataset(data, bins=30, show_plot=False):
    # Compute sample statistics
    sample_mean = np.mean(...)
    sample_var = np.var(..., ddof=1)  # Use Bessel's correction
    sample_std = ...
    
    # Create the plot
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Create histogram
    
    # Overlay fitted normal PDF
    
    # Add vertical lines for mean and mean±std
    
    # Labels and formatting
    
    if show_plot:
        plt.show()
    
    return fig, (sample_mean, sample_var, sample_std)
```

```python
# Cell to view plot - test the function with sample data
rng = np.random.default_rng(rng_seed)
test_data = rng.normal(loc=5.0, scale=2.0, size=1000)  # Generate test data from known normal distribution
fig, stats = analyze_normal_dataset(test_data, bins=25, show_plot=True)
print(f"Computed statistics: Mean={stats[0]:.3f}, Variance={stats[1]:.3f}, Std={stats[2]:.3f}")
print(f"True parameters: Mean=5.0, Variance=4.0, Std=2.0")
```

```python
grader.check("q1")
```

## **Question 2**: Random Walk Ensemble Simulation

Write a function `plot_random_walks(p, N, M, show_plot=False)` that simulates and visualizes an ensemble of 1D random walks.

Your function should:
1. Simulate M independent random walks, each with N steps
2. For each walk, at each step, move up (+1) with probability p or down (-1) with probability (1-p)
3. Plot all M walks on the same figure as a function of step number
4. Show the ensemble behavior and individual trajectories

**Algorithm:**
- For each of the M walks:
  - Generate N random steps using Bernoulli distribution with probability p
  - Convert Bernoulli outcomes (0,1) to step directions (-1,+1)
  - Compute cumulative sum to get walk positions over time
- Plot all walks with step number on x-axis and position on y-axis

**Requirements:**
- Use the seeded random number generator: `rng = np.random.default_rng(rng_seed)`
- For each walk, generate steps using `rng.binomial(1, p, N)` to get 0s and 1s
- Convert to step directions: `steps = 2 * bernoulli_steps - 1` (maps 0→-1, 1→+1)
- Calculate walk positions using `np.cumsum()` with initial position 0
- Create the plot with:
  - X-axis: Step number (0 to N)
  - Y-axis: Position
  - Plot each walk as a line with `alpha=0.7` for transparency
  - Use different colors or let matplotlib cycle through colors automatically
- Add formatting:
  - X-axis label: "Step Number"
  - Y-axis label: "Position"
  - Title: f"Ensemble of {M} Random Walks (p = {p:.2f}, N = {N} steps)"
  - Grid with `alpha=0.3`
  - Set reasonable axis limits based on data range
- Return the matplotlib figure object

**Parameters:**
- `p`: float, probability of taking an upward step (+1)
- `N`: int, number of steps in each walk
- `M`: int, number of walks to simulate  
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object

```python
def plot_random_walks(p, N, M, show_plot=False):
    # Use seeded random number generator
    rng = np.random.default_rng(rng_seed)
    
    # Create the plot
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Simulate M random walks
    for walk_idx in range(M):
        # Generate steps using Bernoulli distribution
        bernoulli_steps = rng.binomial(...)
        
        # Convert to step directions (-1, +1)
        steps = ...
        
        # Calculate cumulative positions
        positions = ...
        
        # Plot this walk
        step_numbers = ...
        ax.plot(...)
    
    # Labels and formatting
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Cell to view plot - test the function with sample parameters
fig1 = plot_random_walks(p=0.5, N=100, M=20, show_plot=True)  # Fair random walk
fig2 = plot_random_walks(p=0.7, N=100, M=20, show_plot=True)  # Biased upward walk
```

```python
grader.check("q2")
```

## **Question 3**: Monte Carlo Integration

### Part a: Area Under Curve Estimation

Write a function `monte_carlo_integration(func, a, b, n_points, show_plot=False)` that estimates the area under a curve using Monte Carlo methods.

Your function should:
1. Generate random points in the rectangular domain [a, b] × [0, max_y]
2. Determine which points fall under the curve defined by the function
3. Calculate the area using the ratio of points under the curve to total points
4. Create a visualization showing the function, sampling points, and integration region

**Algorithm:**
- Find the maximum value of the function in the interval [a, b] to define the sampling rectangle
- Generate random points uniformly in the rectangle [a, b] × [0, max_y]
- For each point (x, y), check if y ≤ func(x) (point is under the curve)
- Estimate area = (rectangle area) × (fraction of points under curve)
- Area = (b - a) × max_y × (points_under_curve / total_points)

**Requirements:**
- Use the seeded random number generator: `rng = np.random.default_rng(rng_seed)`
- To find max_y, evaluate the function on a dense grid: `x_eval = np.linspace(a, b, 1000)` and `max_y = np.max(func(x_eval))`
- Generate random points using `rng.uniform()`:
  - X coordinates: `rng.uniform(a, b, n_points)`
  - Y coordinates: `rng.uniform(0, max_y, n_points)`
- For each point, check if `y_random <= func(x_random)`
- Count points under the curve and calculate area estimate
- Create the plot with:
  - Plot the function curve using `x_plot = np.linspace(a, b, 1000)` and `ax.plot(x_plot, func(x_plot), 'b-', linewidth=2, label='Function')`
  - Plot points under curve as green dots: `ax.scatter(x_under, y_under, c='green', s=10, alpha=0.6, label='Under curve')`
  - Plot points above curve as red dots: `ax.scatter(x_above, y_above, c='red', s=10, alpha=0.6, label='Above curve')`
  - Fill area under curve: `ax.fill_between(x_plot, func(x_plot), alpha=0.3, color='lightblue', label='Integration region')`
- Add formatting:
  - X-axis label: "x"
  - Y-axis label: "f(x)" 
  - Title: f"Monte Carlo Integration: Area ≈ {area_estimate:.4f}"
  - Legend with `ax.legend()`
  - Grid with `alpha=0.3`
- Return both the area estimate and the figure object

**Parameters:**
- `func`: callable, function to integrate (assumed positive on [a, b])
- `a`: float, left boundary of integration
- `b`: float, right boundary of integration  
- `n_points`: int, number of random points to generate
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `area_estimate`: float, estimated area under the curve
- `fig`: matplotlib figure object

```python
def monte_carlo_integration(func, a, b, n_points, show_plot=False):
    # Use seeded random number generator
    rng = np.random.default_rng(rng_seed)
    
    # Find maximum y value for sampling rectangle
    x_eval = np.linspace(a, b, 1000)
    max_y = np.max(func(...))
    
    # Generate random points
    x_random = rng.uniform(...)
    y_random = rng.uniform(...)
    
    # Check which points are under the curve
    under_curve = ...
    
    # Calculate area estimate
    area_estimate = ...
    
    # Create the plot
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Plot function and points
    
    # Labels and formatting
    
    if show_plot:
        plt.show()
    
    return area_estimate, fig
```

```python
# Cell to view plot - test the function with sample functions
# Test with quadratic function
def test_func1(x):
    return x**2

area1, fig1 = monte_carlo_integration(test_func1, 0, 2, 5000, show_plot=True)
print(f"Estimated area under x² from 0 to 2: {area1:.4f}")
print(f"Analytical result: {8/3:.4f}")

# Test with sine function
def test_func2(x):
    return np.sin(x) + 1  # Adding 1 to ensure positive

area2, fig2 = monte_carlo_integration(test_func2, 0, np.pi, 3000, show_plot=True)
print(f"Estimated area under sin(x)+1 from 0 to π: {area2:.4f}")
print(f"Analytical result: {np.pi + 2:.4f}")
```

```python
grader.check("q3")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: None)
```
>>> rng = np.random.default_rng(rng_seed)
>>> test_data = rng.normal(0, 1, 100)
>>> fig, stats = analyze_normal_dataset(test_data)
>>> import matplotlib.figure
>>> test_result = isinstance(fig, matplotlib.figure.Figure) and isinstance(stats, tuple) and (len(stats) == 3)
>>> plt.close(fig)
>>> test_result
True
```
```
>>> rng = np.random.default_rng(rng_seed)
>>> test_data = rng.normal(10, 3, 10000)
>>> fig, stats = analyze_normal_dataset(test_data)
>>> sample_mean, sample_var, sample_std = stats
>>> mean_close = abs(sample_mean - 10) < 0.2
>>> var_close = abs(sample_var - 9) < 0.5
>>> std_close = abs(sample_std - 3) < 0.2
>>> std_var_consistent = abs(sample_std - np.sqrt(sample_var)) < 1e-10
>>> plt.close(fig)
>>> (mean_close, var_close, std_close, std_var_consistent)
(np.True_, np.True_, np.True_, np.True_)
```


Tests for q2 (points: None)
```
>>> np.random.seed(rng_seed)
>>> fig = plot_random_walks(p=0.5, N=10, M=5)
>>> import matplotlib.figure
>>> test_result = isinstance(fig, matplotlib.figure.Figure)
>>> plt.close(fig)
>>> test_result
True
```
```
>>> np.random.seed(rng_seed)
>>> fig = plot_random_walks(p=0.4, N=15, M=3)
>>> ax = fig.axes[0]
>>> lines = ax.get_lines()
>>> correct_lengths = all((len(line.get_xdata()) == 16 for line in lines))
>>> start_positions = [line.get_ydata()[0] for line in lines]
>>> all_start_zero = all((pos == 0 for pos in start_positions))
>>> xlim = ax.get_xlim()
>>> correct_xlim = xlim[0] <= 0 and xlim[1] >= 15
>>> plt.close(fig)
>>> (correct_lengths, all_start_zero, correct_xlim)
(True, True, np.True_)
```


Tests for q3 (points: None)
```
>>> np.random.seed(rng_seed)
>>> def simple_func(x):
...     return x + 1
>>> area, fig = monte_carlo_integration(simple_func, 0, 1, 100)
>>> import matplotlib.figure
>>> test_result = isinstance(area, (int, float, np.floating)) and isinstance(fig, matplotlib.figure.Figure)
>>> plt.close(fig)
>>> test_result
True
```
```
>>> np.random.seed(rng_seed)
>>> def quadratic_func(x):
...     return x ** 2
>>> area, fig = monte_carlo_integration(quadratic_func, 0, 2, 500)
>>> ax = fig.axes[0]
>>> lines = ax.get_lines()
>>> has_function_line = len(lines) >= 1
>>> collections = ax.collections
>>> has_scatter_points = len(collections) >= 1
>>> has_xlabel = ax.get_xlabel() == 'x'
>>> has_ylabel = ax.get_ylabel() == 'f(x)'
>>> title = ax.get_title()
>>> has_area_in_title = 'Monte Carlo Integration' in title and 'Area' in title
>>> has_legend = ax.get_legend() is not None
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> plt.close(fig)
>>> (has_function_line, has_scatter_points, has_xlabel, has_ylabel, has_area_in_title, has_legend, has_grid)
(True, True, True, True, True, True, True)
```
