```python
# Initialize Otter
import otter
grader = otter.Notebook("ws6.ipynb")
```

```python
rng_seed = 88
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
import scipy as sp
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

## **Question 1**: Statistical Distributions - Sampling and Visualization

In this question, you'll work with different statistical distributions by sampling from them and visualizing the results with histograms overlaid with their exact probability density functions (PDFs) or probability mass functions (PMFs).

### Part a: Normal Distribution

Write a function `plot_normal_distribution(n_samples, mean, std, bins, show_plot=False)` that:
1. Generates `n_samples` random samples from a normal (Gaussian) distribution with specified mean and standard deviation
2. Creates a histogram of the samples 
3. Overlays the exact probability density function (PDF)
4. Adds a vertical line at the theoretical mean

**Requirements:**
- Use `np.random.normal(mean, std, n_samples)` to generate samples
- Create histogram using `ax.hist()` with:
  - `bins=bins` parameter
  - `density=True` to normalize the histogram (so it matches the PDF scale)
  - `alpha=0.7` for transparency
  - `color='skyblue'`
  - `edgecolor='black'`
- For the PDF overlay:
  - Create x-values using `np.linspace(mean - 4*std, mean + 4*std, 1000)`
  - Calculate PDF using `scipy.stats.norm.pdf(x, loc=mean, scale=std)`
  - Plot as a red line with `linewidth=2`
- Add a vertical line at the mean using `ax.axvline(mean, color='green', linestyle='--', linewidth=2, label='Mean')`
- Set labels and formatting:
  - X-axis: "Value"
  - Y-axis: "Density"
  - Title: "Normal Distribution: mu={mean:.1f}, sigma={std:.1f}" (use f-string formatting)
  - Add legend with `ax.legend()`
  - Add grid with `alpha=0.3`
- Return the matplotlib figure object

**Parameters:**
- `n_samples`: int, number of random samples to generate
- `mean`: float, mean of the normal distribution
- `std`: float, standard deviation of the normal distribution  
- `bins`: int, number of bins for the histogram
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object

```python
def plot_normal_distribution(n_samples, mean, std, bins, show_plot=False):
    # Generate samples and create plot
    #seeding rng for predictable testing
    rng = np.random.default_rng(rng_seed)
    samples = rng.normal(...)
    
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Create histogram
    
    # Overlay PDF
    
    # Add mean line
    
    # Labels and formatting
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Cell to view plot
np.random.seed(rng_seed)  # Set seed for reproducibility
plot_normal_distribution(1000, 2.0, 1.5, 30, show_plot=True);
```

```python
grader.check("q1a")
```

### Part b: Uniform Distribution

Write a function `plot_uniform_distribution(n_samples, a, b, bins, show_plot=False)` that:
1. Generates `n_samples` random samples from a uniform distribution on the interval [a, b]
2. Creates a histogram of the samples 
3. Overlays the exact probability density function (PDF)
4. Adds a vertical line at the theoretical mean

**Requirements:**
- Use `np.random.uniform(a, b, n_samples)` to generate samples
- Create histogram using `ax.hist()` with:
  - `bins=bins` parameter
  - `density=True` to normalize the histogram
  - `alpha=0.7` for transparency
  - `color='lightcoral'`
  - `edgecolor='black'`
- For the PDF overlay:
  - Create x-values using `np.linspace(a - 0.5*(b-a), b + 0.5*(b-a), 1000)`
  - Calculate PDF using `scipy.stats.uniform.pdf(x, loc=a, scale=b-a)`
  - Plot as a blue line with `linewidth=2`
- Calculate the theoretical mean: `mean = (a + b) / 2`
- Add a vertical line at the mean using `ax.axvline(mean, color='green', linestyle='--', linewidth=2, label='Mean')`
- Set labels and formatting:
  - X-axis: "Value"
  - Y-axis: "Density"  
  - Title: "Uniform Distribution: [a={a:.1f}, b={b:.1f}]" (use f-string formatting)
  - Add legend with `ax.legend()`
  - Add grid with `alpha=0.3`
- Return the matplotlib figure object

**Parameters:**
- `n_samples`: int, number of random samples to generate
- `a`: float, lower bound of the uniform distribution
- `b`: float, upper bound of the uniform distribution
- `bins`: int, number of bins for the histogram
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object

```python
def plot_uniform_distribution(n_samples, a, b, bins, show_plot=False):
    # Generate samples and create plot
    rng = np.random.default_rng(rng_seed)
    samples = rng.uniform(...)
    
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Create histogram
    
    # Overlay PDF
    
    # Calculate mean and add mean line
    
    # Labels and formatting
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Cell to view plot
np.random.seed(rng_seed)  # Set seed for reproducibility
plot_uniform_distribution(1000, -2.0, 4.0, 25, show_plot=True);
```

```python
grader.check("q1b")
```

### Part c: Poisson Distribution

Write a function `plot_poisson_distribution(n_samples, lam, bins, show_plot=False)` that:
1. Generates `n_samples` random samples from a Poisson distribution with parameter λ (lambda)
2. Creates a histogram of the samples 
3. Overlays the exact probability mass function (PMF)
4. Adds a vertical line at the theoretical mean

**Requirements:**
- Use `np.random.poisson(lam, n_samples)` to generate samples
- Create histogram using `ax.hist()` with:
  - `bins=bins` parameter
  - `density=True` to normalize the histogram
  - `alpha=0.7` for transparency
  - `color='lightgreen'`
  - `edgecolor='black'`
- For the PMF overlay:
  - Create integer x-values using `np.arange(0, int(lam + 5*np.sqrt(lam)) + 1)` (this covers the relevant range)
  - Calculate PMF using `scipy.stats.poisson.pmf(x, mu=lam)`
  - Plot as orange points connected by lines using `ax.plot(x, pmf, 'o-', color='orange', linewidth=2, markersize=4, label='PMF')`
- For Poisson distribution, the theoretical mean equals λ
- Add a vertical line at the mean using `ax.axvline(lam, color='green', linestyle='--', linewidth=2, label='Mean')`
- Set labels and formatting:
  - X-axis: "Value"
  - Y-axis: "Density"  
  - Title: "Poisson Distribution: lambda={lam:.1f}" (use f-string formatting)
  - Add legend with `ax.legend()`
  - Add grid with `alpha=0.3`
- Return the matplotlib figure object

**Parameters:**
- `n_samples`: int, number of random samples to generate
- `lam`: float, rate parameter λ of the Poisson distribution
- `bins`: int, number of bins for the histogram
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object

**Note:** Since Poisson is a discrete distribution, we plot the PMF as points connected by lines rather than a continuous curve.

```python
def plot_poisson_distribution(n_samples, lam, bins, show_plot=False):
    # Generate samples and create plot
    rng = np.random.default_rng(rng_seed)
    samples = rng.poisson(...)
    
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Create histogram
    
    # Overlay PMF
    
    # Add mean line (mean = λ for Poisson)
    
    # Labels and formatting
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Cell to view plot
np.random.seed(rng_seed)  # Set seed for reproducibility
plot_poisson_distribution(1000, 3.5, 12, show_plot=True);
```

```python
grader.check("q1c")
```

## **Question 2**: Monte Carlo Estimation of π

In this question, you'll use the Monte Carlo method to estimate the value of π by randomly scattering points in a unit square and determining which ones fall inside a quarter circle.

### Part a: Monte Carlo π Estimator

Write a function `estimate_pi_monte_carlo(n_points)` that estimates the value of π using the Monte Carlo method.

**Method:**
The approach is based on the ratio of areas:
- A unit square has area 1
- A quarter circle with radius 1 inscribed in the unit square has area π/4
- Therefore, the ratio of points inside the quarter circle to total points should approximate π/4

**Algorithm:**
1. Generate `n_points` random points (x, y) uniformly distributed in the unit square [0,1] × [0,1]
2. For each point, check if it falls inside the quarter circle: x² + y² ≤ 1
3. Count how many points fall inside the quarter circle
4. Estimate π using: π ≈ 4 × (points_inside_circle / total_points)

**Requirements:**
- Use the seeded random number generator: `rng = np.random.default_rng(rng_seed)`
- Generate random points using `rng.uniform(0, 1, (n_points, 2))`
- Calculate distances from origin using vectorized operations
- Count points inside the quarter circle (distance ≤ 1)
- Return the estimated value of π

**Parameters:**
- `n_points`: int, number of random points to generate

**Returns:**
- `pi_estimate`: float, estimated value of π

**Mathematical Background:**
For a quarter circle of radius 1:
- Area of quarter circle = π/4
- Area of unit square = 1
- Probability that a random point falls in quarter circle = (π/4) / 1 = π/4
- Therefore: π = 4 × P(point in circle) ≈ 4 × (count_inside / total_points)

```python
def estimate_pi_monte_carlo(n_points):
    # Use seeded random number generator
    rng = np.random.default_rng(rng_seed)
    
    # Generate random points and estimate π
    
    return pi_estimate
```

```python
# Test the function with a few values
print("Monte Carlo π estimates:")
for n in [100, 1000, 10000, 100000]:
    pi_est = estimate_pi_monte_carlo(n)
    error = abs(pi_est - np.pi)
    print(f"N = {n:6d}: π ≈ {pi_est:.6f}, error = {error:.6f}")
print(f"True π = {np.pi:.6f}")
```

```python
grader.check("q2a")
```

### Part b: Convergence Analysis Plot

Write a function `plot_pi_convergence(n_values, show_plot=False)` that creates a plot showing how the Monte Carlo estimate of π converges as the number of sample points increases.

This function should use the `estimate_pi_monte_carlo()` function from Part a to compute π estimates for different values of N and visualize the convergence behavior.

**Requirements:**
- Use the `estimate_pi_monte_carlo()` function from Part a to compute π estimates for each N in `n_values`
- Create a plot with:
  - X-axis: Number of points (log scale) - use `ax.set_xscale('log')`
  - Y-axis: Estimated value of π
  - Plot the estimates as blue circles connected by lines: `ax.plot(n_values, pi_estimates, 'bo-', linewidth=2, markersize=6)`
- Add a horizontal reference line for the true value of π:
  - Use `ax.axhline(np.pi, color='red', linestyle='--', linewidth=2, label='True π')`
- Set labels and formatting:
  - X-axis: "Number of Points"
  - Y-axis: "Estimated π"
  - Title: "Monte Carlo Convergence to π"
  - Add legend with `ax.legend()`
  - Add grid with `alpha=0.3`
- Set y-axis limits to show the convergence clearly: `ax.set_ylim([2.8, 3.5])`
- Return the matplotlib figure object

**Parameters:**
- `n_values`: array-like, list of N values to test (should be integers)
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object

**Note:** The log scale on the x-axis helps visualize how the estimate improves with increasing sample size. You should expect to see the estimates generally converging toward the true value of π ≈ 3.14159 as N increases.

```python
def plot_pi_convergence(n_values, show_plot=False):
    # Calculate π estimates for different N values
    np.random.seed(rng_seed)
    
    # Create plot with log scale x-axis
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Plot estimates and true π line
    
    # Labels and formatting
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Cell to view plot
np.random.seed(rng_seed)
n_test_values = [100, 300, 1000, 3000, 10000, 30000, 100000]
plot_pi_convergence(n_test_values, show_plot=True);
```

```python
grader.check("q2b")
```

## **Question 3**: Biased Coin Simulation

In this question, you'll simulate a biased coin and analyze how the observed probability of heads converges to the theoretical value as the number of flips increases.

### Part a: Biased Coin Flip Simulator

Write a function `simulate_coin_flips(p, n_flips)` that simulates flipping a biased coin.

**Method:**
A biased coin has probability `p` of landing heads and probability `(1-p)` of landing tails. This follows a Bernoulli distribution where:
- Success (heads) occurs with probability `p`
- Failure (tails) occurs with probability `1-p`

**Requirements:**
- Use the seeded random number generator: `rng = np.random.default_rng(rng_seed)`
- Generate random samples from the Bernoulli distribution using `rng.binomial(...)`
- Return a list or array where 1 represents heads and 0 represents tails
- The function should work for any probability `p` between 0 and 1

**Parameters:**
- `p`: float, probability of heads (between 0 and 1)
- `n_flips`: int, number of coin flips to simulate

**Returns:**
- `flips`: numpy array or list, sequence of 1s (heads) and 0s (tails)

```python
def simulate_coin_flips(p, n_flips):
    # Use seeded random number generator
    rng = np.random.default_rng(rng_seed)
    
    # Generate coin flips from Bernoulli distribution
    
    return flips
```

```python
# Test the function with different probabilities
np.random.seed(rng_seed)
print("Biased coin simulation tests:")
for p in [0.3, 0.5, 0.7]:
    flips = simulate_coin_flips(p, 1000)
    observed_p = np.mean(flips)
    error = abs(observed_p - p)
    print(f"p = {p:.1f}: observed = {observed_p:.3f}, error = {error:.3f}")
    print(f"  First 20 flips: {flips[:20]}")
```

```python
grader.check("q3a")
```

### Part b: Coin Flip Convergence Analysis

Write a function `plot_coin_convergence(p, n_values, show_plot=False)` that creates a plot showing how the observed proportion of heads converges to the theoretical probability as the number of flips increases.

This function should use the `simulate_coin_flips()` function from Part a to simulate coin flips for different numbers of flips and visualize the convergence behavior.

**Requirements:**
- Use the `simulate_coin_flips()` function from Part a to simulate coin flips for each N in `n_values`
- For each N, calculate the observed proportion of heads: `proportion = np.mean(flips)`
- Create a plot with:
  - X-axis: Number of flips (log scale) - use `ax.set_xscale('log')`
  - Y-axis: Observed proportion of heads
  - Plot the proportions as green circles connected by lines: `ax.plot(n_values, proportions, 'go-', linewidth=2, markersize=6)`
- Add a horizontal reference line for the theoretical probability:
  - Use `ax.axhline(p, color='red', linestyle='--', linewidth=2, label=f'Theoretical p = {p:.2f}')`
- Set labels and formatting:
  - X-axis: "Number of Flips"
  - Y-axis: "Observed Proportion of Heads"
  - Title: f"Coin Flip Convergence (p = {p:.2f})"
  - Add legend with `ax.legend()`
  - Add grid with `alpha=0.3`
- Set y-axis limits based on the probability: `ax.set_ylim([max(0, p-0.3), min(1, p+0.3)])`
- Return the matplotlib figure object

**Parameters:**
- `p`: float, true probability of heads for the biased coin
- `n_values`: array-like, list of N values to test (number of flips)
- `show_plot`: bool, default False. If True, call `plt.show()`

**Returns:**
- `fig`: matplotlib figure object

**Note:** This demonstrates the Law of Large Numbers - as the number of trials increases, the observed frequency approaches the theoretical probability.

```python
def plot_coin_convergence(p, n_values, show_plot=False):
    # Calculate observed proportions for different N values
    np.random.seed(rng_seed)
    
    # Create plot with log scale x-axis
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Plot proportions and theoretical line
    
    # Labels and formatting
    
    if show_plot:
        plt.show()
    
    return fig
```

```python
# Cell to view plot
np.random.seed(rng_seed)
n_flip_values = [50, 200, 500, 2000, 5000, 20000, 50000]
plot_coin_convergence(0.3, n_flip_values, show_plot=True);
```

```python
grader.check("q3b")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> np.random.seed(rng_seed)
>>> fig_test = plot_normal_distribution(100, 0, 1, 20)
>>> import matplotlib.figure
>>> temp = isinstance(fig_test, matplotlib.figure.Figure)
>>> plt.close(fig_test)
>>> temp
True
```
```
>>> np.random.seed(rng_seed)
>>> fig_normal = plot_normal_distribution(500, 1, 2, 25)
>>> ax = fig_normal.axes[0]
>>> has_histogram = len(ax.patches) > 0
>>> lines = ax.get_lines()
>>> has_pdf_line = len(lines) >= 1
>>> has_vertical_line = len(ax.collections) > 0 or any(('axvline' in str(type(child)) for child in ax.get_children()))
>>> has_xlabel = ax.get_xlabel() == 'Value'
>>> has_ylabel = ax.get_ylabel() == 'Density'
>>> plt.close(fig_normal)
>>> (has_histogram, has_pdf_line, has_xlabel, has_ylabel)
(True, True, True, True)
```
```
>>> np.random.seed(42)
>>> fig_title = plot_normal_distribution(200, 3.5, 0.8, 15)
>>> ax = fig_title.axes[0]
>>> title = ax.get_title()
>>> has_correct_title = 'Normal Distribution' in title and '3.5' in title and ('0.8' in title)
>>> has_legend = ax.get_legend() is not None
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> plt.close(fig_title)
>>> (has_correct_title, has_legend, has_grid)
(True, True, True)
```


Tests for q1b (points: None)
```
>>> np.random.seed(rng_seed)
>>> fig_test = plot_uniform_distribution(100, 0, 5, 20)
>>> import matplotlib.figure
>>> temp = isinstance(fig_test, matplotlib.figure.Figure)
>>> plt.close(fig_test)
>>> temp
True
```
```
>>> np.random.seed(rng_seed)
>>> fig_uniform = plot_uniform_distribution(500, 1, 7, 20)
>>> ax = fig_uniform.axes[0]
>>> has_histogram = len(ax.patches) > 0
>>> lines = ax.get_lines()
>>> has_pdf_line = len(lines) >= 1
>>> has_xlabel = ax.get_xlabel() == 'Value'
>>> has_ylabel = ax.get_ylabel() == 'Density'
>>> plt.close(fig_uniform)
>>> (has_histogram, has_pdf_line, has_xlabel, has_ylabel)
(True, True, True, True)
```
```
>>> np.random.seed(rng_seed)
>>> fig_title = plot_uniform_distribution(200, -1.5, 2.5, 15)
>>> ax = fig_title.axes[0]
>>> title = ax.get_title()
>>> has_correct_title = 'Uniform Distribution' in title and '-1.5' in title and ('2.5' in title)
>>> has_legend = ax.get_legend() is not None
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> plt.close(fig_title)
>>> (has_correct_title, has_legend, has_grid)
(True, True, True)
```


Tests for q1c (points: None)
```
>>> np.random.seed(rng_seed)
>>> fig_test = plot_poisson_distribution(100, 2.5, 15)
>>> import matplotlib.figure
>>> temp = isinstance(fig_test, matplotlib.figure.Figure)
>>> plt.close(fig_test)
>>> temp
True
```
```
>>> np.random.seed(rng_seed)
>>> fig_poisson = plot_poisson_distribution(500, 4.0, 20)
>>> ax = fig_poisson.axes[0]
>>> has_histogram = len(ax.patches) > 0
>>> lines = ax.get_lines()
>>> has_pmf_line = len(lines) >= 1
>>> has_xlabel = ax.get_xlabel() == 'Value'
>>> has_ylabel = ax.get_ylabel() == 'Density'
>>> plt.close(fig_poisson)
>>> (has_histogram, has_pmf_line, has_xlabel, has_ylabel)
(True, True, True, True)
```
```
>>> np.random.seed(rng_seed)
>>> fig_title = plot_poisson_distribution(200, 2.8, 12)
>>> ax = fig_title.axes[0]
>>> title = ax.get_title()
>>> has_correct_title = 'Poisson Distribution' in title and '2.8' in title
>>> has_legend = ax.get_legend() is not None
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> lines = ax.get_lines()
>>> has_markers = len(lines) > 0 and any((line.get_marker() != 'None' for line in lines))
>>> plt.close(fig_title)
>>> (has_correct_title, has_legend, has_grid, has_markers)
(True, True, True, True)
```


Tests for q2a (points: None)
```
>>> np.random.seed(rng_seed)
>>> pi_est_small = estimate_pi_monte_carlo(1000)
>>> test_result = isinstance(pi_est_small, (int, float, np.floating, np.integer)) and 2.0 < pi_est_small < 4.0
>>> test_result
np.True_
```
```
>>> np.random.seed(rng_seed)
>>> pi_est_10k = estimate_pi_monte_carlo(10000)
>>> pi_est_100k = estimate_pi_monte_carlo(100000)
>>> error_10k = abs(pi_est_10k - np.pi)
>>> error_100k = abs(pi_est_100k - np.pi)
>>> bool(error_100k < error_10k)
True
```
```
>>> np.random.seed(rng_seed)
>>> pi_est1 = estimate_pi_monte_carlo(5000)
>>> pi_est2 = estimate_pi_monte_carlo(5000)
>>> test_result = np.isclose(pi_est1, pi_est2, rtol=1e-10)
>>> bool(test_result)
True
```


Tests for q2b (points: None)
```
>>> np.random.seed(rng_seed)
>>> test_n_values = [100, 1000, 10000]
>>> fig_test = plot_pi_convergence(test_n_values)
>>> import matplotlib.figure
>>> temp = isinstance(fig_test, matplotlib.figure.Figure)
>>> plt.close(fig_test)
>>> temp
True
```
```
>>> np.random.seed(rng_seed)
>>> test_n_values = [500, 2000, 8000]
>>> fig_conv = plot_pi_convergence(test_n_values)
>>> ax = fig_conv.axes[0]
>>> lines = ax.get_lines()
>>> has_data_line = len(lines) >= 1
>>> has_reference_line = len(lines) >= 2
>>> has_log_scale = ax.get_xscale() == 'log'
>>> has_xlabel = ax.get_xlabel() == 'Number of Points'
>>> has_ylabel = ax.get_ylabel() == 'Estimated π'
>>> plt.close(fig_conv)
>>> (has_data_line, has_reference_line, has_log_scale, has_xlabel, has_ylabel)
(True, True, True, True, True)
```
```
>>> np.random.seed(rng_seed)
>>> test_n_values = [200, 1000, 5000]
>>> fig_format = plot_pi_convergence(test_n_values)
>>> ax = fig_format.axes[0]
>>> has_title = 'Monte Carlo' in ax.get_title() and 'π' in ax.get_title()
>>> has_legend = ax.get_legend() is not None
>>> ylim = ax.get_ylim()
>>> correct_ylim = ylim[0] <= 2.8 and ylim[1] >= 3.5
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> plt.close(fig_format)
>>> (has_title, has_legend, correct_ylim, has_grid)
(True, True, np.True_, True)
```


Tests for q3a (points: None)
```
>>> np.random.seed(rng_seed)
>>> flips_test = simulate_coin_flips(0.6, 100)
>>> test_result = len(flips_test) == 100 and isinstance(flips_test, np.ndarray) and np.all(np.isin(flips_test, [0, 1]))
>>> test_result
np.True_
```
```
>>> np.random.seed(rng_seed)
>>> flips_fair = simulate_coin_flips(0.5, 10000)
>>> flips_biased = simulate_coin_flips(0.8, 10000)
>>> fair_prop = np.mean(flips_fair)
>>> biased_prop = np.mean(flips_biased)
>>> test_result = 0.45 < fair_prop < 0.55 and 0.75 < biased_prop < 0.85 and (biased_prop > fair_prop)
>>> bool(test_result)
True
```
```
>>> np.random.seed(rng_seed)
>>> flips1 = simulate_coin_flips(0.3, 500)
>>> flips2 = simulate_coin_flips(0.3, 500)
>>> test_result = np.array_equal(flips1, flips2)
>>> bool(test_result)
True
```


Tests for q3b (points: None)
```
>>> test_n_values = [100, 1000, 10000]
>>> fig_test = plot_coin_convergence(0.4, test_n_values)
>>> import matplotlib.figure
>>> temp = isinstance(fig_test, matplotlib.figure.Figure)
>>> plt.close(fig_test)
>>> temp
True
```
```
>>> test_n_values = [500, 2000, 8000]
>>> fig_conv = plot_coin_convergence(0.7, test_n_values)
>>> ax = fig_conv.axes[0]
>>> lines = ax.get_lines()
>>> has_data_line = len(lines) >= 1
>>> has_reference_line = len(lines) >= 2
>>> has_log_scale = ax.get_xscale() == 'log'
>>> has_xlabel = ax.get_xlabel() == 'Number of Flips'
>>> has_ylabel = ax.get_ylabel() == 'Observed Proportion of Heads'
>>> plt.close(fig_conv)
>>> (has_data_line, has_reference_line, has_log_scale, has_xlabel, has_ylabel)
(True, True, True, True, True)
```
```
>>> test_n_values = [200, 1000, 5000]
>>> fig_format = plot_coin_convergence(0.25, test_n_values)
>>> ax = fig_format.axes[0]
>>> title = ax.get_title()
>>> has_correct_title = 'Coin Flip Convergence' in title and '0.25' in title
>>> has_legend = ax.get_legend() is not None
>>> ylim = ax.get_ylim()
>>> reasonable_ylim = ylim[0] >= 0 and ylim[1] <= 1 and (ylim[1] - ylim[0] <= 0.7)
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> plt.close(fig_format)
>>> (has_correct_title, has_legend, reasonable_ylim, has_grid)
(True, True, np.True_, True)
```
