```python
# Initialize Otter
import otter
grader = otter.Notebook("ws8.ipynb")
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import pandas as pd
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

## **Question 1**: Determining the Order of a Matrix

In this question, you'll implement a function to determine the **order** of a square matrix. The order of a matrix is the smallest positive integer $n$ such that $A^n = I$, where $I$ is the identity matrix.

### Background: Matrix Order Theory

For a square matrix $A$, the **order** (also called the **period**) is defined as the smallest positive integer $n$ such that:

$$A^n = I$$

where $I$ is the identity matrix of the same size as $A$.

**Key Concepts:**

- If $A^n = I$ for some positive integer $n$, then $A$ is called a **periodic matrix** of order $n$
- Not all matrices have a finite order (some matrices never return to the identity)
- Rotation matrices often have finite orders. For example, a 90° rotation in 2D has order 4, since rotating four times returns to the original orientation
- Permutation matrices and certain orthogonal matrices commonly have finite orders

**Mathematical Properties:**

1. **Matrix Power**: The $n$-th power of matrix $A$ is computed by multiplying $A$ by itself $n$ times:
   $$A^n = \underbrace{A \cdot A \cdot ... \cdot A}_{n \text{ times}}$$

2. **Identity Matrix**: The identity matrix $I$ is a square matrix with 1's on the diagonal and 0's elsewhere:
   $$I = \begin{pmatrix} 1 & 0 & 0 \\ 0 & 1 & 0 \\ 0 & 0 & 1 \end{pmatrix} \quad \text{(for 3×3)}$$

3. **Matrix Equality with Tolerance**: Due to numerical precision limitations in computers, we check if two matrices are "close enough" using:
   $$|A^n - I| < \text{tolerance}$$
   where $|M|$ represents the magnitude of the matrix difference (we'll use `np.allclose()` for this)

### Your Task

Write a function `find_matrix_order(A, N, tolerance=1e-10)` that:
1. Takes a square matrix $A$ and checks if $A^n = I$ for $n = 1, 2, 3, ..., N$
2. Returns the smallest $n$ for which $A^n$ is approximately equal to the identity matrix (within the specified tolerance)
3. Returns `None` if no order is found up to $N$

**Requirements:**
- Use `np.linalg.matrix_power(...)` to compute $A^n$
- Use `np.eye(...)` to create the identity matrix of the same size as $A$
- Use `np.allclose(...)` to check if $A^n \approx I$ within the tolerance
- Check orders from $n=1$ up to $n=N$ (inclusive)
- Return the **smallest** $n$ that satisfies the condition
- Return `None` if no order is found within the range $[1, N]$

**Parameters:**
- `A`: numpy array, square matrix to analyze
- `N`: int, maximum order to check (check orders from 1 to N inclusive)
- `tolerance`: float, tolerance for checking matrix equality (default 1e-10)

**Returns:**
- `order`: int or None, the smallest order $n$ such that $A^n \approx I$, or None if no order found

**Example:**
For a 90° rotation matrix in 2D:
```python
A = np.array([[0, -1], [1, 0]])  # 90° counterclockwise rotation
order = find_matrix_order(A, 10)
# Should return 4, since rotating 4 times = 360° = identity
```

```python
def find_matrix_order(A, N, tolerance=1e-10):
    # your code here
```

```python
grader.check("q1")
```

## **Question 2**: Polynomial Fitting with Least Squares

In this question, you'll implement polynomial curve fitting using the **least squares method** with the **design matrix approach**. This is a fundamental technique in data analysis for fitting polynomial models to experimental data.

### Background: Least Squares Polynomial Fitting

Given a set of data points $(x_1, y_1), (x_2, y_2), ..., (x_m, y_m)$, we want to fit an $N$-th degree polynomial:

$$p(x) = a_0 + a_1 x + a_2 x^2 + ... + a_N x^N$$

The **least squares method** finds the coefficients $\mathbf{a} = [a_0, a_1, ..., a_N]^T$ that minimize the sum of squared residuals:

$$\text{minimize} \quad \sum_{i=1}^{m} (y_i - p(x_i))^2$$

### The Design Matrix Approach

The key to solving this problem is constructing the **design matrix** $\mathbf{X}$, where each row corresponds to a data point and each column corresponds to a power of $x$:

$$\mathbf{X} = \begin{pmatrix}
1 & x_1 & x_1^2 & \cdots & x_1^N \\
1 & x_2 & x_2^2 & \cdots & x_2^N \\
\vdots & \vdots & \vdots & \ddots & \vdots \\
1 & x_m & x_m^2 & \cdots & x_m^N
\end{pmatrix}$$

This is an $m \times (N+1)$ matrix, where $m$ is the number of data points.

**Matrix Formulation:**

The polynomial fitting problem can be written in matrix form as:
$$\mathbf{X} \mathbf{a} = \mathbf{y}$$

where $\mathbf{y} = [y_1, y_2, ..., y_m]^T$ is the vector of observed values.

**The Normal Equations:**

The least squares solution is found by solving the **normal equations**:

$$\mathbf{X}^T \mathbf{X} \mathbf{a} = \mathbf{X}^T \mathbf{y}$$

Solving for $\mathbf{a}$:
$$\mathbf{a} = (\mathbf{X}^T \mathbf{X})^{-1} \mathbf{X}^T \mathbf{y}$$

**In NumPy**, we can use `np.linalg.lstsq(X, y, rcond=None)` which solves this efficiently and handles numerical stability issues.

### Your Task

Write a function `fit_polynomial(filename, N, show_plot=False)` that:
1. Loads data from a CSV file containing x and y columns
2. Constructs the design matrix $\mathbf{X}$ for an $N$-th degree polynomial
3. Solves for the polynomial coefficients using least squares
4. Plots the data points and the fitted polynomial curve
5. Returns the figure object and the fitted coefficients

**Requirements:**
- Load the CSV file using `pd.read_csv(filename)`
- Extract x and y values from the DataFrame
- Construct the design matrix $\mathbf{X}$ with shape $(m, N+1)$ where:
  - Column 0: all ones (for $a_0$)
  - Column 1: $x$ values (for $a_1$)
  - Column 2: $x^2$ values (for $a_2$)
  - ...
  - Column $k$: $x^k$ values (for $a_k$)
- Use `np.linalg.lstsq(X, y, rcond=None)` to solve for coefficients
- Create a smooth curve for plotting using `np.linspace()` with at least 200 points
- Plot configuration:
  - Scatter plot of data with `'o'` markers, `alpha=0.6`, `s=40`, label='Data'
  - Line plot of fit with `linewidth=2`, `color='red'`, label='Fitted Polynomial'
  - X-axis label: "x"
  - Y-axis label: "y"
  - Title: "Polynomial Fit (Degree N)" where N is the actual degree
  - Grid with `alpha=0.3`
  - Legend
- Only call `plt.show()` if `show_plot=True`
- Return both the figure object and the coefficient array

**Parameters:**
- `filename`: string, path to CSV file with 'x' and 'y' columns
- `N`: int, degree of polynomial to fit (0 ≤ N)
- `show_plot`: bool, default False. If True, display the plot

**Returns:**
- `fig`: matplotlib figure object
- `coefficients`: numpy array of shape (N+1,), coefficients [a₀, a₁, ..., aₙ]

**Example:**
For a quadratic fit (N=2) to data, the function returns:
- A figure showing data points and the fitted parabola
- Coefficients array like `[1.5, -0.3, 0.8]` representing $y = 1.5 - 0.3x + 0.8x^2$

```python
def fit_polynomial(filename, N, show_plot=False):
    # Load the data
    
    
    # Construct the design matrix X with shape (m, N+1)
    
    # Solve for coefficients using np.linalg.lstsq
    
    # create array to plot fitted curve
    
    # Create the plot
    fig, ax = plt.subplots(figsize=(10, 6))
    
    # Plot data and fit
    ...
    
    # Set labels, title, grid, legend
    ...
    
    # Show plot if requested
    if show_plot:
        plt.show()
    
    return fig, coefficients
```

```python
# Example: Fit polynomials of different degrees to the data
fig1, coef1 = fit_polynomial('polynomial_data.csv', N=1, show_plot=True)
print(f"Linear fit coefficients: {coef1}")
print(f"Polynomial: y = {coef1[0]:.3f} + {coef1[1]:.3f}x")
print()

fig2, coef2 = fit_polynomial('polynomial_data.csv', N=2, show_plot=True)
print(f"Quadratic fit coefficients: {coef2}")
print(f"Polynomial: y = {coef2[0]:.3f} + {coef2[1]:.3f}x + {coef2[2]:.3f}x²")
print()

fig3, coef3 = fit_polynomial('polynomial_data.csv', N=3, show_plot=True)
print(f"Cubic fit coefficients: {coef3}")
print(f"Polynomial: y = {coef3[0]:.3f} + {coef3[1]:.3f}x + {coef3[2]:.3f}x² + {coef3[3]:.3f}x³")
print("\nTrue polynomial: y = 2 + 3x - 0.5x² + 0.1x³")
```

```python
grader.check("q2")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1 (points: None)
```
>>> A_test = np.array([[0, -1], [1, 0]])
>>> result = find_matrix_order(A_test, 10)
>>> test_result = isinstance(result, int) or result is None
>>> test_result
True
```
```
>>> A_rotation = np.array([[0, -1], [1, 0]])
>>> order_rotation = find_matrix_order(A_rotation, 10)
>>> test_result = order_rotation == 4
>>> test_result
True
```
```
>>> A_identity = np.eye(3)
>>> order_identity = find_matrix_order(A_identity, 10)
>>> test_result = order_identity == 1
>>> test_result
True
```
```
>>> A_no_order = np.array([[2, 0], [0, 0.5]])
>>> order_no_order = find_matrix_order(A_no_order, 10)
>>> test_result = order_no_order is None
>>> test_result
True
```
```
>>> A_180 = np.array([[-1, 0], [0, -1]])
>>> order_180 = find_matrix_order(A_180, 10)
>>> test_result = order_180 == 2
>>> test_result
True
```
```
>>> A_perm = np.array([[0, 1, 0], [0, 0, 1], [1, 0, 0]])
>>> order_perm = find_matrix_order(A_perm, 10)
>>> test_result = order_perm == 3
>>> test_result
True
```


Tests for q2 (points: None)
```
>>> fig_test, coef_test = fit_polynomial('polynomial_data.csv', N=2)
>>> import matplotlib.figure
>>> test_result = isinstance(fig_test, matplotlib.figure.Figure) and isinstance(coef_test, np.ndarray)
>>> plt.close(fig_test)
>>> test_result
True
```
```
>>> fig_test, coef_test = fit_polynomial('polynomial_data.csv', N=3)
>>> test_result = len(coef_test) == 4
>>> plt.close(fig_test)
>>> test_result
True
```
```
>>> test_data_const = pd.DataFrame({'x': [1, 2, 3, 4, 5], 'y': [5, 5, 5, 5, 5]})
>>> test_data_const.to_csv('test_const.csv', index=False)
>>> fig_const, coef_const = fit_polynomial('test_const.csv', N=0)
>>> test_result = len(coef_const) == 1 and abs(coef_const[0] - 5.0) < 0.01
>>> plt.close(fig_const)
>>> bool(test_result)
True
```
```
>>> test_data_linear = pd.DataFrame({'x': [1, 2, 3, 4, 5], 'y': [5, 8, 11, 14, 17]})
>>> test_data_linear.to_csv('test_linear.csv', index=False)
>>> fig_linear, coef_linear = fit_polynomial('test_linear.csv', N=1)
>>> test_result = len(coef_linear) == 2 and abs(coef_linear[0] - 2.0) < 0.01 and (abs(coef_linear[1] - 3.0) < 0.01)
>>> plt.close(fig_linear)
>>> bool(test_result)
True
```
```
>>> fig_components, coef_components = fit_polynomial('polynomial_data.csv', N=2)
>>> ax = fig_components.axes[0]
>>> has_scatter = len(ax.collections) > 0
>>> has_line = len(ax.lines) > 0
>>> has_xlabel = ax.get_xlabel() == 'x'
>>> has_ylabel = ax.get_ylabel() == 'y'
>>> has_title = 'Polynomial Fit' in ax.get_title()
>>> has_legend = ax.get_legend() is not None
>>> plt.close(fig_components)
>>> test_result = has_scatter and has_line and has_xlabel and has_ylabel and has_title and has_legend
>>> test_result
True
```
```
>>> fig_cubic, coef_cubic = fit_polynomial('polynomial_data.csv', N=3)
>>> test_a0 = abs(coef_cubic[0] - 2.0) < 1.0
>>> test_a1 = abs(coef_cubic[1] - 3.0) < 1.0
>>> test_a2 = abs(coef_cubic[2] - -0.5) < 1.0
>>> test_a3 = abs(coef_cubic[3] - 0.1) < 1.0
>>> plt.close(fig_cubic)
>>> test_result = test_a0 and test_a1 and test_a2 and test_a3
>>> bool(test_result)
True
```
