```python
# Initialize Otter
import otter
grader = otter.Notebook("ws5.ipynb")
```

```python
#imports
import numpy as np
import matplotlib.pyplot as plt
#below line allows matplotlib plots to appear in cell output
%matplotlib inline
```

## **Question 1**: Data Input/Output

In this question, you'll work with data from a physics simulation of projectile motion with air resistance. The data is stored in a CSV file and you'll learn to convert it to NumPy's compressed format.

### Part a: Reading CSV and Saving as NPZ

[np.savez documentation](https://numpy.org/doc/2.3/reference/generated/numpy.savez.html)

Write a function `convert_csv_to_npz(csv_filename, npz_filename)` that:
1. Reads data from a CSV file using `np.genfromtxt()` or `np.loadtxt()`
2. Extracts the columns into separate numpy arrays
3. Saves the arrays to a compressed `.npz` file using `np.savez()`

The CSV file `projectile_simulation.csv` contains simulation data for projectile motion with air resistance. It has 3 header lines (starting with '#') followed by a column header line, then the data rows with 6 columns:
- `time_s`: Time in seconds
- `x_position_m`: Horizontal position in meters
- `y_position_m`: Vertical position in meters
- `x_velocity_m/s`: Horizontal velocity in m/s
- `y_velocity_m/s`: Vertical velocity in m/s
- `kinetic_energy_J`: Kinetic energy in Joules

**Requirements:**
- Skip the first 4 lines (3 comment lines + 1 column header line)
- Use `delimiter=','` when loading the CSV
- Save to the NPZ file with these exact array names:
  - `'time'` for time data
  - `'x_pos'` for x position
  - `'y_pos'` for y position
  - `'x_vel'` for x velocity
  - `'y_vel'` for y velocity
  - `'ke'` for kinetic energy

**Parameters:**
- `csv_filename`: string, path to the CSV file to read
- `npz_filename`: string, path to the NPZ file to create

**Returns:**
- None (but creates the .npz file) (no return statement needed)

```python
def convert_csv_to_npz(csv_filename, npz_filename):
    # Write your code here!
    pass
```

```python
grader.check("q1a")
```

### Part b: Loading NPZ and Computing Physics Quantities

Write a function `analyze_projectile_data(npz_filename)` that:
1. Loads the NPZ file created in Part a
2. Performs calculations on the projectile motion data
3. Returns key physics quantities

The function should calculate and return:
- **Maximum height**: The maximum value of y_position (in meters)
- **Range**: The maximum value of x_position (in meters)
- **Flight time**: The total time from launch to when the projectile returns to ground level (y_position ≈ 0)
- **Energy loss**: The difference between initial and final kinetic energy (in Joules)

**Requirements:**
- Load the NPZ file using `np.load()`
- Access arrays using the keys: `'time'`, `'x_pos'`, `'y_pos'`, `'ke'`
- For flight time: find the last time where y_position is still positive (> 0.01 m to account for numerical precision)
- Calculate energy_loss = initial_ke - final_ke

**Parameters:**
- `npz_filename`: string, path to the NPZ file to load

**Returns:**
- Tuple `(max_height, range_distance, flight_time, energy_loss)` where:
  - `max_height`: float, maximum vertical position in meters
  - `range_distance`: float, maximum horizontal position in meters
  - `flight_time`: float, time of flight in seconds
  - `energy_loss`: float, kinetic energy lost due to air resistance in Joules

Note: Make sure you call your .npz file `projectile_data.npz` so that test cases work.

```python
def analyze_projectile_data(npz_filename):
    # Load the NPZ file
    data = np.load(npz_filename)

    # Write your code here!
    
    return (max_height, range_distance, flight_time, energy_loss)
```

```python
grader.check("q1b")
```

## **Question 2**: Polynomial and 2D Function Plotting

### Part a: General Polynomial Plotter

Write a function `plot_polynomial(coeff, show_plot=False)` that plots a polynomial of arbitrary degree based on the coefficients provided.

The polynomial is defined as:
$$P(x) = c_0 + c_1 x + c_2 x^2 + c_3 x^3 + \cdots + c_n x^n$$

where `coeff = [c_0, c_1, c_2, ..., c_n]` is an array of coefficients. The degree of the polynomial is determined by the length of the coefficient array.

**Requirements:**
- Use `np.linspace(-3, 3, 200)` to create 200 x-values in the range [-3, 3]
- Calculate the polynomial using vectorized NumPy operations (use powers of x)
- Create a plot with:
  - X-axis label: "x"
  - Y-axis label: "P(x)"
  - Title: "Polynomial of Degree N" (where N is the actual degree)
  - Grid with alpha=0.3
  - Blue line with linewidth=2
- Set xlim to [-3, 3]
- The function should return the matplotlib figure object

**Parameters:**
- `coeff`: numpy array or list of floats, polynomial coefficients [c_0, c_1, c_2, ...]
- `show_plot`: boolean, default False. If True, call `plt.show()` to display the plot

**Returns:**
- `fig`: matplotlib figure object

**Hint:** To evaluate the polynomial, use a loop or `np.polyval()` with reversed coefficients, or calculate each term c_i * x**i and sum them.

```python
def plot_polynomial(coeff, show_plot=False):
   
    #do calculation to make array to plot

    #plotting
    fig, ax = plt.subplots(figsize=(10, 6))
    ...
    
    # Show plot if requested
    if show_plot:
        plt.show()
    
    return fig
```

```python
#cell to view plot
coeff = [1,-5,2,-4] #feel free to change this
plot_polynomial(coeff, show_plot=True);
```

```python
grader.check("q2a")
```

### Part b: 2D Function Heatmap

Write a function `plot_2d_function(show_plot=False)` that creates a heatmap visualization of a 2D function.

The function to visualize is:
$$z = f(x, y) = \sin(2x) \cos(3y) + \cos(x) \sin(y)$$

This combination of sine and cosine functions creates an interesting interference pattern.

**Requirements:**
- Create a grid of (x, y) values:
  - x ranges from -π to π with 100 points
  - y ranges from -π to π with 100 points
  - Use `np.meshgrid()` to create the 2D grid
- Calculate z = f(x, y) on the grid
- Create a heatmap using `plt.imshow()` with:
  - Colormap: 'RdBu_r' (red-blue reversed)
  - `origin='lower'` to have y increasing upwards
  - `extent=[-np.pi, np.pi, -np.pi, np.pi]` to set axis limits correctly
  - `aspect='auto'` for proper scaling
- Add a colorbar with label "f(x, y)"
can be done with:  
```
cbar = plt.colorbar(im, ax=ax)
cbar.set_label('f(x, y)')
```
(`im` is the output of `ax.imshow`)
- Set axis labels:
  - X-axis: "x"
  - Y-axis: "y"
- Set title: "Heatmap: sin(2x)cos(3y) + cos(x)sin(y)"
- The function should return the matplotlib figure object

**Parameters:**
- `show_plot`: boolean, default False. If True, call `plt.show()` to display the plot

**Returns:**
- `fig`: matplotlib figure object

**Hint:** After creating the meshgrid X, Y, you can compute the function directly using NumPy: `Z = np.sin(2*X) * np.cos(3*Y) + np.cos(X) * np.sin(Y)`

```python
def plot_2d_function(show_plot=False):
   
    #do calculation to make array to plot

    #plotting
    fig, ax = plt.subplots(figsize=(10, 8))
    ...
    
    # Show plot if requested
    if show_plot:
        plt.show()
    
    return fig

```

```python
plot_2d_function(show_plot=True);
```

```python
grader.check("q2b")
```

## **Question 3**: Scatter Plot with Distance-Based Colormap

In this question, you'll visualize 2D position data using a scatter plot where the color of each point depends on its distance from the origin.

Write a function `plot_position_scatter(filename, show_plot=False)` that creates a scatter plot of 2D position data with colors representing the distance from the origin.

The data file contains two columns:
- Column 1: x-position
- Column 2: y-position

Each point should be colored according to its distance from the origin: $r = \sqrt{x^2 + y^2}$

**Requirements:**
- Load the data using `np.loadtxt()` with `skiprows=3` (to skip the header lines)
- Extract x and y positions from the columns
- Calculate the distance from origin for each point: $r = \sqrt{x^2 + y^2}$
- Create a scatter plot using `ax.scatter()` with:
  - `c=distances` to color by distance
  - `cmap='viridis'` for the colormap
  - `s=30` for marker size
  - `alpha=0.6` for transparency
- Add a colorbar with label "Distance from Origin"
- Set axis labels:
  - X-axis: "x position"
  - Y-axis: "y position"
- Set title: "2D Position Data Colored by Distance"
- Set equal aspect ratio using `ax.set_aspect('equal')`
- Add a grid with `alpha=0.3`
- The function should return the matplotlib figure object

**Parameters:**
- `filename`: string, path to the data file to read
- `show_plot`: boolean, default False. If True, call `plt.show()` to display the plot

**Returns:**
- `fig`: matplotlib figure object


```python
def plot_position_scatter(filename, show_plot=False):
   
    # Load data and calculate distances
    
    # Create scatter plot
    fig, ax = plt.subplots(figsize=(10, 10))
    ...
    
    # Show plot if requested
    if show_plot:
        plt.show()
    
    return fig

```

```python
# Cell to view plot
plot_position_scatter('position_data.txt', show_plot=True);
```

```python
grader.check("q3")
```

## Required disclosure of use of AI technology

Please indicate whether you used AI to complete this homework. If you did, explain how you used it in the python cell below, as a comment.

```python
"""
# write ai disclosure here:

"""
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit.

Upload the .zip file to Gradescope!

```python
grader.export(pdf=False, force_save=True, run_tests=True)
```

Tests for q1a (points: None)
```
>>> import os
>>> try:
...     convert_csv_to_npz('projectile_simulation.csv', 'projectile_data.npz')
...     test_result = os.path.exists('projectile_data.npz')
...     if test_result:
...         data = np.load('projectile_data.npz')
...         test_result = all((key in data for key in ['time', 'x_pos', 'y_pos', 'x_vel', 'y_vel', 'ke']))
... except:
...     test_result = False
>>> test_result
True
```
```
>>> try:
...     data = np.load('projectile_data.npz')
...     test_result = len(data['time']) == 100 and len(data['x_pos']) == 100 and (data['time'][0] == 0.0) and (data['ke'][0] > 200) and (data['ke'][0] < 230)
... except:
...     test_result = False
>>> bool(test_result)
True
```


Tests for q1b (points: None)
```
>>> try:
...     result = analyze_projectile_data('projectile_data.npz')
...     test_result = isinstance(result, tuple) and len(result) == 4 and all((isinstance(x, (int, float, np.floating, np.integer)) for x in result))
... except:
...     test_result = False
>>> test_result
True
```
```
>>> try:
...     max_height, range_distance, flight_time, energy_loss = analyze_projectile_data('projectile_data.npz')
...     test_result = 15.0 < max_height < 18.0 and 50.0 < range_distance < 60.0 and (3.0 < flight_time < 5.0) and (200.0 < energy_loss < 230.0)
... except:
...     test_result = False
>>> bool(test_result)
True
```
```
>>> try:
...     max_height, range_distance, flight_time, energy_loss = analyze_projectile_data('projectile_data.npz')
...     test_result = energy_loss > 0 and range_distance > max_height and (flight_time > 0)
... except:
...     test_result = False
>>> bool(test_result)
True
```


Tests for q2a (points: None)
```
>>> fig_test = plot_polynomial([1, 2, -1])
>>> import matplotlib.figure
>>> temp = isinstance(fig_test, matplotlib.figure.Figure)
>>> plt.close(fig_test)
>>> temp
True
```
```
>>> fig_quad = plot_polynomial([2, 0, -1])
>>> ax = fig_quad.axes[0]
>>> has_one_line = len(ax.get_lines()) == 1
>>> xlim = ax.get_xlim()
>>> has_correct_xlim = np.isclose(xlim[0], -3, atol=0.1) and np.isclose(xlim[1], 3, atol=0.1)
>>> has_xlabel = ax.get_xlabel() == 'x'
>>> has_ylabel = ax.get_ylabel() == 'P(x)'
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> plt.close(fig_quad)
>>> (has_one_line, has_correct_xlim, has_xlabel, has_ylabel, has_grid)
(True, np.True_, True, True, True)
```
```
>>> fig_linear = plot_polynomial([1, 2])
>>> ax = fig_linear.axes[0]
>>> line_data = ax.get_lines()[0]
>>> x_data = line_data.get_xdata()
>>> y_data = line_data.get_ydata()
>>> idx_zero = np.argmin(np.abs(x_data))
>>> correct_intercept = np.isclose(y_data[idx_zero], 1.0, atol=0.1)
>>> slope_avg = np.mean(np.diff(y_data) / np.diff(x_data))
>>> correct_slope = np.isclose(slope_avg, 2.0, atol=0.1)
>>> has_correct_title = 'Polynomial of Degree 1' in ax.get_title()
>>> plt.close(fig_linear)
>>> (correct_intercept, correct_slope, has_correct_title)
(np.True_, np.True_, True)
```


Tests for q2b (points: None)
```
>>> fig_2d = plot_2d_function()
>>> import matplotlib.figure
>>> ax = fig_2d.axes[0]
>>> images = ax.get_images()
>>> plt.close(fig_2d)
>>> isinstance(fig_2d, matplotlib.figure.Figure) and len(images) == 1
True
```
```
>>> fig_2d_test = plot_2d_function()
>>> ax = fig_2d_test.axes[0]
>>> has_xlabel = ax.get_xlabel() == 'x'
>>> has_ylabel = ax.get_ylabel() == 'y'
>>> title = ax.get_title()
>>> has_title = 'sin(2x)cos(3y)' in title or 'Heatmap' in title
>>> has_colorbar = len(fig_2d_test.axes) >= 2
>>> im = ax.get_images()[0]
>>> extent = im.get_extent()
>>> correct_extent = np.isclose(extent[0], -np.pi, atol=0.1) and np.isclose(extent[1], np.pi, atol=0.1) and np.isclose(extent[2], -np.pi, atol=0.1) and np.isclose(extent[3], np.pi, atol=0.1)
>>> plt.close(fig_2d_test)
>>> (has_xlabel, has_ylabel, has_title, has_colorbar, correct_extent)
(True, True, True, True, np.True_)
```


Tests for q3 (points: None)
```
>>> fig_scatter = plot_position_scatter('position_data.txt')
>>> import matplotlib.figure
>>> ax = fig_scatter.axes[0]
>>> collections = ax.collections
>>> plt.close(fig_scatter)
>>> isinstance(fig_scatter, matplotlib.figure.Figure) and len(collections) > 0
True
```
```
>>> fig_test = plot_position_scatter('position_data.txt')
>>> ax = fig_test.axes[0]
>>> has_xlabel = ax.get_xlabel() == 'x position'
>>> has_ylabel = ax.get_ylabel() == 'y position'
>>> has_title = '2D Position Data' in ax.get_title() or 'Distance' in ax.get_title()
>>> has_colorbar = len(fig_test.axes) >= 2
>>> has_equal_aspect = ax.get_aspect() == 1.0 or ax.get_aspect() == 'equal'
>>> plt.close(fig_test)
>>> (has_xlabel, has_ylabel, has_title, has_colorbar, has_equal_aspect)
(True, True, True, True, True)
```
```
>>> fig_dist = plot_position_scatter('position_data.txt')
>>> ax = fig_dist.axes[0]
>>> scatter_collection = ax.collections[0]
>>> color_array = scatter_collection.get_array()
>>> has_colors = color_array is not None and len(color_array) == 500
>>> if has_colors:
...     reasonable_range = np.all(color_array >= 0) and np.all(color_array <= 1.5)
... else:
...     reasonable_range = False
>>> plt.close(fig_dist)
>>> (has_colors, reasonable_range)
(True, np.True_)
```
```
>>> fig_count = plot_position_scatter('position_data.txt')
>>> ax = fig_count.axes[0]
>>> scatter_collection = ax.collections[0]
>>> offsets = scatter_collection.get_offsets()
>>> correct_count = len(offsets) == 500
>>> has_grid = ax.xaxis.get_gridlines()[0].get_visible()
>>> plt.close(fig_count)
>>> (correct_count, has_grid)
(True, True)
```
